# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>((k2_relset_1(X1,X2,X3)=X2&(r2_relset_1(X1,X1,k1_partfun1(X1,X2,X2,X1,X3,X4),k6_partfun1(X1))&v2_funct_1(X3)))=>(X1=k1_xboole_0|(X2=k1_xboole_0|X4=k2_funct_1(X3)))))), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', t36_funct_2)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', d1_funct_2)).
fof(t31_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X3)&k2_relset_1(X1,X2,X3)=X2)=>(v1_funct_1(k2_funct_1(X3))&(v1_funct_2(k2_funct_1(X3),X2,X1)&m1_subset_1(k2_funct_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1))))))), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', t31_funct_2)).
fof(t61_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k5_relat_1(X1,k2_funct_1(X1))=k6_relat_1(k1_relat_1(X1))&k5_relat_1(k2_funct_1(X1),X1)=k6_relat_1(k2_relat_1(X1))))), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', t61_funct_1)).
fof(redefinition_k6_partfun1, axiom, ![X1]:k6_partfun1(X1)=k6_relat_1(X1), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', redefinition_k6_partfun1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', redefinition_k1_relset_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', fc6_relat_1)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', redefinition_k1_partfun1)).
fof(t22_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(~(X2=k1_xboole_0)=>(k2_relset_1(X1,X2,X3)=X2<=>![X4]:(~(X4=k1_xboole_0)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X2,X4)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X2,X4)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))))=>(r2_relset_1(X1,X4,k1_partfun1(X1,X2,X2,X4,X3,X5),k1_partfun1(X1,X2,X2,X4,X3,X6))=>r2_relset_1(X2,X4,X5,X6)))))))), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', t22_funct_2)).
fof(redefinition_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_relset_1(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/7ee40296591e4b1a9cd72f70f5aee64b.E_input', redefinition_r2_relset_1)).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>((k2_relset_1(X1,X2,X3)=X2&(r2_relset_1(X1,X1,k1_partfun1(X1,X2,X2,X1,X3,X4),k6_partfun1(X1))&v2_funct_1(X3)))=>(X1=k1_xboole_0|(X2=k1_xboole_0|X4=k2_funct_1(X3))))))), inference(assume_negation,[status(cth)],[t36_funct_2])).
fof(c_0_12, plain, ![X11, X12, X13]:((((~v1_funct_2(X13,X11,X12)|X11=k1_relset_1(X11,X12,X13)|X12=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X11!=k1_relset_1(X11,X12,X13)|v1_funct_2(X13,X11,X12)|X12=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))))&((~v1_funct_2(X13,X11,X12)|X11=k1_relset_1(X11,X12,X13)|X11!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X11!=k1_relset_1(X11,X12,X13)|v1_funct_2(X13,X11,X12)|X11!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))))&((~v1_funct_2(X13,X11,X12)|X13=k1_xboole_0|X11=k1_xboole_0|X12!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X13!=k1_xboole_0|v1_funct_2(X13,X11,X12)|X11=k1_xboole_0|X12!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_13, negated_conjecture, ((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk2_0,esk1_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))))&((k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0&(r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,esk4_0),k6_partfun1(esk1_0))&v2_funct_1(esk3_0)))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&esk4_0!=k2_funct_1(esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_14, plain, ![X14, X15, X16]:((v1_funct_1(k2_funct_1(X16))|(~v2_funct_1(X16)|k2_relset_1(X14,X15,X16)!=X15)|(~v1_funct_1(X16)|(~v1_funct_2(X16,X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))))&((v1_funct_2(k2_funct_1(X16),X15,X14)|(~v2_funct_1(X16)|k2_relset_1(X14,X15,X16)!=X15)|(~v1_funct_1(X16)|(~v1_funct_2(X16,X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15))))))&(m1_subset_1(k2_funct_1(X16),k1_zfmisc_1(k2_zfmisc_1(X15,X14)))|(~v2_funct_1(X16)|k2_relset_1(X14,X15,X16)!=X15)|(~v1_funct_1(X16)|(~v1_funct_2(X16,X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_funct_2])])])).
fof(c_0_15, plain, ![X17]:((k5_relat_1(X17,k2_funct_1(X17))=k6_relat_1(k1_relat_1(X17))|~v2_funct_1(X17)|(~v1_relat_1(X17)|~v1_funct_1(X17)))&(k5_relat_1(k2_funct_1(X17),X17)=k6_relat_1(k2_relat_1(X17))|~v2_funct_1(X17)|(~v1_relat_1(X17)|~v1_funct_1(X17)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t61_funct_1])])])).
fof(c_0_16, plain, ![X18]:k6_partfun1(X18)=k6_relat_1(X18), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
fof(c_0_17, plain, ![X19, X20, X21]:(~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X19,X20)))|k1_relset_1(X19,X20,X21)=k1_relat_1(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_18, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_22, plain, ![X22, X23]:(~v1_relat_1(X22)|(~m1_subset_1(X23,k1_zfmisc_1(X22))|v1_relat_1(X23))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_23, plain, ![X24, X25]:v1_relat_1(k2_zfmisc_1(X24,X25)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_24, plain, ![X26, X27, X28, X29, X30, X31]:(~v1_funct_1(X30)|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X26,X27)))|(~v1_funct_1(X31)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X28,X29))))|k1_partfun1(X26,X27,X28,X29,X30,X31)=k5_relat_1(X30,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_25, plain, (m1_subset_1(k2_funct_1(X1),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v2_funct_1(X1)|k2_relset_1(X3,X2,X1)!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_26, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_27, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_28, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_29, plain, (v1_funct_1(k2_funct_1(X1))|~v2_funct_1(X1)|k2_relset_1(X2,X3,X1)!=X3|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_30, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_relat_1(k1_relat_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_31, plain, (k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_32, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_33, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk3_0)=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_34, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_20]), c_0_19]), c_0_28])])).
cnf(c_0_38, negated_conjecture, (v1_funct_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_26]), c_0_27]), c_0_20]), c_0_19]), c_0_28])])).
cnf(c_0_39, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_partfun1(k1_relat_1(X1))|~v1_funct_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_20]), c_0_33])).
cnf(c_0_41, negated_conjecture, (v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_20]), c_0_35])])).
cnf(c_0_42, negated_conjecture, (k1_partfun1(X1,X2,esk2_0,esk1_0,X3,k2_funct_1(esk3_0))=k5_relat_1(X3,k2_funct_1(esk3_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])])).
cnf(c_0_43, negated_conjecture, (k5_relat_1(esk3_0,k2_funct_1(esk3_0))=k6_partfun1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_27]), c_0_28])])).
fof(c_0_44, plain, ![X32, X33, X34, X35, X36, X37]:((k2_relset_1(X32,X33,X34)!=X33|(X35=k1_xboole_0|(~v1_funct_1(X36)|(~v1_funct_2(X36,X33,X35)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X33,X35))))|(~v1_funct_1(X37)|(~v1_funct_2(X37,X33,X35)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X33,X35))))|(~r2_relset_1(X32,X35,k1_partfun1(X32,X33,X33,X35,X34,X36),k1_partfun1(X32,X33,X33,X35,X34,X37))|r2_relset_1(X33,X35,X36,X37)))))|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&((esk5_3(X32,X33,X34)!=k1_xboole_0|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&(((v1_funct_1(esk6_3(X32,X33,X34))|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&((v1_funct_2(esk6_3(X32,X33,X34),X33,esk5_3(X32,X33,X34))|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&(m1_subset_1(esk6_3(X32,X33,X34),k1_zfmisc_1(k2_zfmisc_1(X33,esk5_3(X32,X33,X34))))|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))))&(((v1_funct_1(esk7_3(X32,X33,X34))|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&((v1_funct_2(esk7_3(X32,X33,X34),X33,esk5_3(X32,X33,X34))|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&(m1_subset_1(esk7_3(X32,X33,X34),k1_zfmisc_1(k2_zfmisc_1(X33,esk5_3(X32,X33,X34))))|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))))&((r2_relset_1(X32,esk5_3(X32,X33,X34),k1_partfun1(X32,X33,X33,esk5_3(X32,X33,X34),X34,esk6_3(X32,X33,X34)),k1_partfun1(X32,X33,X33,esk5_3(X32,X33,X34),X34,esk7_3(X32,X33,X34)))|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&(~r2_relset_1(X33,esk5_3(X32,X33,X34),esk6_3(X32,X33,X34),esk7_3(X32,X33,X34))|k2_relset_1(X32,X33,X34)=X33|X33=k1_xboole_0|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t22_funct_2])])])])])])).
cnf(c_0_45, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,esk4_0),k6_partfun1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_46, negated_conjecture, (k6_partfun1(esk1_0)=k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_20]), c_0_43]), c_0_28])])).
cnf(c_0_47, plain, (v1_funct_2(k2_funct_1(X1),X2,X3)|~v2_funct_1(X1)|k2_relset_1(X3,X2,X1)!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_48, plain, ![X41, X42, X43, X44]:((~r2_relset_1(X41,X42,X43,X44)|X43=X44|(~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))))&(X43!=X44|r2_relset_1(X41,X42,X43,X44)|(~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_relset_1])])])).
cnf(c_0_49, plain, (X4=k1_xboole_0|r2_relset_1(X2,X4,X5,X6)|X2=k1_xboole_0|k2_relset_1(X1,X2,X3)!=X2|~v1_funct_1(X5)|~v1_funct_2(X5,X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~v1_funct_1(X6)|~v1_funct_2(X6,X2,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~r2_relset_1(X1,X4,k1_partfun1(X1,X2,X2,X4,X3,X5),k1_partfun1(X1,X2,X2,X4,X3,X6))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_50, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,esk4_0),k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,k2_funct_1(esk3_0)))), inference(rw,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_52, negated_conjecture, (v1_funct_2(k2_funct_1(esk3_0),esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_26]), c_0_27]), c_0_20]), c_0_19]), c_0_28])])).
cnf(c_0_53, negated_conjecture, (v1_funct_2(esk4_0,esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_55, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_56, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, negated_conjecture, (r2_relset_1(esk2_0,esk1_0,esk4_0,k2_funct_1(esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_26]), c_0_37]), c_0_51]), c_0_20]), c_0_52]), c_0_53]), c_0_19]), c_0_38]), c_0_54]), c_0_28])]), c_0_55]), c_0_21])).
cnf(c_0_58, negated_conjecture, (esk4_0!=k2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_59, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_37]), c_0_51])]), c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 37
# Proof object formula steps           : 23
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 49
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 39
# Processed clauses                    : 167
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 163
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 7
# Generated clauses                    : 136
# ...of the previous two non-trivial   : 132
# Contextual simplify-reflections      : 13
# Paramodulations                      : 131
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 108
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 71
# Current number of unprocessed clauses: 42
# ...number of literals in the above   : 218
# Current number of archived formulas  : 0
# Current number of archived clauses   : 51
# Clause-clause subsumption calls (NU) : 1373
# Rec. Clause-clause subsumption calls : 414
# Non-unit clause-clause subsumptions  : 21
# Unit Clause-clause subsumption calls : 60
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9416

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28692 pages
