# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t75_enumset1)).
fof(t65_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,k1_tarski(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_tarski(X2))))))=>(r2_hidden(X3,X1)=>k1_funct_1(X4,X3)=X2)), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t65_funct_2)).
fof(t202_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>![X3]:(r2_hidden(X3,k2_relat_1(X2))=>r2_hidden(X3,X1))), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t202_relat_1)).
fof(t12_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X1),k2_relat_1(X2)))), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t12_funct_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', cc1_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', redefinition_k1_relset_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', d1_funct_2)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', d6_enumset1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/7f43db999dde47939e3bd27099eade00.E_input', t2_xboole_1)).
fof(c_0_18, plain, ![X11, X12, X13, X14, X15, X16]:(((~r2_hidden(X13,X12)|X13=X11|X12!=k1_tarski(X11))&(X14!=X11|r2_hidden(X14,X12)|X12!=k1_tarski(X11)))&((~r2_hidden(esk1_2(X15,X16),X16)|esk1_2(X15,X16)!=X15|X16=k1_tarski(X15))&(r2_hidden(esk1_2(X15,X16),X16)|esk1_2(X15,X16)=X15|X16=k1_tarski(X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_19, plain, ![X18]:k2_tarski(X18,X18)=k1_tarski(X18), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_20, plain, ![X19, X20]:k1_enumset1(X19,X19,X20)=k2_tarski(X19,X20), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, plain, ![X21, X22, X23]:k2_enumset1(X21,X21,X22,X23)=k1_enumset1(X21,X22,X23), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_22, plain, ![X24, X25, X26, X27]:k3_enumset1(X24,X24,X25,X26,X27)=k2_enumset1(X24,X25,X26,X27), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_23, plain, ![X28, X29, X30, X31, X32]:k4_enumset1(X28,X28,X29,X30,X31,X32)=k3_enumset1(X28,X29,X30,X31,X32), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_24, plain, ![X33, X34, X35, X36, X37, X38]:k5_enumset1(X33,X33,X34,X35,X36,X37,X38)=k4_enumset1(X33,X34,X35,X36,X37,X38), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_25, plain, ![X39, X40, X41, X42, X43, X44, X45]:k6_enumset1(X39,X39,X40,X41,X42,X43,X44,X45)=k5_enumset1(X39,X40,X41,X42,X43,X44,X45), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_26, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,k1_tarski(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k1_tarski(X2))))))=>(r2_hidden(X3,X1)=>k1_funct_1(X4,X3)=X2))), inference(assume_negation,[status(cth)],[t65_funct_2])).
cnf(c_0_27, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_35, plain, ![X50, X51, X52]:(~v1_relat_1(X51)|~v5_relat_1(X51,X50)|(~r2_hidden(X52,k2_relat_1(X51))|r2_hidden(X52,X50))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t202_relat_1])])])).
fof(c_0_36, plain, ![X53, X54]:(~v1_relat_1(X54)|~v1_funct_1(X54)|(~r2_hidden(X53,k1_relat_1(X54))|r2_hidden(k1_funct_1(X54,X53),k2_relat_1(X54)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_1])])).
fof(c_0_37, negated_conjecture, ((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,esk2_0,k1_tarski(esk3_0))&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_tarski(esk3_0))))))&(r2_hidden(esk4_0,esk2_0)&k1_funct_1(esk5_0,esk4_0)!=esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
cnf(c_0_38, plain, (X1=X3|X2!=k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34])).
cnf(c_0_39, plain, (r2_hidden(X3,X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)|~r2_hidden(X3,k2_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_41, plain, ![X55, X56, X57]:((v4_relat_1(X57,X55)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56))))&(v5_relat_1(X57,X56)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k1_tarski(esk3_0))))), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_43, plain, ![X58, X59, X60]:(~m1_subset_1(X60,k1_zfmisc_1(k2_zfmisc_1(X58,X59)))|v1_relat_1(X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_44, plain, ![X61, X62, X63]:(~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X62)))|k1_relset_1(X61,X62,X63)=k1_relat_1(X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_45, plain, (X1=X2|~r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~v5_relat_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_47, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34])).
cnf(c_0_49, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_50, plain, ![X64, X65, X66]:((((~v1_funct_2(X66,X64,X65)|X64=k1_relset_1(X64,X65,X66)|X65=k1_xboole_0|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65))))&(X64!=k1_relset_1(X64,X65,X66)|v1_funct_2(X66,X64,X65)|X65=k1_xboole_0|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))))&((~v1_funct_2(X66,X64,X65)|X64=k1_relset_1(X64,X65,X66)|X64!=k1_xboole_0|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65))))&(X64!=k1_relset_1(X64,X65,X66)|v1_funct_2(X66,X64,X65)|X64!=k1_xboole_0|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65))))))&((~v1_funct_2(X66,X64,X65)|X66=k1_xboole_0|X64=k1_xboole_0|X65!=k1_xboole_0|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65))))&(X66!=k1_xboole_0|v1_funct_2(X66,X64,X65)|X64=k1_xboole_0|X65!=k1_xboole_0|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X64,X65)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_51, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52, negated_conjecture, (v1_funct_2(esk5_0,esk2_0,k1_tarski(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_53, plain, (k1_funct_1(X1,X2)=X3|~v5_relat_1(X1,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_54, negated_conjecture, (v5_relat_1(esk5_0,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_55, negated_conjecture, (v1_relat_1(esk5_0)), inference(spm,[status(thm)],[c_0_49, c_0_48])).
cnf(c_0_56, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_57, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, negated_conjecture, (k1_relset_1(esk2_0,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0),esk5_0)=k1_relat_1(esk5_0)), inference(spm,[status(thm)],[c_0_51, c_0_48])).
cnf(c_0_59, negated_conjecture, (v1_funct_2(esk5_0,esk2_0,k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34])).
fof(c_0_60, plain, ![X67, X68, X69, X70, X71, X72, X73, X74, X75, X76, X77, X78, X79, X80, X81, X82, X83, X84, X85, X86]:(((~r2_hidden(X76,X75)|(X76=X67|(X76=X68|(X76=X69|(X76=X70|(X76=X71|(X76=X72|(X76=X73|X76=X74)))))))|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))&((X77!=X67|r2_hidden(X77,X75)|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))&((X77!=X68|r2_hidden(X77,X75)|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))&((X77!=X69|r2_hidden(X77,X75)|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))&((X77!=X70|r2_hidden(X77,X75)|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))&((X77!=X71|r2_hidden(X77,X75)|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))&((X77!=X72|r2_hidden(X77,X75)|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))&((X77!=X73|r2_hidden(X77,X75)|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))&(X77!=X74|r2_hidden(X77,X75)|X75!=k6_enumset1(X67,X68,X69,X70,X71,X72,X73,X74))))))))))&(((esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)!=X78|~r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85))&((esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)!=X79|~r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85))&((esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)!=X80|~r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85))&((esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)!=X81|~r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85))&((esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)!=X82|~r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85))&((esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)!=X83|~r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85))&((esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)!=X84|~r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85))&(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)!=X85|~r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85)))))))))&(r2_hidden(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86),X86)|(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)=X78|(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)=X79|(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)=X80|(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)=X81|(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)=X82|(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)=X83|(esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)=X84|esk6_9(X78,X79,X80,X81,X82,X83,X84,X85,X86)=X85)))))))|X86=k6_enumset1(X78,X79,X80,X81,X82,X83,X84,X85)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
cnf(c_0_61, negated_conjecture, (k1_funct_1(esk5_0,X1)=esk3_0|~r2_hidden(X1,k1_relat_1(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56])])).
cnf(c_0_62, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0|k1_relat_1(esk5_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_48]), c_0_58]), c_0_59])])).
fof(c_0_63, plain, ![X88, X89]:(~r2_hidden(X88,X89)|~r1_tarski(X89,X88)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_64, plain, ![X90]:r1_tarski(k1_xboole_0,X90), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_65, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X2,X4,X5,X6,X7,X8,X9,X10)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66, negated_conjecture, (k1_funct_1(esk5_0,esk4_0)!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_67, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0|k1_funct_1(esk5_0,X1)=esk3_0|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_69, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_70, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_71, plain, (r2_hidden(X1,k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_65])])).
cnf(c_0_72, negated_conjecture, (k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])])).
cnf(c_0_73, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_74, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 75
# Proof object clause steps            : 38
# Proof object formula steps           : 37
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 48
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 41
# Processed clauses                    : 157
# ...of these trivial                  : 1
# ...subsumed                          : 25
# ...remaining for further processing  : 131
# Other redundant clauses eliminated   : 50
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 13
# Generated clauses                    : 353
# ...of the previous two non-trivial   : 303
# Contextual simplify-reflections      : 2
# Paramodulations                      : 309
# Factorizations                       : 3
# Equation resolutions                 : 51
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 63
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 46
# Current number of unprocessed clauses: 221
# ...number of literals in the above   : 1476
# Current number of archived formulas  : 0
# Current number of archived clauses   : 60
# Clause-clause subsumption calls (NU) : 920
# Rec. Clause-clause subsumption calls : 429
# Non-unit clause-clause subsumptions  : 26
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 72
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9766

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.003 s
# Total time               : 0.016 s
# Maximum resident set size: 28716 pages
