# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', t63_subset_1)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', d8_lattice3)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', t7_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', d10_yellow_0)).
fof(fc1_finset_1, axiom, ![X1]:v1_finset_1(k1_tarski(X1)), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', fc1_finset_1)).
fof(t57_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r2_yellow_0(X1,X5)&X4=k2_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k2_yellow_0(X1,X4),X3)))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)<=>r1_lattice3(X1,X3,X4))))))), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', t57_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', t24_orders_2)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', fc2_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', fc1_xboole_0)).
fof(t9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', t9_yellow_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/7f5bece0e6c246c880d72dbbedc27e0d.E_input', t3_subset)).
fof(c_0_11, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>(![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r2_yellow_0(X1,X5)&X4=k2_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k2_yellow_0(X1,X4),X3)))))), introduced(definition)).
fof(c_0_12, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>(![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r2_yellow_0(X1,X5)&X4=k2_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k2_yellow_0(X1,X4),X3)))))), inference(split_equiv,[status(thm)],[c_0_11])).
fof(c_0_13, plain, ![X33, X34]:(~r2_hidden(X33,X34)|m1_subset_1(k1_tarski(X33),k1_zfmisc_1(X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
fof(c_0_14, plain, ![X28, X29, X30, X31]:((~r1_lattice3(X28,X29,X30)|(~m1_subset_1(X31,u1_struct_0(X28))|(~r2_hidden(X31,X29)|r1_orders_2(X28,X30,X31)))|~m1_subset_1(X30,u1_struct_0(X28))|~l1_orders_2(X28))&((m1_subset_1(esk6_3(X28,X29,X30),u1_struct_0(X28))|r1_lattice3(X28,X29,X30)|~m1_subset_1(X30,u1_struct_0(X28))|~l1_orders_2(X28))&((r2_hidden(esk6_3(X28,X29,X30),X29)|r1_lattice3(X28,X29,X30)|~m1_subset_1(X30,u1_struct_0(X28))|~l1_orders_2(X28))&(~r1_orders_2(X28,X30,esk6_3(X28,X29,X30))|r1_lattice3(X28,X29,X30)|~m1_subset_1(X30,u1_struct_0(X28))|~l1_orders_2(X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
fof(c_0_15, plain, ![X20, X21, X22]:((~r1_lattice3(X20,k1_tarski(X22),X21)|r1_orders_2(X20,X21,X22)|~m1_subset_1(X22,u1_struct_0(X20))|~m1_subset_1(X21,u1_struct_0(X20))|~l1_orders_2(X20))&((~r1_orders_2(X20,X21,X22)|r1_lattice3(X20,k1_tarski(X22),X21)|~m1_subset_1(X22,u1_struct_0(X20))|~m1_subset_1(X21,u1_struct_0(X20))|~l1_orders_2(X20))&((~r2_lattice3(X20,k1_tarski(X22),X21)|r1_orders_2(X20,X22,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~m1_subset_1(X21,u1_struct_0(X20))|~l1_orders_2(X20))&(~r1_orders_2(X20,X22,X21)|r2_lattice3(X20,k1_tarski(X22),X21)|~m1_subset_1(X22,u1_struct_0(X20))|~m1_subset_1(X21,u1_struct_0(X20))|~l1_orders_2(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
fof(c_0_16, plain, ![X11, X12, X13, X14]:(((r1_lattice3(X11,X12,X13)|X13!=k2_yellow_0(X11,X12)|~r2_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&(~m1_subset_1(X14,u1_struct_0(X11))|(~r1_lattice3(X11,X12,X14)|r1_orders_2(X11,X14,X13))|X13!=k2_yellow_0(X11,X12)|~r2_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11)))&((m1_subset_1(esk5_3(X11,X12,X13),u1_struct_0(X11))|~r1_lattice3(X11,X12,X13)|X13=k2_yellow_0(X11,X12)|~r2_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&((r1_lattice3(X11,X12,esk5_3(X11,X12,X13))|~r1_lattice3(X11,X12,X13)|X13=k2_yellow_0(X11,X12)|~r2_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&(~r1_orders_2(X11,esk5_3(X11,X12,X13),X13)|~r1_lattice3(X11,X12,X13)|X13=k2_yellow_0(X11,X12)|~r2_yellow_0(X11,X12)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_17, plain, ![X35, X36, X37, X38, X39, X41]:((~v1_finset_1(X38)|~m1_subset_1(X38,k1_zfmisc_1(X36))|(X38=k1_xboole_0|r2_yellow_0(X37,X38))|~epred1_3(X37,X36,X35))&((((v1_finset_1(esk7_4(X35,X36,X37,X39))|~r2_hidden(X39,X35)|~m1_subset_1(X39,u1_struct_0(X37))|~epred1_3(X37,X36,X35))&(m1_subset_1(esk7_4(X35,X36,X37,X39),k1_zfmisc_1(X36))|~r2_hidden(X39,X35)|~m1_subset_1(X39,u1_struct_0(X37))|~epred1_3(X37,X36,X35)))&((r2_yellow_0(X37,esk7_4(X35,X36,X37,X39))|~r2_hidden(X39,X35)|~m1_subset_1(X39,u1_struct_0(X37))|~epred1_3(X37,X36,X35))&(X39=k2_yellow_0(X37,esk7_4(X35,X36,X37,X39))|~r2_hidden(X39,X35)|~m1_subset_1(X39,u1_struct_0(X37))|~epred1_3(X37,X36,X35))))&(~v1_finset_1(X41)|~m1_subset_1(X41,k1_zfmisc_1(X36))|(X41=k1_xboole_0|r2_hidden(k2_yellow_0(X37,X41),X35))|~epred1_3(X37,X36,X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])).
cnf(c_0_18, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_20, plain, ![X23]:v1_finset_1(k1_tarski(X23)), inference(variable_rename,[status(thm)],[fc1_finset_1])).
fof(c_0_21, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)<=>r1_lattice3(X1,X3,X4)))))))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[t57_waybel_0]), c_0_11])).
cnf(c_0_22, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (r1_lattice3(X1,X2,esk5_3(X1,X2,X3))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (X1=k1_xboole_0|r2_yellow_0(X3,X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~epred1_3(X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (r1_lattice3(X1,X2,X3)|m1_subset_1(k1_tarski(esk6_3(X1,X2,X3)),k1_zfmisc_1(X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_27, plain, (v1_finset_1(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(epred1_3(esk1_0,esk2_0,esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&((~r1_lattice3(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0))&(r1_lattice3(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_21])])])])).
fof(c_0_29, plain, ![X16, X17]:(v2_struct_0(X16)|(~v3_orders_2(X16)|~l1_orders_2(X16))|(~m1_subset_1(X17,u1_struct_0(X16))|r1_orders_2(X16,X17,X17))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_30, plain, (X1=k1_xboole_0|r2_hidden(k2_yellow_0(X3,X1),X4)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~epred1_3(X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_31, plain, (X3=k2_yellow_0(X1,X2)|~r1_orders_2(X1,esk5_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_32, plain, (X1=k2_yellow_0(X2,k1_tarski(X3))|r1_orders_2(X2,esk5_3(X2,k1_tarski(X3),X1),X3)|~r1_lattice3(X2,k1_tarski(X3),X1)|~r2_yellow_0(X2,k1_tarski(X3))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
cnf(c_0_33, plain, (k1_tarski(esk6_3(X1,X2,X3))=k1_xboole_0|r1_lattice3(X1,X2,X3)|r2_yellow_0(X4,k1_tarski(esk6_3(X1,X2,X3)))|~epred1_3(X4,X2,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_34, negated_conjecture, (epred1_3(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37, plain, (k1_tarski(esk6_3(X1,X2,X3))=k1_xboole_0|r1_lattice3(X1,X2,X3)|r2_hidden(k2_yellow_0(X4,k1_tarski(esk6_3(X1,X2,X3))),X5)|~epred1_3(X4,X2,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_26]), c_0_27])])).
cnf(c_0_38, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|~r1_lattice3(X1,k1_tarski(X2),X2)|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39, negated_conjecture, (k1_tarski(esk6_3(X1,esk2_0,X2))=k1_xboole_0|r1_lattice3(X1,esk2_0,X2)|r2_yellow_0(esk1_0,k1_tarski(esk6_3(X1,esk2_0,X2)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_40, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (r1_orders_2(X1,esk6_3(X1,X2,X3),esk6_3(X1,X2,X3))|r1_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_43, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_44, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_45, negated_conjecture, (k1_tarski(esk6_3(X1,esk2_0,X2))=k1_xboole_0|r1_lattice3(X1,esk2_0,X2)|r2_hidden(k2_yellow_0(esk1_0,k1_tarski(esk6_3(X1,esk2_0,X2))),esk3_0)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_37, c_0_34])).
cnf(c_0_46, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(esk6_3(X1,esk2_0,X2)))=esk6_3(X1,esk2_0,X2)|k1_tarski(esk6_3(X1,esk2_0,X2))=k1_xboole_0|r1_lattice3(X1,esk2_0,X2)|~r1_lattice3(esk1_0,k1_tarski(esk6_3(X1,esk2_0,X2)),esk6_3(X1,esk2_0,X2))|~m1_subset_1(esk6_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
cnf(c_0_47, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_48, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,X1,esk4_0),esk6_3(esk1_0,X1,esk4_0))|r1_lattice3(esk1_0,X1,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_40]), c_0_43])]), c_0_44])).
fof(c_0_49, plain, ![X10]:~v1_xboole_0(k1_tarski(X10)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
cnf(c_0_50, negated_conjecture, (k1_tarski(esk6_3(X1,esk2_0,X2))=k1_xboole_0|r1_lattice3(X1,esk2_0,X2)|r2_hidden(esk6_3(X1,esk2_0,X2),esk3_0)|~r1_lattice3(esk1_0,k1_tarski(esk6_3(X1,esk2_0,X2)),esk6_3(X1,esk2_0,X2))|~m1_subset_1(esk6_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(esk6_3(esk1_0,X1,esk4_0)),esk6_3(esk1_0,X1,esk4_0))|r1_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk6_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_40])])).
cnf(c_0_52, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_53, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_0)|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_54, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_55, negated_conjecture, (k1_tarski(esk6_3(esk1_0,esk2_0,esk4_0))=k1_xboole_0|r1_lattice3(esk1_0,esk2_0,esk4_0)|r2_hidden(esk6_3(esk1_0,esk2_0,esk4_0),esk3_0)|~m1_subset_1(esk6_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_42]), c_0_40])])).
cnf(c_0_56, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_57, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,X1)|r1_lattice3(esk1_0,esk2_0,esk4_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_42]), c_0_40])])).
cnf(c_0_58, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_0)|r2_hidden(esk6_3(esk1_0,esk2_0,esk4_0),esk3_0)|~m1_subset_1(esk6_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56])])).
cnf(c_0_59, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_60, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk6_3(esk1_0,esk2_0,esk4_0))|r1_lattice3(esk1_0,esk2_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_61, plain, ![X24, X25, X26, X27]:((~r1_lattice3(X24,X26,X27)|r1_lattice3(X24,X25,X27)|~m1_subset_1(X27,u1_struct_0(X24))|~r1_tarski(X25,X26)|~l1_orders_2(X24))&(~r2_lattice3(X24,X26,X27)|r2_lattice3(X24,X25,X27)|~m1_subset_1(X27,u1_struct_0(X24))|~r1_tarski(X25,X26)|~l1_orders_2(X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])).
cnf(c_0_62, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_42]), c_0_40])])).
cnf(c_0_63, plain, (r1_lattice3(X1,X4,X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X4,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_64, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_36]), c_0_42]), c_0_40])])).
fof(c_0_65, plain, ![X18, X19]:((~m1_subset_1(X18,k1_zfmisc_1(X19))|r1_tarski(X18,X19))&(~r1_tarski(X18,X19)|m1_subset_1(X18,k1_zfmisc_1(X19)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_66, negated_conjecture, (r1_lattice3(esk1_0,X1,esk4_0)|~r1_tarski(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_42]), c_0_40])])).
cnf(c_0_67, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_68, negated_conjecture, (r1_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_69, plain, (m1_subset_1(esk7_4(X1,X2,X3,X4),k1_zfmisc_1(X2))|~r2_hidden(X4,X1)|~m1_subset_1(X4,u1_struct_0(X3))|~epred1_3(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_70, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_71, plain, (r1_lattice3(esk1_0,esk7_4(X1,esk2_0,X2,X3),esk4_0)|~epred1_3(X2,esk2_0,X1)|~r2_hidden(X3,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_72, plain, (r1_orders_2(esk1_0,esk4_0,X1)|X1!=k2_yellow_0(esk1_0,esk7_4(X2,esk2_0,X3,X4))|~epred1_3(X3,esk2_0,X2)|~r2_hidden(X4,X2)|~r2_yellow_0(esk1_0,esk7_4(X2,esk2_0,X3,X4))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_42]), c_0_40])])).
cnf(c_0_73, plain, (X1=k2_yellow_0(X2,esk7_4(X3,X4,X2,X1))|~r2_hidden(X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_74, plain, (r2_yellow_0(X1,esk7_4(X2,X3,X1,X4))|~r2_hidden(X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred1_3(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_75, plain, (r1_orders_2(esk1_0,esk4_0,X1)|~epred1_3(esk1_0,esk2_0,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73])]), c_0_74])).
cnf(c_0_76, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,X1)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_75, c_0_34])).
cnf(c_0_77, negated_conjecture, (~r1_lattice3(esk1_0,esk2_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_78, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk6_3(X1,esk3_0,X2))|r1_lattice3(X1,esk3_0,X2)|~m1_subset_1(esk6_3(X1,esk3_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_76, c_0_19])).
cnf(c_0_79, negated_conjecture, (~r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_64])])).
cnf(c_0_80, negated_conjecture, (~m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_78]), c_0_42]), c_0_40])]), c_0_79])).
cnf(c_0_81, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_36]), c_0_42]), c_0_40])]), c_0_79]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 82
# Proof object clause steps            : 57
# Proof object formula steps           : 25
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 38
# Processed clauses                    : 706
# ...of these trivial                  : 0
# ...subsumed                          : 84
# ...remaining for further processing  : 622
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 22
# Backward-rewritten                   : 134
# Generated clauses                    : 1195
# ...of the previous two non-trivial   : 1121
# Contextual simplify-reflections      : 31
# Paramodulations                      : 1135
# Factorizations                       : 0
# Equation resolutions                 : 60
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 428
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 412
# Current number of unprocessed clauses: 474
# ...number of literals in the above   : 4846
# Current number of archived formulas  : 0
# Current number of archived clauses   : 194
# Clause-clause subsumption calls (NU) : 101539
# Rec. Clause-clause subsumption calls : 1790
# Non-unit clause-clause subsumptions  : 137
# Unit Clause-clause subsumption calls : 448
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 77366

# -------------------------------------------------
# User time                : 0.122 s
# System time              : 0.000 s
# Total time               : 0.122 s
# Maximum resident set size: 28832 pages
