# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', t47_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', t48_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', commutativity_k3_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', d5_xboole_0)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', t41_xboole_1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', t98_xboole_1)).
fof(d6_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1)), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', d6_xboole_0)).
fof(l98_xboole_1, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', l98_xboole_1)).
fof(t100_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/80064adea60e4feaa008f2889c7756ab.E_input', t100_xboole_1)).
fof(c_0_9, plain, ![X17, X18]:k4_xboole_0(X17,k3_xboole_0(X17,X18))=k4_xboole_0(X17,X18), inference(variable_rename,[status(thm)],[t47_xboole_1])).
fof(c_0_10, plain, ![X15, X16]:k4_xboole_0(X15,k4_xboole_0(X15,X16))=k3_xboole_0(X15,X16), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_11, plain, ![X9, X10]:k3_xboole_0(X9,X10)=k3_xboole_0(X10,X9), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_12, plain, ![X22, X23, X24, X25, X26, X27, X28, X29]:((((r2_hidden(X25,X22)|~r2_hidden(X25,X24)|X24!=k4_xboole_0(X22,X23))&(~r2_hidden(X25,X23)|~r2_hidden(X25,X24)|X24!=k4_xboole_0(X22,X23)))&(~r2_hidden(X26,X22)|r2_hidden(X26,X23)|r2_hidden(X26,X24)|X24!=k4_xboole_0(X22,X23)))&((~r2_hidden(esk3_3(X27,X28,X29),X29)|(~r2_hidden(esk3_3(X27,X28,X29),X27)|r2_hidden(esk3_3(X27,X28,X29),X28))|X29=k4_xboole_0(X27,X28))&((r2_hidden(esk3_3(X27,X28,X29),X27)|r2_hidden(esk3_3(X27,X28,X29),X29)|X29=k4_xboole_0(X27,X28))&(~r2_hidden(esk3_3(X27,X28,X29),X28)|r2_hidden(esk3_3(X27,X28,X29),X29)|X29=k4_xboole_0(X27,X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_13, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_14]), c_0_14])).
cnf(c_0_19, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22, plain, (~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r2_hidden(X1,k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_19, c_0_18])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, plain, (k4_xboole_0(k4_xboole_0(X1,X1),k4_xboole_0(X1,X1))=k4_xboole_0(k4_xboole_0(X1,X1),X1)), inference(spm,[status(thm)],[c_0_21, c_0_21])).
cnf(c_0_26, plain, (r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk3_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_27, plain, ![X19, X20, X21]:k4_xboole_0(k4_xboole_0(X19,X20),X21)=k4_xboole_0(X19,k2_xboole_0(X20,X21)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_28, plain, ![X11, X12]:k2_xboole_0(X11,X12)=k5_xboole_0(X11,k4_xboole_0(X12,X11)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_29, plain, ![X7, X8]:k5_xboole_0(X7,X8)=k2_xboole_0(k4_xboole_0(X7,X8),k4_xboole_0(X8,X7)), inference(variable_rename,[status(thm)],[d6_xboole_0])).
fof(c_0_30, plain, ![X13, X14]:k5_xboole_0(X13,X14)=k4_xboole_0(k2_xboole_0(X13,X14),k3_xboole_0(X13,X14)), inference(variable_rename,[status(thm)],[l98_xboole_1])).
cnf(c_0_31, plain, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(X2,X3),X2))), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_18]), c_0_21]), c_0_23])).
cnf(c_0_32, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk3_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(X2,X2),X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_25]), c_0_23])).
cnf(c_0_34, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk3_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_26, c_0_24])).
cnf(c_0_35, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (k5_xboole_0(X1,X2)=k4_xboole_0(k2_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (k4_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X1),X3)=k4_xboole_0(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, plain, (k4_xboole_0(k4_xboole_0(X1,X1),X1)=k4_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(X2,X1),k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_37, c_0_36])).
cnf(c_0_43, plain, (k5_xboole_0(X1,X2)=k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_36]), c_0_14])).
cnf(c_0_44, plain, (k4_xboole_0(k4_xboole_0(X1,X1),X2)=k4_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_45, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(assume_negation,[status(cth)],[t100_xboole_1])).
cnf(c_0_46, plain, (k4_xboole_0(k4_xboole_0(X1,k4_xboole_0(X2,X3)),k4_xboole_0(X3,X2))=k4_xboole_0(X1,k5_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_47, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X1,X1)),k4_xboole_0(X1,X1))=k5_xboole_0(X1,k4_xboole_0(X1,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_21]), c_0_44])).
cnf(c_0_48, plain, (k4_xboole_0(k5_xboole_0(X1,k4_xboole_0(X2,X1)),k5_xboole_0(X1,X2))=k4_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_43]), c_0_41]), c_0_39]), c_0_18]), c_0_17])).
fof(c_0_49, negated_conjecture, k4_xboole_0(esk1_0,esk2_0)!=k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_45])])])).
cnf(c_0_50, plain, (~r2_hidden(X1,k4_xboole_0(k4_xboole_0(k4_xboole_0(X2,X3),X4),X2))), inference(spm,[status(thm)],[c_0_31, c_0_41])).
cnf(c_0_51, plain, (k5_xboole_0(X1,k4_xboole_0(X1,X1))=k4_xboole_0(X1,k4_xboole_0(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_47]), c_0_48])).
cnf(c_0_52, plain, (k4_xboole_0(X1,X1)=k4_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_40, c_0_40])).
cnf(c_0_53, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k4_xboole_0(X3,X3)), inference(spm,[status(thm)],[c_0_31, c_0_34])).
cnf(c_0_54, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)!=k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_55, plain, (k5_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(k4_xboole_0(X2,k4_xboole_0(X2,X1)),X1),k4_xboole_0(X1,X2)))=k5_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_42, c_0_21])).
cnf(c_0_56, plain, (k4_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X3),X1)=k4_xboole_0(X4,X4)), inference(spm,[status(thm)],[c_0_50, c_0_34])).
cnf(c_0_57, plain, (k5_xboole_0(X1,k4_xboole_0(X2,X2))=k4_xboole_0(X1,k4_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_58, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X3,X3))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_53]), c_0_17])).
cnf(c_0_59, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)!=k5_xboole_0(esk1_0,k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0)))), inference(rw,[status(thm)],[c_0_54, c_0_14])).
cnf(c_0_60, plain, (k5_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_44]), c_0_57]), c_0_58]), c_0_18]), c_0_17])).
cnf(c_0_61, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_60]), c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 43
# Proof object formula steps           : 19
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 12
# Processed clauses                    : 1668
# ...of these trivial                  : 55
# ...subsumed                          : 1425
# ...remaining for further processing  : 188
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 26
# Generated clauses                    : 15972
# ...of the previous two non-trivial   : 14829
# Contextual simplify-reflections      : 2
# Paramodulations                      : 15959
# Factorizations                       : 10
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 145
#    Positive orientable unit clauses  : 46
#    Positive unorientable unit clauses: 8
#    Negative unit clauses             : 40
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 13026
# ...number of literals in the above   : 18586
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 667
# Rec. Clause-clause subsumption calls : 568
# Non-unit clause-clause subsumptions  : 57
# Unit Clause-clause subsumption calls : 489
# Rewrite failures with RHS unbound    : 342
# BW rewrite match attempts            : 956
# BW rewrite match successes           : 71
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 248553

# -------------------------------------------------
# User time                : 0.105 s
# System time              : 0.016 s
# Total time               : 0.122 s
# Maximum resident set size: 25968 pages
