# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/80840f6dab2c4ed093360a85b5d02599.E_input', t3_xboole_0)).
fof(d20_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(v5_funct_1(X2,X1)<=>![X3]:(r2_hidden(X3,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X3),k1_funct_1(X1,X3)))))), file('data/MPTP/proofs/80840f6dab2c4ed093360a85b5d02599.E_input', d20_funct_1)).
fof(d4_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))=>(X3=k1_funct_1(X1,X2)<=>r2_hidden(k4_tarski(X2,X3),X1)))&(~(r2_hidden(X2,k1_relat_1(X1)))=>(X3=k1_funct_1(X1,X2)<=>X3=k1_xboole_0)))), file('data/MPTP/proofs/80840f6dab2c4ed093360a85b5d02599.E_input', d4_funct_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/80840f6dab2c4ed093360a85b5d02599.E_input', d3_tarski)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/80840f6dab2c4ed093360a85b5d02599.E_input', t65_xboole_1)).
fof(t175_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&(v1_funct_1(X2)&v5_funct_1(X2,X1)))=>r1_tarski(k1_relat_1(X2),k1_relat_1(X1)))), file('data/MPTP/proofs/80840f6dab2c4ed093360a85b5d02599.E_input', t175_funct_1)).
fof(c_0_6, plain, ![X16, X17, X19, X20, X21]:(((r2_hidden(esk4_2(X16,X17),X16)|r1_xboole_0(X16,X17))&(r2_hidden(esk4_2(X16,X17),X17)|r1_xboole_0(X16,X17)))&(~r2_hidden(X21,X19)|~r2_hidden(X21,X20)|~r1_xboole_0(X19,X20))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_7, plain, ![X22, X23, X24]:((~v5_funct_1(X23,X22)|(~r2_hidden(X24,k1_relat_1(X23))|r2_hidden(k1_funct_1(X23,X24),k1_funct_1(X22,X24)))|(~v1_relat_1(X23)|~v1_funct_1(X23))|(~v1_relat_1(X22)|~v1_funct_1(X22)))&((r2_hidden(esk5_2(X22,X23),k1_relat_1(X23))|v5_funct_1(X23,X22)|(~v1_relat_1(X23)|~v1_funct_1(X23))|(~v1_relat_1(X22)|~v1_funct_1(X22)))&(~r2_hidden(k1_funct_1(X23,esk5_2(X22,X23)),k1_funct_1(X22,esk5_2(X22,X23)))|v5_funct_1(X23,X22)|(~v1_relat_1(X23)|~v1_funct_1(X23))|(~v1_relat_1(X22)|~v1_funct_1(X22))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_funct_1])])])])])).
fof(c_0_8, plain, ![X7, X8, X9]:(((X9!=k1_funct_1(X7,X8)|r2_hidden(k4_tarski(X8,X9),X7)|~r2_hidden(X8,k1_relat_1(X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(~r2_hidden(k4_tarski(X8,X9),X7)|X9=k1_funct_1(X7,X8)|~r2_hidden(X8,k1_relat_1(X7))|(~v1_relat_1(X7)|~v1_funct_1(X7))))&((X9!=k1_funct_1(X7,X8)|X9=k1_xboole_0|r2_hidden(X8,k1_relat_1(X7))|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(X9!=k1_xboole_0|X9=k1_funct_1(X7,X8)|r2_hidden(X8,k1_relat_1(X7))|(~v1_relat_1(X7)|~v1_funct_1(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_funct_1])])])])])).
cnf(c_0_9, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (r2_hidden(k1_funct_1(X1,X3),k1_funct_1(X2,X3))|~v5_funct_1(X1,X2)|~r2_hidden(X3,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, ![X10, X11, X12, X13, X14]:((~r1_tarski(X10,X11)|(~r2_hidden(X12,X10)|r2_hidden(X12,X11)))&((r2_hidden(esk3_2(X13,X14),X13)|r1_tarski(X13,X14))&(~r2_hidden(esk3_2(X13,X14),X14)|r1_tarski(X13,X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_12, plain, (X1=k1_xboole_0|r2_hidden(X3,k1_relat_1(X2))|X1!=k1_funct_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (~r2_hidden(k1_funct_1(X1,X2),X3)|~r2_hidden(X2,k1_relat_1(X1))|~r1_xboole_0(X3,k1_funct_1(X4,X2))|~v5_funct_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_14, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (k1_funct_1(X1,X2)=k1_xboole_0|r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X6]:r1_xboole_0(X6,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_17, plain, (~r2_hidden(X1,k1_relat_1(X2))|~r1_xboole_0(k1_funct_1(X3,X1),k1_funct_1(X4,X1))|~v5_funct_1(X2,X4)|~v5_funct_1(X2,X3)|~v1_funct_1(X4)|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_13, c_0_10])).
cnf(c_0_18, plain, (k1_funct_1(X1,esk3_2(X2,k1_relat_1(X1)))=k1_xboole_0|r1_tarski(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&(v1_funct_1(X2)&v5_funct_1(X2,X1)))=>r1_tarski(k1_relat_1(X2),k1_relat_1(X1))))), inference(assume_negation,[status(cth)],[t175_funct_1])).
cnf(c_0_21, plain, (r1_tarski(X1,k1_relat_1(X2))|~r2_hidden(esk3_2(X1,k1_relat_1(X2)),k1_relat_1(X3))|~v5_funct_1(X3,X2)|~v5_funct_1(X3,X4)|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X4)|~v1_relat_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
cnf(c_0_22, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_23, negated_conjecture, ((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&(v1_funct_1(esk2_0)&v5_funct_1(esk2_0,esk1_0)))&~r1_tarski(k1_relat_1(esk2_0),k1_relat_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
cnf(c_0_24, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~v5_funct_1(X1,X2)|~v5_funct_1(X1,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25, negated_conjecture, (v5_funct_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_26, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, negated_conjecture, (~r1_tarski(k1_relat_1(esk2_0),k1_relat_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, negated_conjecture, (r1_tarski(k1_relat_1(esk2_0),k1_relat_1(X1))|~v5_funct_1(esk2_0,X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29])])).
cnf(c_0_32, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_25]), c_0_27]), c_0_29])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 20
# Proof object formula steps           : 13
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 453
# ...of these trivial                  : 1
# ...subsumed                          : 284
# ...remaining for further processing  : 168
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 0
# Generated clauses                    : 786
# ...of the previous two non-trivial   : 760
# Contextual simplify-reflections      : 6
# Paramodulations                      : 783
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 134
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 123
# Current number of unprocessed clauses: 305
# ...number of literals in the above   : 3525
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 12361
# Rec. Clause-clause subsumption calls : 1177
# Non-unit clause-clause subsumptions  : 265
# Unit Clause-clause subsumption calls : 50
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 24009

# -------------------------------------------------
# User time                : 0.022 s
# System time              : 0.003 s
# Total time               : 0.025 s
# Maximum resident set size: 28660 pages
