# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', t69_enumset1)).
fof(t77_enumset1, axiom, ![X1, X2]:k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', t77_enumset1)).
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', t39_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', fc2_struct_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', t24_orders_2)).
fof(t18_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', t18_yellow_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', dt_l1_orders_2)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', dt_k1_yellow_0)).
fof(t38_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', t38_yellow_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', d9_lattice3)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', d2_tarski)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', t7_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', d8_lattice3)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/80fbe98a295547119fb26d75dd61aa29.E_input', d10_yellow_0)).
fof(c_0_16, plain, ![X6, X7]:(v1_xboole_0(X6)|~m1_subset_1(X7,X6)|k6_domain_1(X6,X7)=k1_tarski(X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_17, plain, ![X8]:k2_tarski(X8,X8)=k1_tarski(X8), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_18, plain, ![X9, X10]:k2_enumset1(X9,X9,X9,X10)=k2_tarski(X9,X10), inference(variable_rename,[status(thm)],[t77_enumset1])).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_20, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
cnf(c_0_21, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
fof(c_0_25, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X13]:(v2_struct_0(X13)|~l1_struct_0(X13)|~v1_xboole_0(u1_struct_0(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_27, plain, (k6_domain_1(X1,X2)=k2_enumset1(X2,X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_29, plain, ![X56, X57, X58, X59, X60]:(((r1_orders_2(X58,X57,X59)|(X59!=k10_lattice3(X58,X57,X56)|~r1_yellow_0(X58,k2_tarski(X57,X56)))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56))&((r1_orders_2(X58,X56,X59)|(X59!=k10_lattice3(X58,X57,X56)|~r1_yellow_0(X58,k2_tarski(X57,X56)))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56))&(~m1_subset_1(X60,u1_struct_0(X58))|(~r1_orders_2(X58,X57,X60)|~r1_orders_2(X58,X56,X60)|r1_orders_2(X58,X59,X60))|(X59!=k10_lattice3(X58,X57,X56)|~r1_yellow_0(X58,k2_tarski(X57,X56)))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56))))&(((X59=k10_lattice3(X58,X57,X56)|(m1_subset_1(esk8_4(X56,X57,X58,X59),u1_struct_0(X58))|~r1_orders_2(X58,X56,X59)|~r1_orders_2(X58,X57,X59))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56))&(r1_yellow_0(X58,k2_tarski(X57,X56))|(m1_subset_1(esk8_4(X56,X57,X58,X59),u1_struct_0(X58))|~r1_orders_2(X58,X56,X59)|~r1_orders_2(X58,X57,X59))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56)))&((((X59=k10_lattice3(X58,X57,X56)|(r1_orders_2(X58,X57,esk8_4(X56,X57,X58,X59))|~r1_orders_2(X58,X56,X59)|~r1_orders_2(X58,X57,X59))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56))&(r1_yellow_0(X58,k2_tarski(X57,X56))|(r1_orders_2(X58,X57,esk8_4(X56,X57,X58,X59))|~r1_orders_2(X58,X56,X59)|~r1_orders_2(X58,X57,X59))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56)))&((X59=k10_lattice3(X58,X57,X56)|(r1_orders_2(X58,X56,esk8_4(X56,X57,X58,X59))|~r1_orders_2(X58,X56,X59)|~r1_orders_2(X58,X57,X59))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56))&(r1_yellow_0(X58,k2_tarski(X57,X56))|(r1_orders_2(X58,X56,esk8_4(X56,X57,X58,X59))|~r1_orders_2(X58,X56,X59)|~r1_orders_2(X58,X57,X59))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56))))&((X59=k10_lattice3(X58,X57,X56)|(~r1_orders_2(X58,X59,esk8_4(X56,X57,X58,X59))|~r1_orders_2(X58,X56,X59)|~r1_orders_2(X58,X57,X59))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56))&(r1_yellow_0(X58,k2_tarski(X57,X56))|(~r1_orders_2(X58,X59,esk8_4(X56,X57,X58,X59))|~r1_orders_2(X58,X56,X59)|~r1_orders_2(X58,X57,X59))|~m1_subset_1(X59,u1_struct_0(X58))|~epred1_3(X58,X57,X56)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])])).
fof(c_0_30, plain, ![X14, X15]:(v2_struct_0(X14)|(~v3_orders_2(X14)|~l1_orders_2(X14))|(~m1_subset_1(X15,u1_struct_0(X14))|r1_orders_2(X14,X15,X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_31, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t18_yellow_0, c_0_20])).
cnf(c_0_32, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_35, plain, ![X19]:(~l1_orders_2(X19)|l1_struct_0(X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_36, plain, (X1=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,esk8_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X4,esk8_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_40, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_41, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X3,esk8_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_42, plain, ![X16, X17, X18]:(~v5_orders_2(X16)|~l1_orders_2(X16)|(~m1_subset_1(X17,u1_struct_0(X16))|(~m1_subset_1(X18,u1_struct_0(X16))|epred1_3(X16,X17,X18)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])).
fof(c_0_43, plain, ![X20, X21, X22, X23]:(((r2_lattice3(X20,X21,X22)|X22!=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&(~m1_subset_1(X23,u1_struct_0(X20))|(~r2_lattice3(X20,X21,X23)|r1_orders_2(X20,X22,X23))|X22!=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20)))&((m1_subset_1(esk3_3(X20,X21,X22),u1_struct_0(X20))|~r2_lattice3(X20,X21,X22)|X22=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&((r2_lattice3(X20,X21,esk3_3(X20,X21,X22))|~r2_lattice3(X20,X21,X22)|X22=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))&(~r1_orders_2(X20,X22,esk3_3(X20,X21,X22))|~r2_lattice3(X20,X21,X22)|X22=k1_yellow_0(X20,X21)|~r1_yellow_0(X20,X21)|~m1_subset_1(X22,u1_struct_0(X20))|~l1_orders_2(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_44, plain, ![X25, X26]:(~l1_orders_2(X25)|m1_subset_1(k1_yellow_0(X25,X26),u1_struct_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_45, plain, (r1_orders_2(X2,X5,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r1_orders_2(X2,X4,X1)|X5!=k10_lattice3(X2,X3,X4)|~r1_yellow_0(X2,k2_tarski(X3,X4))|~m1_subset_1(X5,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_46, plain, ![X27, X28]:((r1_yellow_0(X27,k6_domain_1(u1_struct_0(X27),X28))|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v3_orders_2(X27)|(~v5_orders_2(X27)|~l1_orders_2(X27)))))&(r2_yellow_0(X27,k6_domain_1(u1_struct_0(X27),X28))|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v3_orders_2(X27)|(~v5_orders_2(X27)|~l1_orders_2(X27)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_yellow_0])])])])])).
cnf(c_0_47, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_48, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_49, plain, (k10_lattice3(X1,X2,X3)=X3|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])]), c_0_34])).
cnf(c_0_51, plain, (k10_lattice3(X1,X2,X3)=X2|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X2)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_36, c_0_41])).
cnf(c_0_52, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_54, plain, ![X29, X30, X31, X32]:((~r2_lattice3(X29,X30,X31)|(~m1_subset_1(X32,u1_struct_0(X29))|(~r2_hidden(X32,X30)|r1_orders_2(X29,X32,X31)))|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&((m1_subset_1(esk4_3(X29,X30,X31),u1_struct_0(X29))|r2_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&((r2_hidden(esk4_3(X29,X30,X31),X30)|r2_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))&(~r1_orders_2(X29,esk4_3(X29,X30,X31),X31)|r2_lattice3(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~l1_orders_2(X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_55, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_56, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_57, plain, (r1_orders_2(X2,X5,X1)|X5!=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~epred1_3(X2,X3,X4)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_yellow_0(X2,k2_enumset1(X3,X3,X3,X4))), inference(rw,[status(thm)],[c_0_45, c_0_23])).
cnf(c_0_58, plain, (r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_59, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_40])])).
cnf(c_0_60, negated_conjecture, (k10_lattice3(esk1_0,X1,X2)=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_61, negated_conjecture, (k10_lattice3(esk1_0,X1,X2)=X1|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_50])).
cnf(c_0_62, negated_conjecture, (epred1_3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_40])])).
cnf(c_0_63, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_55]), c_0_56])).
cnf(c_0_65, plain, (r1_orders_2(X1,k10_lattice3(X1,X2,X3),X4)|~epred1_3(X1,X2,X3)|~r1_yellow_0(X1,k2_enumset1(X2,X2,X2,X3))|~r1_orders_2(X1,X3,X4)|~r1_orders_2(X1,X2,X4)|~m1_subset_1(k10_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_66, negated_conjecture, (r1_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_40]), c_0_53]), c_0_39]), c_0_28])]), c_0_34])).
cnf(c_0_67, negated_conjecture, (X1=X2|~r1_orders_2(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
cnf(c_0_68, plain, (r1_orders_2(X1,X2,k1_yellow_0(X1,X3))|~r2_hidden(X2,X3)|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_56])).
cnf(c_0_69, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_70, plain, ![X34, X35, X36, X37, X38, X39, X40, X41]:(((~r2_hidden(X37,X36)|(X37=X34|X37=X35)|X36!=k2_tarski(X34,X35))&((X38!=X34|r2_hidden(X38,X36)|X36!=k2_tarski(X34,X35))&(X38!=X35|r2_hidden(X38,X36)|X36!=k2_tarski(X34,X35))))&(((esk5_3(X39,X40,X41)!=X39|~r2_hidden(esk5_3(X39,X40,X41),X41)|X41=k2_tarski(X39,X40))&(esk5_3(X39,X40,X41)!=X40|~r2_hidden(esk5_3(X39,X40,X41),X41)|X41=k2_tarski(X39,X40)))&(r2_hidden(esk5_3(X39,X40,X41),X41)|(esk5_3(X39,X40,X41)=X39|esk5_3(X39,X40,X41)=X40)|X41=k2_tarski(X39,X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_71, negated_conjecture, (r1_orders_2(esk1_0,k10_lattice3(esk1_0,esk2_0,esk2_0),X1)|~epred1_3(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(k10_lattice3(esk1_0,esk2_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_72, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|~r2_hidden(X1,X2)|~r1_yellow_0(esk1_0,X2)|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X2),X1)|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_40])])).
cnf(c_0_73, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_69]), c_0_56])).
cnf(c_0_74, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_75, negated_conjecture, (k10_lattice3(esk1_0,esk2_0,esk2_0)=X1|~epred1_3(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,X1,k10_lattice3(esk1_0,esk2_0,esk2_0))|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(k10_lattice3(esk1_0,esk2_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_67, c_0_71])).
cnf(c_0_76, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_77, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|~r2_hidden(X1,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_40])])).
cnf(c_0_78, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X4,X4,X4,X2)), inference(rw,[status(thm)],[c_0_74, c_0_23])).
fof(c_0_79, plain, ![X43, X44, X45]:((~r1_lattice3(X43,k1_tarski(X45),X44)|r1_orders_2(X43,X44,X45)|~m1_subset_1(X45,u1_struct_0(X43))|~m1_subset_1(X44,u1_struct_0(X43))|~l1_orders_2(X43))&((~r1_orders_2(X43,X44,X45)|r1_lattice3(X43,k1_tarski(X45),X44)|~m1_subset_1(X45,u1_struct_0(X43))|~m1_subset_1(X44,u1_struct_0(X43))|~l1_orders_2(X43))&((~r2_lattice3(X43,k1_tarski(X45),X44)|r1_orders_2(X43,X45,X44)|~m1_subset_1(X45,u1_struct_0(X43))|~m1_subset_1(X44,u1_struct_0(X43))|~l1_orders_2(X43))&(~r1_orders_2(X43,X45,X44)|r2_lattice3(X43,k1_tarski(X45),X44)|~m1_subset_1(X45,u1_struct_0(X43))|~m1_subset_1(X44,u1_struct_0(X43))|~l1_orders_2(X43))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_80, negated_conjecture, (esk2_0=X1|~epred1_3(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_60]), c_0_28])])).
cnf(c_0_81, plain, (X1=X4|X1=X3|X2!=k2_enumset1(X3,X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_76, c_0_23])).
cnf(c_0_82, negated_conjecture, (X1=k1_yellow_0(esk1_0,X2)|~r2_hidden(X1,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_56]), c_0_40])])).
cnf(c_0_83, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_78])])).
cnf(c_0_84, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_85, negated_conjecture, (esk2_0=X1|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_62]), c_0_28])])).
cnf(c_0_86, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_87, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_enumset1(X3,X3,X3,X2))), inference(er,[status(thm)],[c_0_81])).
cnf(c_0_88, plain, (k2_enumset1(k1_yellow_0(X1,X2),k1_yellow_0(X1,X2),k1_yellow_0(X1,X2),k1_yellow_0(X1,X2))=k6_domain_1(u1_struct_0(X1),k1_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_27, c_0_56])).
cnf(c_0_89, negated_conjecture, (k1_yellow_0(esk1_0,k2_enumset1(X1,X1,X1,X2))=X2|~r2_lattice3(esk1_0,k2_enumset1(X1,X1,X1,X2),X2)|~r1_yellow_0(esk1_0,k2_enumset1(X1,X1,X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_90, plain, (r2_lattice3(X1,k2_enumset1(X2,X2,X2,X2),X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_22]), c_0_23])).
cnf(c_0_91, negated_conjecture, (esk2_0=X1|~r1_orders_2(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_50]), c_0_28])])).
cnf(c_0_92, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k2_enumset1(X2,X2,X2,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_22]), c_0_23])).
cnf(c_0_93, plain, (X1=k1_yellow_0(X2,X3)|v1_xboole_0(u1_struct_0(X2))|~r2_hidden(X1,k6_domain_1(u1_struct_0(X2),k1_yellow_0(X2,X3)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_94, negated_conjecture, (k1_yellow_0(esk1_0,k2_enumset1(X1,X1,X1,X1))=X1|~r1_yellow_0(esk1_0,k2_enumset1(X1,X1,X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_40])]), c_0_50])).
cnf(c_0_95, negated_conjecture, (k1_yellow_0(esk1_0,X1)=esk2_0|~r2_lattice3(esk1_0,X1,esk2_0)|~r1_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_73]), c_0_40]), c_0_28])])).
cnf(c_0_96, plain, (r1_orders_2(X1,X2,k1_yellow_0(X1,k2_enumset1(X2,X2,X2,X2)))|~r1_yellow_0(X1,k2_enumset1(X2,X2,X2,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_64]), c_0_56])).
cnf(c_0_97, negated_conjecture, (X1=X2|v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(X1,k6_domain_1(u1_struct_0(esk1_0),X2))|~r1_yellow_0(esk1_0,k2_enumset1(X2,X2,X2,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_40])])).
fof(c_0_98, plain, ![X46, X47, X48, X49]:((~r1_lattice3(X46,X47,X48)|(~m1_subset_1(X49,u1_struct_0(X46))|(~r2_hidden(X49,X47)|r1_orders_2(X46,X48,X49)))|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46))&((m1_subset_1(esk6_3(X46,X47,X48),u1_struct_0(X46))|r1_lattice3(X46,X47,X48)|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46))&((r2_hidden(esk6_3(X46,X47,X48),X47)|r1_lattice3(X46,X47,X48)|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46))&(~r1_orders_2(X46,X48,esk6_3(X46,X47,X48))|r1_lattice3(X46,X47,X48)|~m1_subset_1(X48,u1_struct_0(X46))|~l1_orders_2(X46))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_99, negated_conjecture, (k1_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~r2_lattice3(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|~m1_subset_1(k1_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_66]), c_0_40]), c_0_28])])).
cnf(c_0_100, negated_conjecture, (X1=esk2_0|v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_59]), c_0_66]), c_0_28])])).
cnf(c_0_101, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_102, negated_conjecture, (k1_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~r2_lattice3(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_56]), c_0_40])])).
cnf(c_0_103, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_104, plain, ![X51, X52, X53, X54]:(((r1_lattice3(X51,X52,X53)|X53!=k2_yellow_0(X51,X52)|~r2_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51))&(~m1_subset_1(X54,u1_struct_0(X51))|(~r1_lattice3(X51,X52,X54)|r1_orders_2(X51,X54,X53))|X53!=k2_yellow_0(X51,X52)|~r2_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51)))&((m1_subset_1(esk7_3(X51,X52,X53),u1_struct_0(X51))|~r1_lattice3(X51,X52,X53)|X53=k2_yellow_0(X51,X52)|~r2_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51))&((r1_lattice3(X51,X52,esk7_3(X51,X52,X53))|~r1_lattice3(X51,X52,X53)|X53=k2_yellow_0(X51,X52)|~r2_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51))&(~r1_orders_2(X51,esk7_3(X51,X52,X53),X53)|~r1_lattice3(X51,X52,X53)|X53=k2_yellow_0(X51,X52)|~r2_yellow_0(X51,X52)|~m1_subset_1(X53,u1_struct_0(X51))|~l1_orders_2(X51))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
cnf(c_0_105, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_106, negated_conjecture, (esk6_3(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),X2)=esk2_0|r1_lattice3(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),X2)|v1_xboole_0(u1_struct_0(esk1_0))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_107, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)|~r2_lattice3(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_102]), c_0_66]), c_0_40]), c_0_28])])).
cnf(c_0_108, negated_conjecture, (r2_lattice3(esk1_0,k2_enumset1(X1,X1,X1,X1),X1)|~r1_yellow_0(esk1_0,k2_enumset1(X1,X1,X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_94]), c_0_40])])).
cnf(c_0_109, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_110, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk4_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_111, negated_conjecture, (esk4_3(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),X2)=esk2_0|r2_lattice3(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),X2)|v1_xboole_0(u1_struct_0(esk1_0))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_100, c_0_103])).
cnf(c_0_112, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_113, plain, (r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_114, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_115, negated_conjecture, (r1_lattice3(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),X2)|v1_xboole_0(u1_struct_0(esk1_0))|~r1_orders_2(X1,X2,esk2_0)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_116, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_66]), c_0_28])])).
cnf(c_0_117, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X2,X4)), inference(rw,[status(thm)],[c_0_109, c_0_23])).
cnf(c_0_118, negated_conjecture, (r2_lattice3(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),X2)|v1_xboole_0(u1_struct_0(esk1_0))|~r1_orders_2(X1,esk2_0,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_119, plain, (X3=k2_yellow_0(X1,X2)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_120, plain, (X1=k2_yellow_0(X2,X3)|r1_orders_2(X2,esk7_3(X2,X3,X1),X4)|~r1_lattice3(X2,X3,X1)|~r2_hidden(X4,X3)|~r2_yellow_0(X2,X3)|~l1_orders_2(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_114])).
cnf(c_0_121, negated_conjecture, (r1_lattice3(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_40]), c_0_28])])).
cnf(c_0_122, plain, (r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_123, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_117])])).
cnf(c_0_124, negated_conjecture, (r2_lattice3(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_50]), c_0_40]), c_0_28])])).
cnf(c_0_125, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_126, plain, (X1=k2_yellow_0(X2,X3)|~r1_lattice3(X2,X3,X1)|~r2_hidden(X1,X3)|~r2_yellow_0(X2,X3)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_127, negated_conjecture, (r1_lattice3(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_121]), c_0_34])).
cnf(c_0_128, negated_conjecture, (r2_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_59]), c_0_40]), c_0_53]), c_0_39]), c_0_28])]), c_0_34])).
cnf(c_0_129, negated_conjecture, (k1_yellow_0(esk1_0,k2_enumset1(X1,X1,X1,X2))=X1|~r2_lattice3(esk1_0,k2_enumset1(X1,X1,X1,X2),X1)|~r1_yellow_0(esk1_0,k2_enumset1(X1,X1,X1,X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_82, c_0_123])).
cnf(c_0_130, negated_conjecture, (r2_lattice3(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_124]), c_0_34])).
cnf(c_0_131, negated_conjecture, (k1_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_59]), c_0_59])).
cnf(c_0_132, negated_conjecture, (k2_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_123]), c_0_128]), c_0_40]), c_0_28])])).
cnf(c_0_133, negated_conjecture, (k1_yellow_0(esk1_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_66]), c_0_28])])).
cnf(c_0_134, negated_conjecture, (~l1_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_133])).
cnf(c_0_135, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_48]), c_0_40])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 136
# Proof object clause steps            : 99
# Proof object formula steps           : 37
# Proof object conjectures             : 51
# Proof object clause conjectures      : 48
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 49
# Proof object simplifying inferences  : 103
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 55
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 53
# Processed clauses                    : 713
# ...of these trivial                  : 0
# ...subsumed                          : 309
# ...remaining for further processing  : 404
# Other redundant clauses eliminated   : 32
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 27
# Generated clauses                    : 2842
# ...of the previous two non-trivial   : 2589
# Contextual simplify-reflections      : 30
# Paramodulations                      : 2800
# Factorizations                       : 12
# Equation resolutions                 : 32
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 291
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 279
# Current number of unprocessed clauses: 1972
# ...number of literals in the above   : 16358
# Current number of archived formulas  : 0
# Current number of archived clauses   : 105
# Clause-clause subsumption calls (NU) : 16782
# Rec. Clause-clause subsumption calls : 2052
# Non-unit clause-clause subsumptions  : 362
# Unit Clause-clause subsumption calls : 198
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 102455

# -------------------------------------------------
# User time                : 0.076 s
# System time              : 0.000 s
# Total time               : 0.076 s
# Maximum resident set size: 28672 pages
