# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t25_zfmisc_1, conjecture, ![X1, X2, X3]:~((r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))&(~(X1=X2)&~(X1=X3)))), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t25_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t75_enumset1)).
fof(t67_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k2_tarski(X7,X8)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t67_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t100_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t28_xboole_1)).
fof(t125_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t125_enumset1)).
fof(t64_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_enumset1(X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t64_enumset1)).
fof(t129_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t129_enumset1)).
fof(t127_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t127_enumset1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t92_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t5_boole)).
fof(t60_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_tarski(X6,X7)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t60_enumset1)).
fof(t66_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t66_enumset1)).
fof(t107_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t107_enumset1)).
fof(t119_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X4,X2,X1), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t119_enumset1)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', commutativity_k5_xboole_0)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t4_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', d1_xboole_0)).
fof(l97_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', l97_xboole_1)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t91_xboole_1)).
fof(fc3_xboole_0, axiom, ![X1, X2]:~(v1_xboole_0(k2_tarski(X1,X2))), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', fc3_xboole_0)).
fof(t136_enumset1, axiom, ![X1, X2, X3]:~((~(X1=X2)&(~(X1=X3)&~(k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3))))), file('data/MPTP/proofs/810a9e7b1f084366a857cfe581b6d829.E_input', t136_enumset1)).
fof(c_0_29, negated_conjecture, ~(![X1, X2, X3]:~((r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))&(~(X1=X2)&~(X1=X3))))), inference(assume_negation,[status(cth)],[t25_zfmisc_1])).
fof(c_0_30, negated_conjecture, (r1_tarski(k1_tarski(esk1_0),k2_tarski(esk2_0,esk3_0))&(esk1_0!=esk2_0&esk1_0!=esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_29])])])])).
fof(c_0_31, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_32, plain, ![X14, X15]:k1_enumset1(X14,X14,X15)=k2_tarski(X14,X15), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_33, plain, ![X16, X17, X18]:k2_enumset1(X16,X16,X17,X18)=k1_enumset1(X16,X17,X18), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_34, plain, ![X19, X20, X21, X22]:k3_enumset1(X19,X19,X20,X21,X22)=k2_enumset1(X19,X20,X21,X22), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_35, plain, ![X23, X24, X25, X26, X27]:k4_enumset1(X23,X23,X24,X25,X26,X27)=k3_enumset1(X23,X24,X25,X26,X27), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_36, plain, ![X28, X29, X30, X31, X32, X33]:k5_enumset1(X28,X28,X29,X30,X31,X32,X33)=k4_enumset1(X28,X29,X30,X31,X32,X33), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_37, plain, ![X34, X35, X36, X37, X38, X39, X40]:k6_enumset1(X34,X34,X35,X36,X37,X38,X39,X40)=k5_enumset1(X34,X35,X36,X37,X38,X39,X40), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_38, plain, ![X41, X42, X43, X44, X45, X46, X47, X48]:k6_enumset1(X41,X42,X43,X44,X45,X46,X47,X48)=k2_xboole_0(k4_enumset1(X41,X42,X43,X44,X45,X46),k2_tarski(X47,X48)), inference(variable_rename,[status(thm)],[t67_enumset1])).
fof(c_0_39, plain, ![X49, X50]:k2_xboole_0(X49,X50)=k5_xboole_0(X49,k4_xboole_0(X50,X49)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_40, plain, ![X51, X52]:k4_xboole_0(X51,X52)=k5_xboole_0(X51,k3_xboole_0(X51,X52)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_41, plain, ![X53, X54]:(~r1_tarski(X53,X54)|k3_xboole_0(X53,X54)=X53), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_42, negated_conjecture, (r1_tarski(k1_tarski(esk1_0),k2_tarski(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_43, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_50, plain, ![X55, X56, X57, X58]:k2_enumset1(X55,X56,X57,X58)=k2_enumset1(X58,X57,X56,X55), inference(variable_rename,[status(thm)],[t125_enumset1])).
fof(c_0_51, plain, ![X59, X60, X61, X62, X63, X64, X65, X66]:k6_enumset1(X59,X60,X61,X62,X63,X64,X65,X66)=k2_xboole_0(k1_enumset1(X59,X60,X61),k3_enumset1(X62,X63,X64,X65,X66)), inference(variable_rename,[status(thm)],[t64_enumset1])).
fof(c_0_52, plain, ![X67, X68, X69, X70, X71, X72, X73, X74, X75]:k7_enumset1(X67,X68,X69,X70,X71,X72,X73,X74,X75)=k2_xboole_0(k1_enumset1(X67,X68,X69),k4_enumset1(X70,X71,X72,X73,X74,X75)), inference(variable_rename,[status(thm)],[t129_enumset1])).
fof(c_0_53, plain, ![X76, X77, X78, X79, X80, X81, X82, X83, X84]:k7_enumset1(X76,X77,X78,X79,X80,X81,X82,X83,X84)=k2_xboole_0(k1_tarski(X76),k6_enumset1(X77,X78,X79,X80,X81,X82,X83,X84)), inference(variable_rename,[status(thm)],[t127_enumset1])).
cnf(c_0_54, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k2_tarski(X7,X8))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_55, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_56, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_57, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
fof(c_0_59, plain, ![X85]:k5_xboole_0(X85,X85)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_60, plain, ![X86]:k5_xboole_0(X86,k1_xboole_0)=X86, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_61, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_62, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_enumset1(X1,X2,X3),k3_enumset1(X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_63, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_tarski(X1),k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6),k5_xboole_0(k6_enumset1(X7,X7,X7,X7,X7,X7,X7,X8),k3_xboole_0(k6_enumset1(X7,X7,X7,X7,X7,X7,X7,X8),k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_44]), c_0_55]), c_0_56]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_66, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_67, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_68, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_69, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X4,X4,X4,X4,X4,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_70, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k5_xboole_0(k6_enumset1(X4,X4,X4,X4,X5,X6,X7,X8),k3_xboole_0(k6_enumset1(X4,X4,X4,X4,X5,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_55]), c_0_56]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_71, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k5_xboole_0(k6_enumset1(X4,X4,X4,X5,X6,X7,X8,X9),k3_xboole_0(k6_enumset1(X4,X4,X4,X5,X6,X7,X8,X9),k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_55]), c_0_56]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
fof(c_0_72, plain, ![X87, X88, X89, X90, X91, X92, X93]:k5_enumset1(X87,X88,X89,X90,X91,X92,X93)=k2_xboole_0(k3_enumset1(X87,X88,X89,X90,X91),k2_tarski(X92,X93)), inference(variable_rename,[status(thm)],[t60_enumset1])).
fof(c_0_73, plain, ![X94, X95, X96, X97, X98, X99, X100, X101]:k6_enumset1(X94,X95,X96,X97,X98,X99,X100,X101)=k2_xboole_0(k3_enumset1(X94,X95,X96,X97,X98),k1_enumset1(X99,X100,X101)), inference(variable_rename,[status(thm)],[t66_enumset1])).
cnf(c_0_74, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k5_xboole_0(k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9),k3_xboole_0(k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X9),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_43]), c_0_44]), c_0_55]), c_0_56]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_75, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_68]), c_0_69])).
cnf(c_0_76, plain, (k7_enumset1(X1,X2,X3,X4,X4,X5,X6,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[c_0_70, c_0_71])).
fof(c_0_77, plain, ![X102, X103, X104, X105]:k2_enumset1(X102,X103,X104,X105)=k2_enumset1(X102,X105,X104,X103), inference(variable_rename,[status(thm)],[t107_enumset1])).
cnf(c_0_78, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_tarski(X6,X7))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_79, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_enumset1(X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_80, negated_conjecture, (k6_enumset1(X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_74]), c_0_76]), c_0_76])).
cnf(c_0_81, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_82, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X6,X7),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X6,X7),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_44]), c_0_55]), c_0_56]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_83, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X7,X8),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_55]), c_0_56]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_84, negated_conjecture, (k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)), inference(rw,[status(thm)],[c_0_75, c_0_80])).
cnf(c_0_85, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0,X1,X2)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0,X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_75]), c_0_65])).
cnf(c_0_86, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X4,X3,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_87, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X6,X7)=k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_88, negated_conjecture, (k6_enumset1(X1,esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_84]), c_0_74]), c_0_76]), c_0_76])).
cnf(c_0_89, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1,X2,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0,X2,X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_90, negated_conjecture, (k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_84])).
cnf(c_0_91, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_89]), c_0_87]), c_0_86]), c_0_87]), c_0_84])).
cnf(c_0_92, negated_conjecture, (k6_enumset1(X1,esk2_0,esk1_0,esk1_0,esk1_0,esk3_0,esk3_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_90]), c_0_74]), c_0_76]), c_0_76])).
cnf(c_0_93, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_91]), c_0_74]), c_0_76]), c_0_76])).
cnf(c_0_94, negated_conjecture, (k6_enumset1(X1,esk2_0,esk1_0,esk1_0,esk1_0,esk3_0,esk3_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_95, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k6_enumset1(X4,X4,X3,X2,X1,X5,X6,X7)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_69]), c_0_83])).
cnf(c_0_96, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(X1,X1,esk2_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_87, c_0_94])).
cnf(c_0_97, plain, (k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)=k6_enumset1(X4,X4,X4,X3,X2,X1,X5,X6)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_69]), c_0_65])).
cnf(c_0_98, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,X1,esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
fof(c_0_99, plain, ![X106, X107, X108, X109]:k2_enumset1(X106,X107,X108,X109)=k2_enumset1(X108,X109,X107,X106), inference(variable_rename,[status(thm)],[t119_enumset1])).
cnf(c_0_100, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_101, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X4,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_102, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk2_0,esk1_0,esk3_0,esk3_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_87, c_0_100])).
cnf(c_0_103, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X3,X3,X3,X3,X3,X4,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49])).
cnf(c_0_104, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk2_0,esk1_0,esk3_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,X1,esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_98, c_0_102])).
cnf(c_0_105, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X5,X6)=k6_enumset1(X3,X3,X3,X4,X2,X1,X5,X6)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_103]), c_0_83])).
cnf(c_0_106, plain, (k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X7,X8,X8),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X7,X8,X8),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))=k6_enumset1(X1,X2,X3,X4,X5,X8,X7,X6)), inference(spm,[status(thm)],[c_0_83, c_0_69])).
cnf(c_0_107, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk2_0,esk1_0,esk3_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
fof(c_0_108, plain, ![X110, X111]:k5_xboole_0(X110,X111)=k5_xboole_0(X111,X110), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_109, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X8,X7)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_103]), c_0_106])).
cnf(c_0_110, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0,esk3_0,esk2_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_102, c_0_107])).
fof(c_0_111, plain, ![X112, X113, X115, X116, X117]:((r1_xboole_0(X112,X113)|r2_hidden(esk4_2(X112,X113),k3_xboole_0(X112,X113)))&(~r2_hidden(X117,k3_xboole_0(X115,X116))|~r1_xboole_0(X115,X116))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_112, plain, ![X118, X119, X120]:((~v1_xboole_0(X118)|~r2_hidden(X119,X118))&(r2_hidden(esk5_1(X120),X120)|v1_xboole_0(X120))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_113, plain, ![X122, X123]:r1_xboole_0(k3_xboole_0(X122,X123),k5_xboole_0(X122,X123)), inference(variable_rename,[status(thm)],[l97_xboole_1])).
fof(c_0_114, plain, ![X124, X125, X126]:k5_xboole_0(k5_xboole_0(X124,X125),X126)=k5_xboole_0(X124,k5_xboole_0(X125,X126)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_115, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_116, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k6_enumset1(X3,X3,X4,X2,X1,X5,X6,X7)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_103]), c_0_83])).
cnf(c_0_117, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk2_0,esk1_0,esk3_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_109]), c_0_109])).
cnf(c_0_118, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0,esk2_0,esk3_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_110, c_0_109])).
cnf(c_0_119, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_100, c_0_109])).
cnf(c_0_120, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk2_0,X1,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_98, c_0_109])).
cnf(c_0_121, negated_conjecture, (k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_84, c_0_91])).
cnf(c_0_122, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_123, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
fof(c_0_124, plain, ![X127, X128]:~v1_xboole_0(k2_tarski(X127,X128)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
cnf(c_0_125, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_126, plain, ![X129, X130, X131]:(X129=X130|(X129=X131|k4_xboole_0(k1_enumset1(X129,X130,X131),k1_tarski(X129))=k2_tarski(X130,X131))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_enumset1])])])).
cnf(c_0_127, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_128, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_68, c_0_115])).
cnf(c_0_129, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk2_0,esk1_0,X1,esk3_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_130, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_131, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_105, c_0_120])).
cnf(c_0_132, negated_conjecture, (k6_enumset1(esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_121, c_0_93])).
cnf(c_0_133, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_134, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_135, negated_conjecture, (r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_125, c_0_66])).
cnf(c_0_136, plain, (X1=X2|X1=X3|k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_137, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_67]), c_0_128])).
cnf(c_0_138, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk2_0,esk1_0,X1,esk3_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_139, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_118, c_0_131])).
cnf(c_0_140, negated_conjecture, (k6_enumset1(X1,esk2_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_132]), c_0_74]), c_0_76]), c_0_76])).
cnf(c_0_141, negated_conjecture, (v1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_133, c_0_66])).
cnf(c_0_142, plain, (~v1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_134, c_0_44]), c_0_45]), c_0_46]), c_0_47]), c_0_48]), c_0_49])).
cnf(c_0_143, negated_conjecture, (r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0)))), inference(rw,[status(thm)],[c_0_135, c_0_75])).
cnf(c_0_144, plain, (X1=X3|X1=X2|k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_136, c_0_43]), c_0_44]), c_0_44]), c_0_56]), c_0_45]), c_0_45]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49])).
cnf(c_0_145, plain, (k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))=k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X7,X8),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)))), inference(spm,[status(thm)],[c_0_137, c_0_83])).
cnf(c_0_146, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,X1,esk1_0,esk2_0,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_138]), c_0_139])).
cnf(c_0_147, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X5,X6)=k6_enumset1(X4,X4,X4,X3,X1,X2,X5,X6)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_103]), c_0_83])).
cnf(c_0_148, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_95]), c_0_132])).
cnf(c_0_149, negated_conjecture, (~r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_141, c_0_75]), c_0_142])).
cnf(c_0_150, negated_conjecture, (r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)))), inference(rw,[status(thm)],[c_0_143, c_0_91])).
cnf(c_0_151, plain, (k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3))=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)|X1=X2|X1=X3), inference(rw,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_152, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_146]), c_0_147]), c_0_147])).
cnf(c_0_153, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_148]), c_0_74]), c_0_76]), c_0_76])).
cnf(c_0_154, negated_conjecture, (~r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[c_0_149, c_0_91])).
cnf(c_0_155, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_156, negated_conjecture, (esk1_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_157, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_80]), c_0_93]), c_0_98]), c_0_109]), c_0_152]), c_0_97]), c_0_153]), c_0_154]), c_0_155]), c_0_156]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 158
# Proof object clause steps            : 99
# Proof object formula steps           : 59
# Proof object conjectures             : 49
# Proof object clause conjectures      : 46
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 29
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 237
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 24
# Processed clauses                    : 1480
# ...of these trivial                  : 129
# ...subsumed                          : 927
# ...remaining for further processing  : 424
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 179
# Generated clauses                    : 10864
# ...of the previous two non-trivial   : 9264
# Contextual simplify-reflections      : 0
# Paramodulations                      : 10864
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 220
#    Positive orientable unit clauses  : 128
#    Positive unorientable unit clauses: 31
#    Negative unit clauses             : 53
#    Non-unit-clauses                  : 8
# Current number of unprocessed clauses: 7341
# ...number of literals in the above   : 7766
# Current number of archived formulas  : 0
# Current number of archived clauses   : 213
# Clause-clause subsumption calls (NU) : 3
# Rec. Clause-clause subsumption calls : 1
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 2644
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9790
# BW rewrite match successes           : 1720
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 183036

# -------------------------------------------------
# User time                : 0.092 s
# System time              : 0.000 s
# Total time               : 0.092 s
# Maximum resident set size: 28688 pages
