# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', d10_xboole_0)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t70_enumset1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t12_setfam_1)).
fof(t125_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k2_relat_1(X2),X1)=>k8_relat_1(X1,X2)=X2)), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t125_relat_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t7_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t74_enumset1)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', d6_relat_1)).
fof(t26_wellord1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k2_wellord1(k2_wellord1(X3,X1),X2)=k2_wellord1(X3,k3_xboole_0(X1,X2))), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t26_wellord1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t71_relat_1)).
fof(dt_k6_relat_1, axiom, ![X1]:v1_relat_1(k6_relat_1(X1)), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', dt_k6_relat_1)).
fof(t18_wellord1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k2_wellord1(X2,X1)=k8_relat_1(X1,k7_relat_1(X2,X1))), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t18_wellord1)).
fof(t97_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k1_relat_1(X2),X1)=>k7_relat_1(X2,X1)=X2)), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t97_relat_1)).
fof(t43_funct_1, axiom, ![X1, X2]:k5_relat_1(k6_relat_1(X2),k6_relat_1(X1))=k6_relat_1(k3_xboole_0(X1,X2)), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t43_funct_1)).
fof(dt_k2_wellord1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k2_wellord1(X1,X2))), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', dt_k2_wellord1)).
fof(t20_wellord1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k3_relat_1(k2_wellord1(X2,X1)),k3_relat_1(X2))&r1_tarski(k3_relat_1(k2_wellord1(X2,X1)),X1))), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t20_wellord1)).
fof(t27_wellord1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k2_wellord1(k2_wellord1(X3,X1),X2)=k2_wellord1(k2_wellord1(X3,X2),X1)), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t27_wellord1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t1_xboole_1)).
fof(t29_wellord1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>k2_wellord1(k2_wellord1(X3,X2),X1)=k2_wellord1(X3,X1))), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t29_wellord1)).
fof(t94_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k7_relat_1(X2,X1)=k5_relat_1(k6_relat_1(X1),X2)), file('data/MPTP/proofs/810fef96bf7145ec92d79979d27bcc4d.E_input', t94_relat_1)).
fof(c_0_23, plain, ![X7, X8]:(((r1_tarski(X7,X8)|X7!=X8)&(r1_tarski(X8,X7)|X7!=X8))&(~r1_tarski(X7,X8)|~r1_tarski(X8,X7)|X7=X8)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_24, plain, ![X9, X10]:k3_tarski(k2_tarski(X9,X10))=k2_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_25, plain, ![X11, X12]:k1_enumset1(X11,X11,X12)=k2_tarski(X11,X12), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_26, plain, ![X13, X14]:k1_setfam_1(k2_tarski(X13,X14))=k3_xboole_0(X13,X14), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_27, plain, ![X15, X16]:(~v1_relat_1(X16)|(~r1_tarski(k2_relat_1(X16),X15)|k8_relat_1(X15,X16)=X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t125_relat_1])])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_29, plain, ![X17, X18]:r1_tarski(X17,k2_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_30, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X19, X20, X21]:k2_enumset1(X19,X19,X20,X21)=k1_enumset1(X19,X20,X21), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_33, plain, ![X22, X23, X24, X25]:k3_enumset1(X22,X22,X23,X24,X25)=k2_enumset1(X22,X23,X24,X25), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_34, plain, ![X26, X27, X28, X29, X30]:k4_enumset1(X26,X26,X27,X28,X29,X30)=k3_enumset1(X26,X27,X28,X29,X30), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_35, plain, ![X31, X32, X33, X34, X35, X36]:k5_enumset1(X31,X31,X32,X33,X34,X35,X36)=k4_enumset1(X31,X32,X33,X34,X35,X36), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_36, plain, ![X37]:(~v1_relat_1(X37)|k3_relat_1(X37)=k2_xboole_0(k1_relat_1(X37),k2_relat_1(X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
fof(c_0_37, plain, ![X38, X39, X40]:(~v1_relat_1(X40)|k2_wellord1(k2_wellord1(X40,X38),X39)=k2_wellord1(X40,k3_xboole_0(X38,X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_wellord1])])).
cnf(c_0_38, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, plain, (k8_relat_1(X2,X1)=X1|~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_28])).
fof(c_0_41, plain, ![X41]:(k1_relat_1(k6_relat_1(X41))=X41&k2_relat_1(k6_relat_1(X41))=X41), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_42, plain, ![X42]:v1_relat_1(k6_relat_1(X42)), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
fof(c_0_43, plain, ![X43, X44]:(~v1_relat_1(X44)|k2_wellord1(X44,X43)=k8_relat_1(X43,k7_relat_1(X44,X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_wellord1])])).
fof(c_0_44, plain, ![X45, X46]:(~v1_relat_1(X46)|(~r1_tarski(k1_relat_1(X46),X45)|k7_relat_1(X46,X45)=X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t97_relat_1])])).
fof(c_0_45, plain, ![X47, X48]:k5_relat_1(k6_relat_1(X48),k6_relat_1(X47))=k6_relat_1(k3_xboole_0(X47,X48)), inference(variable_rename,[status(thm)],[t43_funct_1])).
cnf(c_0_46, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_47, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_48, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_49, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_50, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_51, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_52, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_53, plain, ![X49, X50]:(~v1_relat_1(X49)|v1_relat_1(k2_wellord1(X49,X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_wellord1])])).
cnf(c_0_54, plain, (k2_wellord1(k2_wellord1(X1,X2),X3)=k2_wellord1(X1,k3_xboole_0(X2,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_55, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_38, c_0_31])).
cnf(c_0_56, plain, (k8_relat_1(k2_relat_1(X1),X1)=X1|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_57, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_58, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_59, plain, (k2_wellord1(X1,X2)=k8_relat_1(X2,k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_60, plain, (k7_relat_1(X1,X2)=X1|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_61, plain, ![X51, X52]:((r1_tarski(k3_relat_1(k2_wellord1(X52,X51)),k3_relat_1(X52))|~v1_relat_1(X52))&(r1_tarski(k3_relat_1(k2_wellord1(X52,X51)),X51)|~v1_relat_1(X52))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_wellord1])])])).
fof(c_0_62, plain, ![X53, X54, X55]:(~v1_relat_1(X55)|k2_wellord1(k2_wellord1(X55,X53),X54)=k2_wellord1(k2_wellord1(X55,X54),X53)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t27_wellord1])])).
cnf(c_0_63, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k6_relat_1(k3_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_64, plain, ![X56, X57, X58]:(~r1_tarski(X56,X57)|~r1_tarski(X57,X58)|r1_tarski(X56,X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_65, plain, (r1_tarski(X1,k3_tarski(k5_enumset1(X1,X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_66, plain, (k3_relat_1(X1)=k3_tarski(k5_enumset1(k1_relat_1(X1),k1_relat_1(X1),k1_relat_1(X1),k1_relat_1(X1),k1_relat_1(X1),k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_67, plain, (v1_relat_1(k2_wellord1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_68, plain, (k2_wellord1(k2_wellord1(X1,X2),X3)=k2_wellord1(X1,k1_setfam_1(k5_enumset1(X2,X2,X2,X2,X2,X2,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55]), c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_69, plain, (k8_relat_1(X1,k6_relat_1(X1))=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])])).
cnf(c_0_70, plain, (k8_relat_1(X1,X2)=k2_wellord1(X2,X1)|~v1_relat_1(X2)|~r1_tarski(k1_relat_1(X2),X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_71, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_72, plain, (r1_tarski(k3_relat_1(k2_wellord1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_73, plain, (k2_wellord1(k2_wellord1(X1,X2),X3)=k2_wellord1(k2_wellord1(X1,X3),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_74, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k6_relat_1(k1_setfam_1(k5_enumset1(X2,X2,X2,X2,X2,X2,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_55]), c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_75, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_76, plain, (r1_tarski(k1_relat_1(X1),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_77, plain, (v1_relat_1(k2_wellord1(k2_wellord1(X1,X2),X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_78, plain, (k2_wellord1(k6_relat_1(X1),X1)=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_58]), c_0_71]), c_0_40])])).
cnf(c_0_79, plain, (r1_tarski(k3_relat_1(k2_wellord1(k2_wellord1(X1,X2),X3)),X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_67])).
cnf(c_0_80, plain, (v1_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)))), inference(spm,[status(thm)],[c_0_58, c_0_74])).
cnf(c_0_81, plain, (r1_tarski(X1,X2)|~v1_relat_1(X3)|~r1_tarski(X1,k3_relat_1(k2_wellord1(X3,X2)))), inference(spm,[status(thm)],[c_0_75, c_0_72])).
cnf(c_0_82, plain, (r1_tarski(X1,k3_relat_1(X2))|~v1_relat_1(X2)|~r1_tarski(X1,k1_relat_1(X2))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_83, plain, (v1_relat_1(k2_wellord1(k6_relat_1(X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_58])])).
cnf(c_0_84, plain, (r1_tarski(X1,X2)|~v1_relat_1(X3)|~r1_tarski(X1,k3_relat_1(k2_wellord1(k2_wellord1(X3,X2),X4)))), inference(spm,[status(thm)],[c_0_75, c_0_79])).
cnf(c_0_85, plain, (k2_wellord1(k2_wellord1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)),X2),X1)=k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_78]), c_0_74]), c_0_74]), c_0_74]), c_0_80])])).
cnf(c_0_86, plain, (r1_tarski(X1,X2)|~v1_relat_1(X3)|~r1_tarski(X1,k1_relat_1(k2_wellord1(X3,X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_67])).
cnf(c_0_87, plain, (v1_relat_1(k2_wellord1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)),X3))), inference(spm,[status(thm)],[c_0_83, c_0_74])).
cnf(c_0_88, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_relat_1(k5_relat_1(k6_relat_1(X3),k6_relat_1(X2))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_80])])).
cnf(c_0_89, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k5_relat_1(k6_relat_1(X2),k6_relat_1(X3))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_85]), c_0_87])])).
cnf(c_0_90, plain, (k1_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)))=k1_setfam_1(k5_enumset1(X2,X2,X2,X2,X2,X2,X1))), inference(spm,[status(thm)],[c_0_71, c_0_74])).
fof(c_0_91, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>k2_wellord1(k2_wellord1(X3,X2),X1)=k2_wellord1(X3,X1)))), inference(assume_negation,[status(cth)],[t29_wellord1])).
cnf(c_0_92, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k5_relat_1(k6_relat_1(X3),k6_relat_1(X2))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_82]), c_0_80])])).
cnf(c_0_93, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_setfam_1(k5_enumset1(X3,X3,X3,X3,X3,X3,X2)))), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_94, plain, (k2_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)))=k1_setfam_1(k5_enumset1(X2,X2,X2,X2,X2,X2,X1))), inference(spm,[status(thm)],[c_0_57, c_0_74])).
fof(c_0_95, plain, ![X62, X63]:(~v1_relat_1(X63)|k7_relat_1(X63,X62)=k5_relat_1(k6_relat_1(X62),X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_relat_1])])).
fof(c_0_96, negated_conjecture, (v1_relat_1(esk3_0)&(r1_tarski(esk1_0,esk2_0)&k2_wellord1(k2_wellord1(esk3_0,esk2_0),esk1_0)!=k2_wellord1(esk3_0,esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_91])])])).
cnf(c_0_97, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_setfam_1(k5_enumset1(X2,X2,X2,X2,X2,X2,X3)))), inference(spm,[status(thm)],[c_0_92, c_0_90])).
cnf(c_0_98, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_relat_1(k5_relat_1(k6_relat_1(X2),k6_relat_1(X3))))), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_99, plain, (k7_relat_1(X1,X2)=k5_relat_1(k6_relat_1(X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_100, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_relat_1(k5_relat_1(k6_relat_1(X2),k6_relat_1(X3))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_85]), c_0_87])])).
cnf(c_0_101, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_102, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_relat_1(k5_relat_1(k6_relat_1(X3),k6_relat_1(X2))))), inference(spm,[status(thm)],[c_0_97, c_0_94])).
cnf(c_0_103, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_relat_1(k7_relat_1(k6_relat_1(X3),X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_58])])).
cnf(c_0_104, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_relat_1(k7_relat_1(k6_relat_1(X3),X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_99]), c_0_58])])).
cnf(c_0_105, plain, (v1_relat_1(k7_relat_1(k6_relat_1(X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_99]), c_0_58])])).
cnf(c_0_106, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_75, c_0_101])).
cnf(c_0_107, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_relat_1(k7_relat_1(k6_relat_1(X2),X3)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_99]), c_0_58])])).
cnf(c_0_108, plain, (r1_tarski(k2_relat_1(k7_relat_1(k6_relat_1(X1),X2)),X2)), inference(spm,[status(thm)],[c_0_103, c_0_40])).
cnf(c_0_109, plain, (k2_wellord1(k7_relat_1(X1,X2),X2)=k2_wellord1(X1,X2)|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_59, c_0_70])).
cnf(c_0_110, plain, (r1_tarski(k1_relat_1(k7_relat_1(k6_relat_1(X1),X2)),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_76]), c_0_105])])).
cnf(c_0_111, negated_conjecture, (k8_relat_1(esk2_0,X1)=X1|~v1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_106])).
cnf(c_0_112, plain, (r1_tarski(k2_relat_1(k7_relat_1(k6_relat_1(X1),X2)),X1)), inference(spm,[status(thm)],[c_0_107, c_0_40])).
cnf(c_0_113, plain, (k8_relat_1(X1,k7_relat_1(k6_relat_1(X2),X1))=k7_relat_1(k6_relat_1(X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_108]), c_0_105])])).
cnf(c_0_114, plain, (k2_wellord1(k7_relat_1(k6_relat_1(X1),X2),X2)=k2_wellord1(k6_relat_1(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_105]), c_0_58])])).
cnf(c_0_115, negated_conjecture, (k8_relat_1(esk2_0,k7_relat_1(k6_relat_1(esk1_0),X1))=k7_relat_1(k6_relat_1(esk1_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_105])])).
cnf(c_0_116, plain, (k7_relat_1(k6_relat_1(X1),X2)=k2_wellord1(k6_relat_1(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_113]), c_0_114]), c_0_105]), c_0_110])])).
cnf(c_0_117, negated_conjecture, (k8_relat_1(esk2_0,k6_relat_1(esk1_0))=k6_relat_1(esk1_0)|~r1_tarski(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_60]), c_0_58]), c_0_71])])).
cnf(c_0_118, plain, (k1_setfam_1(k5_enumset1(X1,X1,X1,X1,X1,X1,X2))=k1_relat_1(k2_wellord1(k6_relat_1(X1),X2))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_99]), c_0_58])]), c_0_116])).
cnf(c_0_119, negated_conjecture, (k8_relat_1(esk2_0,k6_relat_1(esk1_0))=k6_relat_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_106]), c_0_40])])).
cnf(c_0_120, negated_conjecture, (k2_wellord1(k2_wellord1(esk3_0,esk2_0),esk1_0)!=k2_wellord1(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_121, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_122, plain, (k2_wellord1(X1,k1_relat_1(k2_wellord1(k6_relat_1(X2),X3)))=k2_wellord1(k2_wellord1(X1,X2),X3)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_68, c_0_118])).
cnf(c_0_123, negated_conjecture, (k2_wellord1(k6_relat_1(esk1_0),esk2_0)=k6_relat_1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_119]), c_0_58]), c_0_71]), c_0_101])])).
cnf(c_0_124, negated_conjecture, (k2_wellord1(k2_wellord1(esk3_0,esk1_0),esk2_0)!=k2_wellord1(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_73]), c_0_121])])).
cnf(c_0_125, negated_conjecture, (k2_wellord1(k2_wellord1(X1,esk1_0),esk2_0)=k2_wellord1(X1,esk1_0)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_71])).
cnf(c_0_126, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_121])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 127
# Proof object clause steps            : 80
# Proof object formula steps           : 47
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 85
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 22
# Processed clauses                    : 6145
# ...of these trivial                  : 167
# ...subsumed                          : 4893
# ...remaining for further processing  : 1085
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 492
# Generated clauses                    : 103487
# ...of the previous two non-trivial   : 97915
# Contextual simplify-reflections      : 92
# Paramodulations                      : 103485
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 557
#    Positive orientable unit clauses  : 94
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 458
# Current number of unprocessed clauses: 91116
# ...number of literals in the above   : 243533
# Current number of archived formulas  : 0
# Current number of archived clauses   : 533
# Clause-clause subsumption calls (NU) : 144053
# Rec. Clause-clause subsumption calls : 113706
# Non-unit clause-clause subsumptions  : 4753
# Unit Clause-clause subsumption calls : 1791
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5343
# BW rewrite match successes           : 189
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1682922

# -------------------------------------------------
# User time                : 0.876 s
# System time              : 0.010 s
# Total time               : 0.885 s
# Maximum resident set size: 68660 pages
