# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t37_wellord1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', t37_wellord1)).
fof(l4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v6_relat_2(X1)<=>![X2, X3]:~((r2_hidden(X2,k3_relat_1(X1))&(r2_hidden(X3,k3_relat_1(X1))&(~(X2=X3)&(~(r2_hidden(k4_tarski(X2,X3),X1))&~(r2_hidden(k4_tarski(X3,X2),X1))))))))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', l4_wellord1)).
fof(d4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)<=>(v1_relat_2(X1)&(v8_relat_2(X1)&(v4_relat_2(X1)&(v6_relat_2(X1)&v1_wellord1(X1))))))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', d4_wellord1)).
fof(d1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k1_wellord1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(~(X4=X2)&r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', d1_wellord1)).
fof(l2_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v8_relat_2(X1)<=>![X2, X3, X4]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X3,X4),X1))=>r2_hidden(k4_tarski(X2,X4),X1)))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', l2_wellord1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', d3_tarski)).
fof(s1_xboole_0__e7_18__wellord1, axiom, ![X1, X2]:(v1_relat_1(X1)=>?[X3]:![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,k3_relat_1(X1))&~(r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', s1_xboole_0__e7_18__wellord1)).
fof(l3_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v4_relat_2(X1)<=>![X2, X3]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X3,X2),X1))=>X2=X3))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', l3_wellord1)).
fof(t30_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', t30_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', d10_xboole_0)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', t7_ordinal1)).
fof(t10_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)=>![X2]:~((r1_tarski(X2,k3_relat_1(X1))&(~(X2=k1_xboole_0)&![X3]:~((r2_hidden(X3,X2)&![X4]:(r2_hidden(X4,X2)=>r2_hidden(k4_tarski(X3,X4),X1))))))))), file('data/MPTP/proofs/814bdcc08d6141d68c3ab0a9ef5e2a53.E_input', t10_wellord1)).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2)))))), inference(assume_negation,[status(cth)],[t37_wellord1])).
fof(c_0_13, plain, ![X43, X44, X45]:((~v6_relat_2(X43)|(~r2_hidden(X44,k3_relat_1(X43))|(~r2_hidden(X45,k3_relat_1(X43))|(X44=X45|(r2_hidden(k4_tarski(X44,X45),X43)|r2_hidden(k4_tarski(X45,X44),X43)))))|~v1_relat_1(X43))&((r2_hidden(esk12_1(X43),k3_relat_1(X43))|v6_relat_2(X43)|~v1_relat_1(X43))&((r2_hidden(esk13_1(X43),k3_relat_1(X43))|v6_relat_2(X43)|~v1_relat_1(X43))&((esk12_1(X43)!=esk13_1(X43)|v6_relat_2(X43)|~v1_relat_1(X43))&((~r2_hidden(k4_tarski(esk12_1(X43),esk13_1(X43)),X43)|v6_relat_2(X43)|~v1_relat_1(X43))&(~r2_hidden(k4_tarski(esk13_1(X43),esk12_1(X43)),X43)|v6_relat_2(X43)|~v1_relat_1(X43))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l4_wellord1])])])])])])])).
fof(c_0_14, negated_conjecture, (v1_relat_1(esk3_0)&((v2_wellord1(esk3_0)&(r2_hidden(esk1_0,k3_relat_1(esk3_0))&r2_hidden(esk2_0,k3_relat_1(esk3_0))))&((~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0)))&(r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_15, plain, (X2=X3|r2_hidden(k4_tarski(X2,X3),X1)|r2_hidden(k4_tarski(X3,X2),X1)|~v6_relat_2(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (r2_hidden(esk2_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X27]:(((v1_relat_2(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))&((v8_relat_2(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))&((v4_relat_2(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))&((v6_relat_2(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))&(v1_wellord1(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))))))&(~v1_relat_2(X27)|(~v8_relat_2(X27)|(~v4_relat_2(X27)|(~v6_relat_2(X27)|~v1_wellord1(X27))))|v2_wellord1(X27)|~v1_relat_1(X27))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_wellord1])])])).
fof(c_0_19, plain, ![X17, X18, X19, X20, X21, X22, X23]:((((X20!=X18|~r2_hidden(X20,X19)|X19!=k1_wellord1(X17,X18)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(X20,X18),X17)|~r2_hidden(X20,X19)|X19!=k1_wellord1(X17,X18)|~v1_relat_1(X17)))&(X21=X18|~r2_hidden(k4_tarski(X21,X18),X17)|r2_hidden(X21,X19)|X19!=k1_wellord1(X17,X18)|~v1_relat_1(X17)))&((~r2_hidden(esk6_3(X17,X22,X23),X23)|(esk6_3(X17,X22,X23)=X22|~r2_hidden(k4_tarski(esk6_3(X17,X22,X23),X22),X17))|X23=k1_wellord1(X17,X22)|~v1_relat_1(X17))&((esk6_3(X17,X22,X23)!=X22|r2_hidden(esk6_3(X17,X22,X23),X23)|X23=k1_wellord1(X17,X22)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(esk6_3(X17,X22,X23),X22),X17)|r2_hidden(esk6_3(X17,X22,X23),X23)|X23=k1_wellord1(X17,X22)|~v1_relat_1(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_wellord1])])])])])])])).
cnf(c_0_20, negated_conjecture, (X1=esk2_0|r2_hidden(k4_tarski(esk2_0,X1),esk3_0)|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~v6_relat_2(esk3_0)|~r2_hidden(X1,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_21, plain, (v6_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (v2_wellord1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (X1=X2|r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (X1=esk2_0|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|r2_hidden(k4_tarski(esk2_0,X1),esk3_0)|~r2_hidden(X1,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_17])])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk1_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_26, plain, ![X36, X37, X38, X39]:((~v8_relat_2(X36)|(~r2_hidden(k4_tarski(X37,X38),X36)|~r2_hidden(k4_tarski(X38,X39),X36)|r2_hidden(k4_tarski(X37,X39),X36))|~v1_relat_1(X36))&(((r2_hidden(k4_tarski(esk9_1(X36),esk10_1(X36)),X36)|v8_relat_2(X36)|~v1_relat_1(X36))&(r2_hidden(k4_tarski(esk10_1(X36),esk11_1(X36)),X36)|v8_relat_2(X36)|~v1_relat_1(X36)))&(~r2_hidden(k4_tarski(esk9_1(X36),esk11_1(X36)),X36)|v8_relat_2(X36)|~v1_relat_1(X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l2_wellord1])])])])])).
cnf(c_0_27, plain, (X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk2_0,esk1_0),esk3_0)|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_29, plain, ![X50, X51, X52, X53, X54]:((~r1_tarski(X50,X51)|(~r2_hidden(X52,X50)|r2_hidden(X52,X51)))&((r2_hidden(esk14_2(X53,X54),X53)|r1_tarski(X53,X54))&(~r2_hidden(esk14_2(X53,X54),X54)|r1_tarski(X53,X54)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_30, plain, ![X12, X13, X15, X16]:(((r2_hidden(X15,k3_relat_1(X12))|~r2_hidden(X15,esk5_2(X12,X13))|~v1_relat_1(X12))&(~r2_hidden(k4_tarski(X15,X13),X12)|~r2_hidden(X15,esk5_2(X12,X13))|~v1_relat_1(X12)))&(~r2_hidden(X16,k3_relat_1(X12))|r2_hidden(k4_tarski(X16,X13),X12)|r2_hidden(X16,esk5_2(X12,X13))|~v1_relat_1(X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s1_xboole_0__e7_18__wellord1])])])])])])])).
cnf(c_0_31, plain, (r2_hidden(k4_tarski(X2,X4),X1)|~v8_relat_2(X1)|~r2_hidden(k4_tarski(X2,X3),X1)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_17])])).
cnf(c_0_33, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35, plain, (~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,esk5_2(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~v8_relat_2(esk3_0)|~r2_hidden(k4_tarski(X1,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_17])])).
cnf(c_0_37, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(X1,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|~r2_hidden(esk1_0,esk5_2(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_32]), c_0_17])])).
cnf(c_0_40, negated_conjecture, (esk2_0=esk1_0|X1=esk2_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~v8_relat_2(esk3_0)|~r2_hidden(k4_tarski(X1,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_36]), c_0_17])])).
cnf(c_0_41, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,k1_wellord1(X3,X2))|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_37])).
fof(c_0_42, plain, ![X31, X32, X33]:((~v4_relat_2(X31)|(~r2_hidden(k4_tarski(X32,X33),X31)|~r2_hidden(k4_tarski(X33,X32),X31)|X32=X33)|~v1_relat_1(X31))&(((r2_hidden(k4_tarski(esk7_1(X31),esk8_1(X31)),X31)|v4_relat_2(X31)|~v1_relat_1(X31))&(r2_hidden(k4_tarski(esk8_1(X31),esk7_1(X31)),X31)|v4_relat_2(X31)|~v1_relat_1(X31)))&(esk7_1(X31)!=esk8_1(X31)|v4_relat_2(X31)|~v1_relat_1(X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_wellord1])])])])])).
fof(c_0_43, plain, ![X28, X29, X30]:((r2_hidden(X28,k3_relat_1(X30))|~r2_hidden(k4_tarski(X28,X29),X30)|~v1_relat_1(X30))&(r2_hidden(X29,k3_relat_1(X30))|~r2_hidden(k4_tarski(X28,X29),X30)|~v1_relat_1(X30))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_relat_1])])])).
cnf(c_0_44, plain, (X1!=X2|~r2_hidden(X1,X3)|X3!=k1_wellord1(X4,X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_45, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk2_0))|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk1_0,esk5_2(esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, negated_conjecture, (esk2_0=esk1_0|X1=esk2_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~v8_relat_2(esk3_0)|~r2_hidden(X1,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_17])])).
cnf(c_0_47, plain, (v8_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_48, plain, (r2_hidden(k4_tarski(esk6_3(X1,X2,X3),X2),X1)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k1_wellord1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_49, plain, (X2=X3|~v4_relat_2(X1)|~r2_hidden(k4_tarski(X2,X3),X1)|~r2_hidden(k4_tarski(X3,X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, plain, (r2_hidden(k4_tarski(X1,X3),X2)|r2_hidden(X1,esk5_2(X2,X3))|~r2_hidden(X1,k3_relat_1(X2))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_51, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (~r2_hidden(X1,k1_wellord1(X2,X1))|~v1_relat_1(X2)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_44])])).
cnf(c_0_53, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(esk1_0,esk5_2(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_45]), c_0_17])])).
cnf(c_0_54, negated_conjecture, (esk2_0=esk1_0|X1=esk2_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))|r2_hidden(X1,k1_wellord1(esk3_0,esk2_0))|~r2_hidden(X1,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_22]), c_0_17])])).
cnf(c_0_55, plain, (r2_hidden(esk14_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_56, plain, ![X25, X26]:(((r1_tarski(X25,X26)|X25!=X26)&(r1_tarski(X26,X25)|X25!=X26))&(~r1_tarski(X25,X26)|~r1_tarski(X26,X25)|X25=X26)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_57, negated_conjecture, (X1=k1_wellord1(esk3_0,X2)|r2_hidden(k4_tarski(esk6_3(esk3_0,X2,X1),X2),esk3_0)|r2_hidden(esk6_3(esk3_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_48, c_0_17])).
cnf(c_0_58, plain, (X1=X2|r2_hidden(X2,esk5_2(X3,X1))|~v4_relat_2(X3)|~r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
cnf(c_0_59, negated_conjecture, (esk2_0=esk1_0|~r2_hidden(esk1_0,esk5_2(esk3_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_17])])).
cnf(c_0_60, plain, (r1_tarski(X1,X2)|~r2_hidden(esk14_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_61, negated_conjecture, (esk14_2(k1_wellord1(esk3_0,esk1_0),X1)=esk2_0|esk2_0=esk1_0|r1_tarski(k1_wellord1(esk3_0,esk1_0),X1)|r2_hidden(esk14_2(k1_wellord1(esk3_0,esk1_0),X1),k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
fof(c_0_62, plain, ![X48, X49]:(~r2_hidden(X48,X49)|~r1_tarski(X49,X48)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_63, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(X3,k1_wellord1(X2,X1))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_51, c_0_41])).
cnf(c_0_65, negated_conjecture, (X1=k1_wellord1(esk3_0,X2)|r2_hidden(esk6_3(esk3_0,X2,X1),X1)|r2_hidden(X2,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_57]), c_0_17])])).
cnf(c_0_66, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~v4_relat_2(esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_28]), c_0_17])]), c_0_59])).
cnf(c_0_67, negated_conjecture, (esk14_2(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))=esk2_0|esk2_0=esk1_0|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_68, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_70, negated_conjecture, (k1_wellord1(X1,X2)=k1_wellord1(esk3_0,X3)|r2_hidden(X3,k3_relat_1(esk3_0))|r2_hidden(X2,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_71, negated_conjecture, (esk2_0=esk1_0|~v4_relat_2(esk3_0)|~r2_hidden(k4_tarski(esk2_0,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_66]), c_0_17])])).
cnf(c_0_72, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_73, negated_conjecture, (esk2_0=esk1_0|r1_tarski(k1_wellord1(esk3_0,esk1_0),k1_wellord1(esk3_0,esk2_0))|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_55, c_0_67])).
cnf(c_0_74, plain, (~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, negated_conjecture, (k1_wellord1(esk3_0,X1)=k1_wellord1(esk3_0,X2)|r2_hidden(X1,k3_relat_1(esk3_0))|r2_hidden(X2,k3_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_70, c_0_17])).
cnf(c_0_76, negated_conjecture, (esk2_0=esk1_0|~v4_relat_2(esk3_0)|~r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_41]), c_0_17])])).
cnf(c_0_77, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_32])).
cnf(c_0_78, negated_conjecture, (k1_wellord1(esk3_0,X1)=k1_wellord1(esk3_0,k3_relat_1(esk3_0))|r2_hidden(X1,k3_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
fof(c_0_79, plain, ![X8, X9, X11]:((r2_hidden(esk4_2(X8,X9),X9)|X9=k1_xboole_0|~r1_tarski(X9,k3_relat_1(X8))|~v2_wellord1(X8)|~v1_relat_1(X8))&(~r2_hidden(X11,X9)|r2_hidden(k4_tarski(esk4_2(X8,X9),X11),X8)|X9=k1_xboole_0|~r1_tarski(X9,k3_relat_1(X8))|~v2_wellord1(X8)|~v1_relat_1(X8))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_wellord1])])])])])])).
cnf(c_0_80, plain, (r1_tarski(k1_wellord1(X1,X2),X3)|r2_hidden(X2,k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_64, c_0_55])).
cnf(c_0_81, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(X1,esk5_2(X2,X3))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_82, negated_conjecture, (esk2_0=esk1_0|~v4_relat_2(esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_83, plain, (v4_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_84, negated_conjecture, (r2_hidden(X1,k3_relat_1(esk3_0))|~r2_hidden(X2,k1_wellord1(esk3_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_78]), c_0_17])]), c_0_74])).
cnf(c_0_85, plain, (r2_hidden(esk4_2(X1,X2),X2)|X2=k1_xboole_0|~r1_tarski(X2,k3_relat_1(X1))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_86, negated_conjecture, (r1_tarski(k1_wellord1(esk3_0,X1),X2)|r2_hidden(X1,k3_relat_1(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_78]), c_0_17])]), c_0_74])).
cnf(c_0_87, plain, (r2_hidden(k4_tarski(esk4_2(X3,X2),X1),X3)|X2=k1_xboole_0|~r2_hidden(X1,X2)|~r1_tarski(X2,k3_relat_1(X3))|~v2_wellord1(X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_88, plain, (r1_tarski(esk5_2(X1,X2),X3)|r2_hidden(esk14_2(esk5_2(X1,X2),X3),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_81, c_0_55])).
cnf(c_0_89, negated_conjecture, (esk2_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_22]), c_0_17])])).
cnf(c_0_90, negated_conjecture, (k1_wellord1(esk3_0,X1)=k1_xboole_0|r2_hidden(X1,k3_relat_1(esk3_0))|~v2_wellord1(X2)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])).
cnf(c_0_91, plain, (X1=k1_xboole_0|~r1_tarski(X1,k3_relat_1(X2))|~r2_hidden(esk4_2(X2,X1),esk5_2(X2,X3))|~r2_hidden(X3,X1)|~v2_wellord1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_35, c_0_87])).
cnf(c_0_92, plain, (r1_tarski(esk5_2(X1,X2),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_60, c_0_88])).
cnf(c_0_93, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk1_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_89]), c_0_69]), c_0_89])])).
cnf(c_0_94, negated_conjecture, (k1_wellord1(esk3_0,X1)=k1_xboole_0|r2_hidden(X1,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_22]), c_0_17])])).
cnf(c_0_95, plain, (esk5_2(X1,X2)=k1_xboole_0|~r2_hidden(X2,esk5_2(X1,X2))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_85]), c_0_92])).
cnf(c_0_96, negated_conjecture, (r2_hidden(esk1_0,esk5_2(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_50]), c_0_25]), c_0_17])])).
cnf(c_0_97, negated_conjecture, (k1_wellord1(esk3_0,k3_relat_1(esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_74, c_0_94])).
cnf(c_0_98, negated_conjecture, (esk5_2(esk3_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_22]), c_0_17])])).
cnf(c_0_99, negated_conjecture, (~r2_hidden(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_97]), c_0_17])]), c_0_74])).
cnf(c_0_100, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_98]), c_0_99]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 101
# Proof object clause steps            : 76
# Proof object formula steps           : 25
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 43
# Proof object simplifying inferences  : 66
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 46
# Processed clauses                    : 2751
# ...of these trivial                  : 13
# ...subsumed                          : 1664
# ...remaining for further processing  : 1074
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 320
# Backward-rewritten                   : 417
# Generated clauses                    : 13871
# ...of the previous two non-trivial   : 11818
# Contextual simplify-reflections      : 44
# Paramodulations                      : 13860
# Factorizations                       : 6
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 332
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 318
# Current number of unprocessed clauses: 8815
# ...number of literals in the above   : 61255
# Current number of archived formulas  : 0
# Current number of archived clauses   : 737
# Clause-clause subsumption calls (NU) : 132043
# Rec. Clause-clause subsumption calls : 11355
# Non-unit clause-clause subsumptions  : 1880
# Unit Clause-clause subsumption calls : 2114
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 348164

# -------------------------------------------------
# User time                : 0.296 s
# System time              : 0.000 s
# Total time               : 0.296 s
# Maximum resident set size: 28744 pages
