# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t96_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3))))), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', t96_tmap_1)).
fof(dt_k4_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_funct_1(k4_tmap_1(X1,X2))&(v1_funct_2(k4_tmap_1(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(k4_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', dt_k4_tmap_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', d1_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', redefinition_k1_relset_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', fc6_relat_1)).
fof(t8_grfunc_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(X1,X2)<=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&![X3]:(r2_hidden(X3,k1_relat_1(X1))=>k1_funct_1(X1,X3)=k1_funct_1(X2,X3)))))), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', t8_grfunc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', d10_xboole_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', t1_subset)).
fof(t55_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1))))), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', t55_pre_topc)).
fof(t95_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,u1_struct_0(X2))=>k1_funct_1(k4_tmap_1(X1,X2),X3)=X3)))), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', t95_tmap_1)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/824fed4954ea48199b4dc3749da85f7c.E_input', redefinition_r2_funct_2)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t96_tmap_1])).
fof(c_0_13, plain, ![X9, X10]:((v1_funct_1(k4_tmap_1(X9,X10))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_pre_topc(X10,X9)))&((v1_funct_2(k4_tmap_1(X9,X10),u1_struct_0(X10),u1_struct_0(X9))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_pre_topc(X10,X9)))&(m1_subset_1(k4_tmap_1(X9,X10),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X9))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_pre_topc(X10,X9))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_tmap_1])])])])).
fof(c_0_14, negated_conjecture, ![X8]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))))&((~m1_subset_1(X8,u1_struct_0(esk1_0))|(~r2_hidden(X8,u1_struct_0(esk2_0))|X8=k1_funct_1(esk3_0,X8)))&~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])).
fof(c_0_15, plain, ![X11, X12, X13]:((((~v1_funct_2(X13,X11,X12)|X11=k1_relset_1(X11,X12,X13)|X12=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X11!=k1_relset_1(X11,X12,X13)|v1_funct_2(X13,X11,X12)|X12=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))))&((~v1_funct_2(X13,X11,X12)|X11=k1_relset_1(X11,X12,X13)|X11!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X11!=k1_relset_1(X11,X12,X13)|v1_funct_2(X13,X11,X12)|X11!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))))&((~v1_funct_2(X13,X11,X12)|X13=k1_xboole_0|X11=k1_xboole_0|X12!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12))))&(X13!=k1_xboole_0|v1_funct_2(X13,X11,X12)|X11=k1_xboole_0|X12!=k1_xboole_0|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_16, plain, (m1_subset_1(k4_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (v1_funct_2(k4_tmap_1(X1,X2),u1_struct_0(X2),u1_struct_0(X1))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_22, plain, ![X14, X15, X16]:(~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))|k1_relset_1(X14,X15,X16)=k1_relat_1(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_23, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(k4_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_25, negated_conjecture, (v1_funct_2(k4_tmap_1(esk1_0,esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
fof(c_0_26, plain, ![X17, X18]:(~v1_relat_1(X17)|(~m1_subset_1(X18,k1_zfmisc_1(X17))|v1_relat_1(X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_27, plain, ![X19, X20]:v1_relat_1(k2_zfmisc_1(X19,X20)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_28, plain, ![X21, X22, X23]:(((r1_tarski(k1_relat_1(X21),k1_relat_1(X22))|~r1_tarski(X21,X22)|(~v1_relat_1(X22)|~v1_funct_1(X22))|(~v1_relat_1(X21)|~v1_funct_1(X21)))&(~r2_hidden(X23,k1_relat_1(X21))|k1_funct_1(X21,X23)=k1_funct_1(X22,X23)|~r1_tarski(X21,X22)|(~v1_relat_1(X22)|~v1_funct_1(X22))|(~v1_relat_1(X21)|~v1_funct_1(X21))))&((r2_hidden(esk4_2(X21,X22),k1_relat_1(X21))|~r1_tarski(k1_relat_1(X21),k1_relat_1(X22))|r1_tarski(X21,X22)|(~v1_relat_1(X22)|~v1_funct_1(X22))|(~v1_relat_1(X21)|~v1_funct_1(X21)))&(k1_funct_1(X21,esk4_2(X21,X22))!=k1_funct_1(X22,esk4_2(X21,X22))|~r1_tarski(k1_relat_1(X21),k1_relat_1(X22))|r1_tarski(X21,X22)|(~v1_relat_1(X22)|~v1_funct_1(X22))|(~v1_relat_1(X21)|~v1_funct_1(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_grfunc_1])])])])])).
cnf(c_0_29, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (k1_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])])).
cnf(c_0_31, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (v1_funct_1(k4_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_36, plain, ![X25, X26]:(((r1_tarski(X25,X26)|X25!=X26)&(r1_tarski(X26,X25)|X25!=X26))&(~r1_tarski(X25,X26)|~r1_tarski(X26,X25)|X25=X26)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_37, plain, (r2_hidden(esk4_2(X1,X2),k1_relat_1(X1))|r1_tarski(X1,X2)|~r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (k1_relat_1(k4_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_24])])).
cnf(c_0_39, negated_conjecture, (v1_relat_1(k4_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_24]), c_0_32])])).
cnf(c_0_40, negated_conjecture, (v1_funct_1(k4_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_41, negated_conjecture, (k1_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_34]), c_0_35])])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_43, plain, ![X27, X28]:(~r2_hidden(X27,X28)|m1_subset_1(X27,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_44, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(X1,k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_2(X1,k4_tmap_1(esk1_0,esk2_0)),k1_relat_1(X1))|~r1_tarski(k1_relat_1(X1),u1_struct_0(esk2_0))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40])])).
cnf(c_0_45, negated_conjecture, (k1_relat_1(esk3_0)=u1_struct_0(esk2_0)|u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_41]), c_0_34])])).
cnf(c_0_46, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_47, negated_conjecture, (v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_34]), c_0_32])])).
cnf(c_0_48, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_49, plain, ![X29, X30, X31]:(v2_struct_0(X29)|~l1_pre_topc(X29)|(v2_struct_0(X30)|~m1_pre_topc(X30,X29)|(~m1_subset_1(X31,u1_struct_0(X30))|m1_subset_1(X31,u1_struct_0(X29))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_pre_topc])])])])).
cnf(c_0_50, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_47]), c_0_48])])).
fof(c_0_52, plain, ![X32, X33, X34]:(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|(v2_struct_0(X33)|~m1_pre_topc(X33,X32)|(~m1_subset_1(X34,u1_struct_0(X32))|(~r2_hidden(X34,u1_struct_0(X33))|k1_funct_1(k4_tmap_1(X32,X33),X34)=X34)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t95_tmap_1])])])])).
cnf(c_0_53, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_55, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_56, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k1_funct_1(k4_tmap_1(X1,X2),X3)=X3|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_57, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
cnf(c_0_58, negated_conjecture, (X1=k1_funct_1(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_59, negated_conjecture, (k1_funct_1(k4_tmap_1(X1,esk2_0),esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0))|u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_51]), c_0_55]), c_0_57])).
cnf(c_0_60, negated_conjecture, (k1_funct_1(esk3_0,esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0))|u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_58, c_0_51])).
cnf(c_0_61, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_17]), c_0_18])]), c_0_20])).
cnf(c_0_62, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_63, plain, (r1_tarski(X1,X2)|k1_funct_1(X1,esk4_2(X1,X2))!=k1_funct_1(X2,esk4_2(X1,X2))|~r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_64, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0))|u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_65, negated_conjecture, (k1_funct_1(esk3_0,esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0))|u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_66, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(k1_relat_1(X1),u1_struct_0(esk2_0))|~r1_tarski(X1,esk3_0)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_45]), c_0_47]), c_0_48])])).
cnf(c_0_67, plain, (k1_funct_1(X2,X1)=k1_funct_1(X3,X1)|~r2_hidden(X1,k1_relat_1(X2))|~r1_tarski(X2,X3)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_68, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|~r1_tarski(k1_relat_1(esk3_0),k1_relat_1(k4_tmap_1(esk1_0,esk2_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_39]), c_0_47]), c_0_40]), c_0_48])]), c_0_65])).
cnf(c_0_69, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(k1_relat_1(esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_46]), c_0_47]), c_0_48])])).
cnf(c_0_70, negated_conjecture, (k1_funct_1(esk3_0,X1)=k1_funct_1(X2,X1)|u1_struct_0(esk1_0)=k1_xboole_0|~r1_tarski(esk3_0,X2)|~v1_relat_1(X2)|~r2_hidden(X1,u1_struct_0(esk2_0))|~v1_funct_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_45]), c_0_47]), c_0_48])])).
cnf(c_0_71, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_38]), c_0_69])).
cnf(c_0_72, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(X1,esk3_0)|r2_hidden(esk4_2(X1,esk3_0),k1_relat_1(X1))|~r1_tarski(k1_relat_1(X1),u1_struct_0(esk2_0))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_45]), c_0_47]), c_0_48])])).
cnf(c_0_73, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),X1)=k1_funct_1(esk3_0,X1)|u1_struct_0(esk1_0)=k1_xboole_0|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_39]), c_0_40])])).
cnf(c_0_74, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)|r2_hidden(esk4_2(k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_38]), c_0_46]), c_0_39]), c_0_40])])).
cnf(c_0_75, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_2(k4_tmap_1(esk1_0,esk2_0),esk3_0))=k1_funct_1(esk3_0,esk4_2(k4_tmap_1(esk1_0,esk2_0),esk3_0))|u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_76, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(u1_struct_0(esk2_0),k1_relat_1(X1))|~r1_tarski(esk3_0,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_45]), c_0_47]), c_0_48])])).
cnf(c_0_77, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_78, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)|~r1_tarski(k1_relat_1(k4_tmap_1(esk1_0,esk2_0)),k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_75]), c_0_47]), c_0_39]), c_0_48]), c_0_40])])).
cnf(c_0_79, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(u1_struct_0(esk2_0),k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_46]), c_0_47]), c_0_48])])).
cnf(c_0_80, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=esk3_0|u1_struct_0(esk1_0)=k1_xboole_0|~r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_77, c_0_71])).
cnf(c_0_81, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_38]), c_0_79])).
fof(c_0_82, plain, ![X35, X36, X37, X38]:((~r2_funct_2(X35,X36,X37,X38)|X37=X38|(~v1_funct_1(X37)|(~v1_funct_2(X37,X35,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36))))|(~v1_funct_1(X38)|(~v1_funct_2(X38,X35,X36)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))))))&(X37!=X38|r2_funct_2(X35,X36,X37,X38)|(~v1_funct_1(X37)|(~v1_funct_2(X37,X35,X36)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36))))|(~v1_funct_1(X38)|(~v1_funct_2(X38,X35,X36)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
cnf(c_0_83, negated_conjecture, (~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_84, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=esk3_0|u1_struct_0(esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_86, plain, (X2=k1_relset_1(X2,X3,X1)|~v1_funct_2(X1,X2,X3)|X2!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_87, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_88, plain, (r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(er,[status(thm)],[c_0_85])).
cnf(c_0_89, plain, (X1=k1_relat_1(X2)|X1!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v1_funct_2(X2,X1,X3)), inference(spm,[status(thm)],[c_0_29, c_0_86])).
cnf(c_0_90, plain, (X1=k1_xboole_0|X2=k1_xboole_0|~v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_91, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_34]), c_0_35]), c_0_48])])).
cnf(c_0_92, negated_conjecture, (k1_relat_1(esk3_0)=u1_struct_0(esk2_0)|u1_struct_0(esk2_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_34]), c_0_35])])).
cnf(c_0_93, plain, (X1=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_xboole_0)))|~v1_funct_2(X1,X2,k1_xboole_0)), inference(er,[status(thm)],[c_0_90])).
cnf(c_0_94, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk2_0),k1_xboole_0)), inference(rw,[status(thm)],[c_0_35, c_0_91])).
cnf(c_0_95, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),k1_xboole_0)))), inference(rw,[status(thm)],[c_0_34, c_0_91])).
cnf(c_0_96, negated_conjecture, (k1_relat_1(k4_tmap_1(esk1_0,esk2_0))=u1_struct_0(esk2_0)|u1_struct_0(esk2_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_24]), c_0_25])])).
cnf(c_0_97, negated_conjecture, (r1_tarski(k1_relat_1(X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)!=k1_xboole_0|~r1_tarski(X1,esk3_0)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_92]), c_0_47]), c_0_48])])).
cnf(c_0_98, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95])])).
cnf(c_0_99, negated_conjecture, (r1_tarski(X1,k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_2(X1,k4_tmap_1(esk1_0,esk2_0)),k1_relat_1(X1))|u1_struct_0(esk2_0)!=k1_xboole_0|~r1_tarski(k1_relat_1(X1),u1_struct_0(esk2_0))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_96]), c_0_39]), c_0_40])])).
cnf(c_0_100, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k1_relat_1(X1),k1_xboole_0)|~r1_tarski(X1,esk3_0)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_101, plain, (k1_relat_1(X1)=k1_relat_1(X2)|~r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X2,X1)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)), inference(spm,[status(thm)],[c_0_77, c_0_62])).
cnf(c_0_102, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(X1,k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_2(X1,k4_tmap_1(esk1_0,esk2_0)),k1_relat_1(X1))|~r1_tarski(k1_relat_1(X1),k1_xboole_0)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_99, c_0_98])).
cnf(c_0_103, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k1_relat_1(esk3_0),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_46]), c_0_47]), c_0_48])])).
cnf(c_0_104, negated_conjecture, (k1_relat_1(X1)=u1_struct_0(esk2_0)|u1_struct_0(esk2_0)!=k1_xboole_0|~r1_tarski(k1_relat_1(X1),u1_struct_0(esk2_0))|~r1_tarski(esk3_0,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_92]), c_0_47]), c_0_48])])).
cnf(c_0_105, negated_conjecture, (m1_subset_1(k4_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),k1_xboole_0)))), inference(rw,[status(thm)],[c_0_24, c_0_91])).
cnf(c_0_106, negated_conjecture, (v1_funct_2(k4_tmap_1(esk1_0,esk2_0),u1_struct_0(esk2_0),k1_xboole_0)), inference(rw,[status(thm)],[c_0_25, c_0_91])).
cnf(c_0_107, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_47]), c_0_48])])).
cnf(c_0_108, negated_conjecture, (k1_relat_1(X1)=u1_struct_0(esk2_0)|u1_struct_0(esk2_0)!=k1_xboole_0|~r1_tarski(esk3_0,X1)|~r1_tarski(X1,esk3_0)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_104, c_0_97])).
cnf(c_0_109, negated_conjecture, (r1_tarski(esk3_0,X1)|r2_hidden(esk4_2(esk3_0,X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)!=k1_xboole_0|~r1_tarski(u1_struct_0(esk2_0),k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_92]), c_0_47]), c_0_48])])).
cnf(c_0_110, negated_conjecture, (esk3_0=k1_xboole_0|m1_subset_1(k4_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_105, c_0_98])).
cnf(c_0_111, negated_conjecture, (esk3_0=k1_xboole_0|v1_funct_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_106, c_0_98])).
cnf(c_0_112, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_46]), c_0_47]), c_0_48])]), c_0_98])).
cnf(c_0_113, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,X1)|r2_hidden(esk4_2(esk3_0,X1),k1_xboole_0)|~r1_tarski(k1_xboole_0,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_109, c_0_98])).
cnf(c_0_114, negated_conjecture, (k1_relat_1(k4_tmap_1(esk1_0,esk2_0))=k1_xboole_0|esk3_0=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_110]), c_0_111])).
cnf(c_0_115, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_50, c_0_112])).
cnf(c_0_116, negated_conjecture, (k1_funct_1(k4_tmap_1(X1,esk2_0),X2)=X2|esk3_0=k1_xboole_0|v2_struct_0(X1)|~r2_hidden(X2,k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_98]), c_0_55])).
cnf(c_0_117, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_46]), c_0_39]), c_0_40])])).
cnf(c_0_118, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(X1))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_115]), c_0_55])).
cnf(c_0_119, negated_conjecture, (k1_funct_1(esk3_0,X1)=X1|~r2_hidden(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_xboole_0)), inference(rw,[status(thm)],[c_0_58, c_0_91])).
cnf(c_0_120, negated_conjecture, (k1_funct_1(k4_tmap_1(X1,esk2_0),esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0))|esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118])).
cnf(c_0_121, negated_conjecture, (k1_funct_1(esk3_0,X1)=X1|esk3_0=k1_xboole_0|~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_98]), c_0_50])).
cnf(c_0_122, negated_conjecture, (r1_tarski(u1_struct_0(esk2_0),k1_relat_1(X1))|u1_struct_0(esk2_0)!=k1_xboole_0|~r1_tarski(esk3_0,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_92]), c_0_47]), c_0_48])])).
cnf(c_0_123, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0))|esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_124, negated_conjecture, (k1_funct_1(esk3_0,esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(esk3_0,k4_tmap_1(esk1_0,esk2_0))|esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_121, c_0_117])).
cnf(c_0_125, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k1_xboole_0,k1_relat_1(X1))|~r1_tarski(esk3_0,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_122, c_0_98])).
cnf(c_0_126, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|~r1_tarski(k1_relat_1(esk3_0),k1_relat_1(k4_tmap_1(esk1_0,esk2_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_123]), c_0_39]), c_0_47]), c_0_40]), c_0_48])]), c_0_124])).
cnf(c_0_127, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k1_xboole_0,k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_46]), c_0_47]), c_0_48])])).
cnf(c_0_128, negated_conjecture, (k1_funct_1(esk3_0,X1)=k1_funct_1(X2,X1)|u1_struct_0(esk2_0)!=k1_xboole_0|~r1_tarski(esk3_0,X2)|~v1_relat_1(X2)|~r2_hidden(X1,u1_struct_0(esk2_0))|~v1_funct_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_92]), c_0_47]), c_0_48])])).
cnf(c_0_129, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))|~r1_tarski(k1_relat_1(esk3_0),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_96]), c_0_98])).
cnf(c_0_130, negated_conjecture, (r1_tarski(k4_tmap_1(esk1_0,esk2_0),X1)|r2_hidden(esk4_2(k4_tmap_1(esk1_0,esk2_0),X1),u1_struct_0(esk2_0))|u1_struct_0(esk2_0)!=k1_xboole_0|~r1_tarski(u1_struct_0(esk2_0),k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_96]), c_0_39]), c_0_40])])).
cnf(c_0_131, negated_conjecture, (k1_relat_1(esk3_0)=k1_xboole_0|esk3_0=k1_xboole_0|~r1_tarski(k1_relat_1(esk3_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_77, c_0_127])).
cnf(c_0_132, negated_conjecture, (k1_funct_1(esk3_0,X1)=k1_funct_1(X2,X1)|esk3_0=k1_xboole_0|~r1_tarski(esk3_0,X2)|~v1_relat_1(X2)|~r2_hidden(X1,k1_xboole_0)|~v1_funct_1(X2)), inference(spm,[status(thm)],[c_0_128, c_0_98])).
cnf(c_0_133, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk3_0,k4_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_108]), c_0_46]), c_0_46]), c_0_47]), c_0_48])]), c_0_98])).
cnf(c_0_134, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),X1)|r2_hidden(esk4_2(k4_tmap_1(esk1_0,esk2_0),X1),k1_xboole_0)|~r1_tarski(k1_xboole_0,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_130, c_0_98])).
cnf(c_0_135, negated_conjecture, (k1_relat_1(esk3_0)=k1_xboole_0|esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_131, c_0_103])).
cnf(c_0_136, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),X1)=k1_funct_1(esk3_0,X1)|esk3_0=k1_xboole_0|~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_39]), c_0_40])])).
cnf(c_0_137, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)|r2_hidden(esk4_2(k4_tmap_1(esk1_0,esk2_0),esk3_0),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_46]), c_0_47]), c_0_48])])).
cnf(c_0_138, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_2(k4_tmap_1(esk1_0,esk2_0),esk3_0))=k1_funct_1(esk3_0,esk4_2(k4_tmap_1(esk1_0,esk2_0),esk3_0))|esk3_0=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_139, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)|~r1_tarski(k1_relat_1(k4_tmap_1(esk1_0,esk2_0)),k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_138]), c_0_47]), c_0_39]), c_0_48]), c_0_40])])).
cnf(c_0_140, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=esk3_0|esk3_0=k1_xboole_0|~r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_77, c_0_133])).
cnf(c_0_141, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_114]), c_0_127])).
cnf(c_0_142, negated_conjecture, (~r2_funct_2(u1_struct_0(esk2_0),k1_xboole_0,k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(rw,[status(thm)],[c_0_83, c_0_91])).
cnf(c_0_143, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=esk3_0|esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_144, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=k1_xboole_0|u1_struct_0(esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_106]), c_0_105])])).
cnf(c_0_145, negated_conjecture, (esk3_0=k1_xboole_0|~r2_funct_2(u1_struct_0(esk2_0),k1_xboole_0,esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_146, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~r2_funct_2(u1_struct_0(esk2_0),k1_xboole_0,k1_xboole_0,esk3_0)), inference(spm,[status(thm)],[c_0_142, c_0_144])).
cnf(c_0_147, negated_conjecture, (esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_88]), c_0_95]), c_0_94]), c_0_48])])).
cnf(c_0_148, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~r2_funct_2(u1_struct_0(esk2_0),k1_xboole_0,k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_149, negated_conjecture, (v1_funct_2(k1_xboole_0,u1_struct_0(esk2_0),k1_xboole_0)), inference(rw,[status(thm)],[c_0_94, c_0_147])).
cnf(c_0_150, negated_conjecture, (v1_funct_1(k1_xboole_0)), inference(rw,[status(thm)],[c_0_48, c_0_147])).
cnf(c_0_151, negated_conjecture, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),k1_xboole_0)))), inference(rw,[status(thm)],[c_0_95, c_0_147])).
cnf(c_0_152, negated_conjecture, (k1_relat_1(k1_xboole_0)=u1_struct_0(esk2_0)|u1_struct_0(esk2_0)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_92, c_0_147])).
cnf(c_0_153, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_88]), c_0_149]), c_0_150])]), c_0_151])])).
cnf(c_0_154, negated_conjecture, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_153]), c_0_153])])).
cnf(c_0_155, negated_conjecture, (v1_relat_1(k1_xboole_0)), inference(rw,[status(thm)],[c_0_47, c_0_147])).
cnf(c_0_156, negated_conjecture, (r1_tarski(k1_xboole_0,X1)|r2_hidden(esk4_2(k1_xboole_0,X1),k1_xboole_0)|~r1_tarski(k1_xboole_0,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_154]), c_0_155]), c_0_150])])).
cnf(c_0_157, negated_conjecture, (k1_relat_1(k4_tmap_1(esk1_0,esk2_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_153]), c_0_153])])).
cnf(c_0_158, negated_conjecture, (k1_funct_1(k4_tmap_1(X1,esk2_0),X2)=X2|v2_struct_0(X1)|~r2_hidden(X2,k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_153]), c_0_55])).
cnf(c_0_159, negated_conjecture, (r1_tarski(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))|r2_hidden(esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0)),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_46]), c_0_39]), c_0_40])])).
cnf(c_0_160, negated_conjecture, (r1_tarski(k4_tmap_1(esk1_0,esk2_0),X1)|r2_hidden(esk4_2(k4_tmap_1(esk1_0,esk2_0),X1),k1_xboole_0)|~r1_tarski(k1_xboole_0,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_130, c_0_153]), c_0_153]), c_0_153])])).
cnf(c_0_161, negated_conjecture, (k1_funct_1(k4_tmap_1(X1,esk2_0),esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))|r1_tarski(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))|v2_struct_0(X1)|~m1_subset_1(esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0)),u1_struct_0(X1))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_162, negated_conjecture, (r1_tarski(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))|m1_subset_1(esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0)),k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_159])).
cnf(c_0_163, negated_conjecture, (k1_funct_1(k1_xboole_0,X1)=X1|~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_147]), c_0_153]), c_0_50])).
cnf(c_0_164, negated_conjecture, (r1_tarski(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)|r2_hidden(esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_154]), c_0_46]), c_0_155]), c_0_150])])).
cnf(c_0_165, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))|r1_tarski(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_17]), c_0_91]), c_0_18]), c_0_19])]), c_0_20]), c_0_162])).
cnf(c_0_166, negated_conjecture, (k1_funct_1(k1_xboole_0,esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0)))=esk4_2(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))|r1_tarski(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_163, c_0_159])).
cnf(c_0_167, negated_conjecture, (k1_funct_1(k4_tmap_1(X1,esk2_0),esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0))=esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)|r1_tarski(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)|v2_struct_0(X1)|~m1_subset_1(esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0),u1_struct_0(X1))|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_158, c_0_164])).
cnf(c_0_168, negated_conjecture, (r1_tarski(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)|m1_subset_1(esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_164])).
cnf(c_0_169, negated_conjecture, (r1_tarski(k1_xboole_0,k4_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_165]), c_0_154]), c_0_157]), c_0_46]), c_0_39]), c_0_155]), c_0_40]), c_0_150])]), c_0_166])).
cnf(c_0_170, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0))=esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)|r1_tarski(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_17]), c_0_91]), c_0_18]), c_0_19])]), c_0_20]), c_0_168])).
cnf(c_0_171, negated_conjecture, (k1_funct_1(k1_xboole_0,esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0))=esk4_2(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)|r1_tarski(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_172, negated_conjecture, (~r2_funct_2(u1_struct_0(esk2_0),k1_xboole_0,k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)), inference(rw,[status(thm)],[c_0_142, c_0_147])).
cnf(c_0_173, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=k1_xboole_0|~r1_tarski(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_77, c_0_169])).
cnf(c_0_174, negated_conjecture, (r1_tarski(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_170]), c_0_157]), c_0_154]), c_0_46]), c_0_155]), c_0_39]), c_0_150]), c_0_40])]), c_0_171])).
cnf(c_0_175, negated_conjecture, (~r2_funct_2(k1_xboole_0,k1_xboole_0,k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)), inference(rw,[status(thm)],[c_0_172, c_0_153])).
cnf(c_0_176, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_173, c_0_174])])).
cnf(c_0_177, negated_conjecture, (~r2_funct_2(k1_xboole_0,k1_xboole_0,k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[c_0_175, c_0_176])).
cnf(c_0_178, negated_conjecture, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,k1_xboole_0)))), inference(rw,[status(thm)],[c_0_151, c_0_153])).
cnf(c_0_179, negated_conjecture, (v1_funct_2(k1_xboole_0,k1_xboole_0,k1_xboole_0)), inference(rw,[status(thm)],[c_0_149, c_0_153])).
cnf(c_0_180, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_88]), c_0_178]), c_0_179]), c_0_150])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 181
# Proof object clause steps            : 156
# Proof object formula steps           : 25
# Proof object conjectures             : 135
# Proof object clause conjectures      : 132
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 103
# Proof object simplifying inferences  : 254
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 34
# Processed clauses                    : 724
# ...of these trivial                  : 20
# ...subsumed                          : 303
# ...remaining for further processing  : 401
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 57
# Backward-rewritten                   : 233
# Generated clauses                    : 1079
# ...of the previous two non-trivial   : 939
# Contextual simplify-reflections      : 46
# Paramodulations                      : 1063
# Factorizations                       : 0
# Equation resolutions                 : 16
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 75
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 58
# Current number of unprocessed clauses: 148
# ...number of literals in the above   : 1485
# Current number of archived formulas  : 0
# Current number of archived clauses   : 323
# Clause-clause subsumption calls (NU) : 22861
# Rec. Clause-clause subsumption calls : 3487
# Non-unit clause-clause subsumptions  : 402
# Unit Clause-clause subsumption calls : 275
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 36168

# -------------------------------------------------
# User time                : 0.042 s
# System time              : 0.000 s
# Total time               : 0.042 s
# Maximum resident set size: 28788 pages
