# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t43_xboole_1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t98_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t7_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t3_xboole_0)).
fof(t103_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t103_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t48_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', idempotence_k3_xboole_0)).
fof(t67_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&(r1_tarski(X1,X3)&r1_xboole_0(X2,X3)))=>X1=k1_xboole_0), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t67_xboole_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t65_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t5_boole)).
fof(t9_zfmisc_1, conjecture, ![X1, X2, X3]:(k1_tarski(X1)=k2_tarski(X2,X3)=>X2=X3), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t9_zfmisc_1)).
fof(t41_enumset1, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2)), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t41_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t70_enumset1)).
fof(t63_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X2),k4_enumset1(X3,X4,X5,X6,X7,X8)), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t63_enumset1)).
fof(t100_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X2,X3,X1), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t100_enumset1)).
fof(t98_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t98_enumset1)).
fof(t95_enumset1, axiom, ![X1, X2]:k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t95_enumset1)).
fof(l57_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_tarski(X4,X5)), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', l57_enumset1)).
fof(t86_enumset1, axiom, ![X1, X2, X3, X4, X5]:k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t86_enumset1)).
fof(t136_enumset1, axiom, ![X1, X2, X3]:~((~(X1=X2)&(~(X1=X3)&~(k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3))))), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t136_enumset1)).
fof(t129_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9)), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t129_enumset1)).
fof(d7_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9, X10]:(X10=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)<=>![X11]:(r2_hidden(X11,X10)<=>~((~(X11=X1)&(~(X11=X2)&(~(X11=X3)&(~(X11=X4)&(~(X11=X5)&(~(X11=X6)&(~(X11=X7)&(~(X11=X8)&~(X11=X9)))))))))))), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', d7_enumset1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t4_xboole_0)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t66_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/828b2c6227c141c488db24e5f380ba3e.E_input', t12_xboole_1)).
fof(c_0_26, plain, ![X52, X53, X54]:(~r1_tarski(X52,k2_xboole_0(X53,X54))|r1_tarski(k4_xboole_0(X52,X53),X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
fof(c_0_27, plain, ![X26, X27]:k2_xboole_0(X26,X27)=k5_xboole_0(X26,k4_xboole_0(X27,X26)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_28, plain, ![X24, X25]:r1_tarski(X24,k2_xboole_0(X24,X25)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_29, plain, ![X106, X107, X109, X110, X111]:(((r2_hidden(esk6_2(X106,X107),X106)|r1_xboole_0(X106,X107))&(r2_hidden(esk6_2(X106,X107),X107)|r1_xboole_0(X106,X107)))&(~r2_hidden(X111,X109)|~r2_hidden(X111,X110)|~r1_xboole_0(X109,X110))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_30, plain, ![X83, X84]:r1_xboole_0(k3_xboole_0(X83,X84),k5_xboole_0(X83,X84)), inference(variable_rename,[status(thm)],[t103_xboole_1])).
fof(c_0_31, plain, ![X96, X97]:k4_xboole_0(X96,k4_xboole_0(X96,X97))=k3_xboole_0(X96,X97), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_32, plain, ![X82]:k3_xboole_0(X82,X82)=X82, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_33, plain, ![X15, X16, X17]:(~r1_tarski(X15,X16)|(~r1_tarski(X15,X17)|~r1_xboole_0(X16,X17))|X15=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_xboole_1])])).
fof(c_0_34, plain, ![X18]:r1_xboole_0(X18,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_35, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (X1=k1_xboole_0|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_46, plain, (r1_tarski(X1,k5_xboole_0(X1,k4_xboole_0(X2,X1)))), inference(rw,[status(thm)],[c_0_37, c_0_36])).
cnf(c_0_47, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk6_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_48, plain, (r2_hidden(esk6_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_49, plain, (r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X1))=X1), inference(rw,[status(thm)],[c_0_42, c_0_41])).
cnf(c_0_51, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52, plain, (r1_tarski(k4_xboole_0(X1,X1),X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_54, plain, (r1_xboole_0(X1,k5_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_52])])).
cnf(c_0_56, plain, (r1_xboole_0(k5_xboole_0(X1,X1),X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_57, plain, (r1_tarski(k1_xboole_0,X1)), inference(rw,[status(thm)],[c_0_52, c_0_55])).
fof(c_0_58, plain, ![X93]:k5_xboole_0(X93,k1_xboole_0)=X93, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_59, plain, (X1=k1_xboole_0|~r1_tarski(X1,k5_xboole_0(X2,X2))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_56])).
cnf(c_0_60, plain, (r1_tarski(k4_xboole_0(k1_xboole_0,X1),X2)), inference(spm,[status(thm)],[c_0_45, c_0_57])).
fof(c_0_61, negated_conjecture, ~(![X1, X2, X3]:(k1_tarski(X1)=k2_tarski(X2,X3)=>X2=X3)), inference(assume_negation,[status(cth)],[t9_zfmisc_1])).
cnf(c_0_62, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_63, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_60])])).
fof(c_0_64, plain, ![X47, X48]:k2_tarski(X47,X48)=k2_xboole_0(k1_tarski(X47),k1_tarski(X48)), inference(variable_rename,[status(thm)],[t41_enumset1])).
fof(c_0_65, plain, ![X55]:k2_tarski(X55,X55)=k1_tarski(X55), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_66, plain, ![X45, X46]:k1_enumset1(X45,X45,X46)=k2_tarski(X45,X46), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_67, negated_conjecture, (k1_tarski(esk1_0)=k2_tarski(esk2_0,esk3_0)&esk2_0!=esk3_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_61])])])).
cnf(c_0_68, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_59, c_0_62])).
cnf(c_0_69, plain, (r1_tarski(k4_xboole_0(X1,X2),k1_xboole_0)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_63]), c_0_62])).
fof(c_0_70, plain, ![X28, X29, X30, X31, X32, X33, X34, X35]:k6_enumset1(X28,X29,X30,X31,X32,X33,X34,X35)=k2_xboole_0(k2_tarski(X28,X29),k4_enumset1(X30,X31,X32,X33,X34,X35)), inference(variable_rename,[status(thm)],[t63_enumset1])).
cnf(c_0_71, plain, (k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_73, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, negated_conjecture, (k1_tarski(esk1_0)=k2_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
fof(c_0_76, plain, ![X49, X50, X51]:k1_enumset1(X49,X50,X51)=k1_enumset1(X50,X51,X49), inference(variable_rename,[status(thm)],[t100_enumset1])).
fof(c_0_77, plain, ![X90, X91, X92]:k1_enumset1(X90,X91,X92)=k1_enumset1(X90,X92,X91), inference(variable_rename,[status(thm)],[t98_enumset1])).
fof(c_0_78, plain, ![X94, X95]:k6_enumset1(X94,X94,X94,X94,X94,X94,X94,X95)=k2_tarski(X94,X95), inference(variable_rename,[status(thm)],[t95_enumset1])).
cnf(c_0_79, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X2),k4_enumset1(X3,X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_80, plain, ![X85, X86, X87, X88, X89]:k3_enumset1(X85,X86,X87,X88,X89)=k2_xboole_0(k1_enumset1(X85,X86,X87),k2_tarski(X88,X89)), inference(variable_rename,[status(thm)],[l57_enumset1])).
fof(c_0_81, plain, ![X19, X20, X21, X22, X23]:k6_enumset1(X19,X19,X19,X19,X20,X21,X22,X23)=k3_enumset1(X19,X20,X21,X22,X23), inference(variable_rename,[status(thm)],[t86_enumset1])).
fof(c_0_82, plain, ![X79, X80, X81]:(X79=X80|(X79=X81|k4_xboole_0(k1_enumset1(X79,X80,X81),k1_tarski(X79))=k2_tarski(X80,X81))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_enumset1])])])).
cnf(c_0_83, plain, (k1_enumset1(X1,X1,X2)=k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k1_enumset1(X2,X2,X2),k1_enumset1(X1,X1,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_72]), c_0_72]), c_0_73]), c_0_73]), c_0_73]), c_0_36])).
cnf(c_0_84, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk3_0)=k1_enumset1(esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_72]), c_0_73]), c_0_73])).
cnf(c_0_85, plain, (k4_xboole_0(X1,k5_xboole_0(X1,k4_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_46])).
cnf(c_0_86, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_87, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_88, plain, ![X36, X37, X38, X39, X40, X41, X42, X43, X44]:k7_enumset1(X36,X37,X38,X39,X40,X41,X42,X43,X44)=k2_xboole_0(k1_enumset1(X36,X37,X38),k4_enumset1(X39,X40,X41,X42,X43,X44)), inference(variable_rename,[status(thm)],[t129_enumset1])).
cnf(c_0_89, plain, (k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_90, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k4_enumset1(X3,X4,X5,X6,X7,X8),k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_73]), c_0_36])).
cnf(c_0_91, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_tarski(X4,X5))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_92, plain, (k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_93, plain, (X1=X2|X1=X3|k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_94, negated_conjecture, (k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k4_xboole_0(k1_enumset1(X1,X1,X1),k1_enumset1(esk2_0,esk2_0,esk3_0)))=k1_enumset1(esk1_0,esk1_0,X1)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_95, plain, (k4_xboole_0(k1_enumset1(X1,X1,X1),k1_enumset1(X1,X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_85, c_0_83])).
cnf(c_0_96, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
fof(c_0_97, plain, ![X56, X57, X58, X59, X60, X61, X62, X63, X64, X65, X66, X67, X68, X69, X70, X71, X72, X73, X74, X75, X76, X77]:(((~r2_hidden(X66,X65)|(X66=X56|(X66=X57|(X66=X58|(X66=X59|(X66=X60|(X66=X61|(X66=X62|(X66=X63|X66=X64))))))))|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&((X67!=X56|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&((X67!=X57|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&((X67!=X58|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&((X67!=X59|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&((X67!=X60|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&((X67!=X61|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&((X67!=X62|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&((X67!=X63|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64))&(X67!=X64|r2_hidden(X67,X65)|X65!=k7_enumset1(X56,X57,X58,X59,X60,X61,X62,X63,X64)))))))))))&(((esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X68|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))&((esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X69|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))&((esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X70|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))&((esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X71|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))&((esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X72|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))&((esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X73|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))&((esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X74|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))&((esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X75|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))&(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)!=X76|~r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76))))))))))&(r2_hidden(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77),X77)|(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X68|(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X69|(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X70|(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X71|(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X72|(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X73|(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X74|(esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X75|esk4_10(X68,X69,X70,X71,X72,X73,X74,X75,X76,X77)=X76))))))))|X77=k7_enumset1(X68,X69,X70,X71,X72,X73,X74,X75,X76)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_enumset1])])])])])])])).
cnf(c_0_98, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_99, plain, ![X100, X101, X103, X104, X105]:((r1_xboole_0(X100,X101)|r2_hidden(esk5_2(X100,X101),k3_xboole_0(X100,X101)))&(~r2_hidden(X105,k3_xboole_0(X103,X104))|~r1_xboole_0(X103,X104))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_100, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k4_enumset1(X1,X1,X1,X1,X1,X2),k1_enumset1(X1,X1,X1)))=k1_enumset1(X1,X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_73]), c_0_90])).
cnf(c_0_101, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k4_enumset1(X1,X1,X2,X3,X4,X5),k1_enumset1(X1,X1,X1)))=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X4,X5),k1_enumset1(X1,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_73]), c_0_36]), c_0_92]), c_0_90])).
cnf(c_0_102, plain, (X1=X3|X1=X2|k4_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X1,X1,X1))=k1_enumset1(X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_72]), c_0_73]), c_0_73])).
cnf(c_0_103, negated_conjecture, (k1_enumset1(esk2_0,esk1_0,esk1_0)=k1_enumset1(esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_62]), c_0_96])).
cnf(c_0_104, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k7_enumset1(X4,X5,X6,X7,X2,X8,X9,X10,X11)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_105, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k4_enumset1(X4,X5,X6,X7,X8,X9),k1_enumset1(X1,X2,X3)))), inference(rw,[status(thm)],[c_0_98, c_0_36])).
cnf(c_0_106, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_107, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k1_enumset1(X1,X1,X2),k1_enumset1(X1,X1,X1)))=k1_enumset1(X1,X1,X2)), inference(rw,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_108, negated_conjecture, (k4_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k1_enumset1(esk2_0,esk2_0,esk2_0))=k1_enumset1(esk2_0,esk2_0,esk3_0)|esk1_0=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_84])).
cnf(c_0_109, negated_conjecture, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k1_enumset1(X1,X1,X1)))=k1_enumset1(X1,X1,esk1_0)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_110, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k5_xboole_0(k1_enumset1(X4,X5,X6),k4_xboole_0(k4_enumset1(X7,X2,X8,X9,X10,X11),k1_enumset1(X4,X5,X6)))), inference(rw,[status(thm)],[c_0_104, c_0_105])).
fof(c_0_111, plain, ![X112]:((r1_xboole_0(X112,X112)|X112!=k1_xboole_0)&(X112=k1_xboole_0|~r1_xboole_0(X112,X112))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_112, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_106, c_0_41])).
cnf(c_0_113, negated_conjecture, (k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk2_0),k1_enumset1(esk2_0,esk2_0,esk3_0))=k1_enumset1(esk2_0,esk2_0,esk3_0)|esk1_0=esk2_0), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_114, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk3_0)=k1_enumset1(esk2_0,esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_107, c_0_109])).
cnf(c_0_115, plain, (r2_hidden(X1,k5_xboole_0(k1_enumset1(X2,X3,X4),k4_xboole_0(k4_enumset1(X5,X1,X6,X7,X8,X9),k1_enumset1(X2,X3,X4))))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_110])])).
cnf(c_0_116, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_117, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_112, c_0_48])).
cnf(c_0_118, negated_conjecture, (k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk2_0),k1_enumset1(esk2_0,esk2_0,esk1_0))=k1_enumset1(esk2_0,esk2_0,esk1_0)|esk1_0=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_114]), c_0_114])).
cnf(c_0_119, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_50, c_0_55])).
cnf(c_0_120, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,k5_xboole_0(k1_enumset1(X3,X4,X5),k4_xboole_0(k4_enumset1(X6,X1,X7,X8,X9,X10),k1_enumset1(X3,X4,X5))))), inference(spm,[status(thm)],[c_0_38, c_0_115])).
cnf(c_0_121, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k4_enumset1(X1,X1,X1,X1,X2,X2),k1_enumset1(X1,X1,X1)))=k1_enumset1(X1,X1,X2)), inference(spm,[status(thm)],[c_0_83, c_0_101])).
cnf(c_0_122, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_123, negated_conjecture, (esk1_0=esk2_0|r1_xboole_0(k1_enumset1(esk2_0,esk2_0,esk2_0),k1_enumset1(esk2_0,esk2_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_118]), c_0_95]), c_0_119])).
cnf(c_0_124, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k7_enumset1(X4,X5,X6,X7,X8,X9,X10,X11,X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_125, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,k1_enumset1(X1,X1,X3))), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_126, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=k1_xboole_0|esk1_0=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_95]), c_0_119])).
cnf(c_0_127, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k5_xboole_0(k1_enumset1(X4,X5,X6),k4_xboole_0(k4_enumset1(X7,X8,X9,X10,X11,X2),k1_enumset1(X4,X5,X6)))), inference(rw,[status(thm)],[c_0_124, c_0_105])).
cnf(c_0_128, plain, (r1_tarski(k1_enumset1(X1,X1,X1),k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X4,X5),k1_enumset1(X1,X2,X3))))), inference(spm,[status(thm)],[c_0_46, c_0_101])).
cnf(c_0_129, plain, (k4_xboole_0(k1_enumset1(X1,X1,X1),k1_enumset1(X2,X1,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_130, negated_conjecture, (esk1_0=esk2_0|~r2_hidden(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_44])])).
cnf(c_0_131, plain, (r2_hidden(X1,k5_xboole_0(k1_enumset1(X2,X3,X4),k4_xboole_0(k4_enumset1(X5,X6,X7,X8,X9,X1),k1_enumset1(X2,X3,X4))))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_127])])).
fof(c_0_132, plain, ![X98, X99]:(~r1_tarski(X98,X99)|k2_xboole_0(X98,X99)=X99), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_133, plain, (r1_tarski(k1_enumset1(X1,X1,X1),k1_enumset1(X1,X2,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_62])).
cnf(c_0_134, negated_conjecture, (esk1_0=esk2_0), inference(spm,[status(thm)],[c_0_130, c_0_131])).
cnf(c_0_135, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k7_enumset1(X2,X4,X5,X6,X7,X8,X9,X10,X11)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_136, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_137, plain, (r1_tarski(k1_enumset1(X1,X1,X1),k1_enumset1(X2,X2,X1))), inference(spm,[status(thm)],[c_0_133, c_0_86])).
cnf(c_0_138, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk3_0)=k1_enumset1(esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_114, c_0_134])).
cnf(c_0_139, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k5_xboole_0(k1_enumset1(X2,X4,X5),k4_xboole_0(k4_enumset1(X6,X7,X8,X9,X10,X11),k1_enumset1(X2,X4,X5)))), inference(rw,[status(thm)],[c_0_135, c_0_105])).
cnf(c_0_140, plain, (k5_xboole_0(X1,k4_xboole_0(X2,X1))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_136, c_0_36])).
cnf(c_0_141, negated_conjecture, (r1_tarski(k1_enumset1(esk3_0,esk3_0,esk3_0),k1_enumset1(esk2_0,esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_142, plain, (r2_hidden(X1,k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k4_enumset1(X4,X5,X6,X7,X8,X9),k1_enumset1(X1,X2,X3))))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_139])])).
cnf(c_0_143, negated_conjecture, (k1_enumset1(esk2_0,esk3_0,esk3_0)=k1_enumset1(esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_83]), c_0_96])).
cnf(c_0_144, negated_conjecture, (esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_145, plain, (~r2_hidden(X1,k4_xboole_0(X2,X2))|~r1_xboole_0(X2,k4_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_112, c_0_50])).
cnf(c_0_146, plain, (r2_hidden(X1,k1_enumset1(X1,X1,X2))), inference(spm,[status(thm)],[c_0_142, c_0_121])).
cnf(c_0_147, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk3_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_143]), c_0_95]), c_0_144])).
cnf(c_0_148, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_145, c_0_55]), c_0_55]), c_0_44])])).
cnf(c_0_149, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_148]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 150
# Proof object clause steps            : 97
# Proof object formula steps           : 53
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 41
# Proof object simplifying inferences  : 66
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 43
# Processed clauses                    : 1929
# ...of these trivial                  : 152
# ...subsumed                          : 1415
# ...remaining for further processing  : 362
# Other redundant clauses eliminated   : 29
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 94
# Generated clauses                    : 10289
# ...of the previous two non-trivial   : 7744
# Contextual simplify-reflections      : 0
# Paramodulations                      : 10269
# Factorizations                       : 0
# Equation resolutions                 : 29
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 212
#    Positive orientable unit clauses  : 71
#    Positive unorientable unit clauses: 12
#    Negative unit clauses             : 20
#    Non-unit-clauses                  : 109
# Current number of unprocessed clauses: 5095
# ...number of literals in the above   : 18942
# Current number of archived formulas  : 0
# Current number of archived clauses   : 146
# Clause-clause subsumption calls (NU) : 7260
# Rec. Clause-clause subsumption calls : 6590
# Non-unit clause-clause subsumptions  : 1087
# Unit Clause-clause subsumption calls : 572
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 686
# BW rewrite match successes           : 129
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 140470

# -------------------------------------------------
# User time                : 0.108 s
# System time              : 0.007 s
# Total time               : 0.114 s
# Maximum resident set size: 28740 pages
