# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t18_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))<=>r2_waybel_7(X1,X2,X3))))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', t18_yellow19)).
fof(dt_k1_connsp_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k1_connsp_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', dt_k1_connsp_2)).
fof(dt_k3_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&l1_waybel_0(k3_yellow19(X1,X2,X3),X1)))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', dt_k3_yellow19)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', d2_yellow_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', d4_subset_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', dt_l1_pre_topc)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', t5_subset)).
fof(t30_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r2_hidden(X2,k1_connsp_2(X1,X2)))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', t30_connsp_2)).
fof(t15_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', t15_yellow19)).
fof(fc4_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v3_orders_2(k3_yellow19(X1,X2,X3))&(v4_orders_2(k3_yellow19(X1,X2,X3))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', fc4_yellow19)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', d3_struct_0)).
fof(fc5_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&v7_waybel_0(k3_yellow19(X1,X2,X3))))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', fc5_yellow19)).
fof(t13_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k10_yellow_6(X1,X2))<=>r2_waybel_7(X1,k2_yellow19(X1,X2),X3))))), file('data/MPTP/proofs/83c527eca1d64e26b5d2b1c7f35fb99d.E_input', t13_yellow19)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k10_yellow_6(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))<=>r2_waybel_7(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t18_yellow19])).
fof(c_0_15, plain, ![X15, X16]:(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))|~m1_subset_1(X16,u1_struct_0(X15))|m1_subset_1(k1_connsp_2(X15,X16),k1_zfmisc_1(u1_struct_0(X15)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_connsp_2])])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v1_xboole_0(esk2_0)&(v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))&(v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&(v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))))))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r2_hidden(esk3_0,k10_yellow_6(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0)))|~r2_waybel_7(esk1_0,esk2_0,esk3_0))&(r2_hidden(esk3_0,k10_yellow_6(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0)))|r2_waybel_7(esk1_0,esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_17, plain, ![X7, X8, X9]:((~v2_struct_0(k3_yellow19(X7,X8,X9))|(v2_struct_0(X7)|~l1_struct_0(X7)|(v1_xboole_0(X8)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(v1_xboole_0(X9)|(~v2_waybel_0(X9,k3_yellow_1(X8))|(~v13_waybel_0(X9,k3_yellow_1(X8))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X8))))))))))&((v6_waybel_0(k3_yellow19(X7,X8,X9),X7)|(v2_struct_0(X7)|~l1_struct_0(X7)|(v1_xboole_0(X8)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(v1_xboole_0(X9)|(~v2_waybel_0(X9,k3_yellow_1(X8))|(~v13_waybel_0(X9,k3_yellow_1(X8))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X8))))))))))&(l1_waybel_0(k3_yellow19(X7,X8,X9),X7)|(v2_struct_0(X7)|~l1_struct_0(X7)|(v1_xboole_0(X8)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(v1_xboole_0(X9)|(~v2_waybel_0(X9,k3_yellow_1(X8))|(~v13_waybel_0(X9,k3_yellow_1(X8))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X8)))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_yellow19])])])])).
fof(c_0_18, plain, ![X13]:k3_yellow_1(X13)=k3_lattice3(k1_lattice3(X13)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_19, plain, ![X32]:m1_subset_1(k2_subset_1(X32),k1_zfmisc_1(X32)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_20, plain, ![X21]:k2_subset_1(X21)=X21, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_21, plain, ![X14]:(~l1_pre_topc(X14)|l1_struct_0(X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_22, plain, ![X22, X23, X24]:(~r2_hidden(X22,X23)|(~m1_subset_1(X23,k1_zfmisc_1(X24))|~v1_xboole_0(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_23, plain, (v2_struct_0(X1)|m1_subset_1(k1_connsp_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_28, plain, ![X19, X20]:(v2_struct_0(X19)|(~v2_pre_topc(X19)|~l1_pre_topc(X19))|(~m1_subset_1(X20,u1_struct_0(X19))|r2_hidden(X20,k1_connsp_2(X19,X20)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_connsp_2])])])])).
fof(c_0_29, plain, ![X17, X18]:(v2_struct_0(X17)|~l1_struct_0(X17)|(v1_xboole_0(X18)|(~v1_subset_1(X18,u1_struct_0(k3_yellow_1(k2_struct_0(X17))))|(~v2_waybel_0(X18,k3_yellow_1(k2_struct_0(X17)))|(~v13_waybel_0(X18,k3_yellow_1(k2_struct_0(X17)))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X17))))))))|X18=k2_yellow19(X17,k3_yellow19(X17,k2_struct_0(X17),X18)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_yellow19])])])])).
fof(c_0_30, plain, ![X10, X11, X12]:((~v2_struct_0(k3_yellow19(X10,X11,X12))|(v2_struct_0(X10)|~l1_struct_0(X10)|(v1_xboole_0(X11)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v1_xboole_0(X12)|(~v2_waybel_0(X12,k3_yellow_1(X11))|(~v13_waybel_0(X12,k3_yellow_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X11))))))))))&((v3_orders_2(k3_yellow19(X10,X11,X12))|(v2_struct_0(X10)|~l1_struct_0(X10)|(v1_xboole_0(X11)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v1_xboole_0(X12)|(~v2_waybel_0(X12,k3_yellow_1(X11))|(~v13_waybel_0(X12,k3_yellow_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X11))))))))))&((v4_orders_2(k3_yellow19(X10,X11,X12))|(v2_struct_0(X10)|~l1_struct_0(X10)|(v1_xboole_0(X11)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v1_xboole_0(X12)|(~v2_waybel_0(X12,k3_yellow_1(X11))|(~v13_waybel_0(X12,k3_yellow_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X11))))))))))&(v6_waybel_0(k3_yellow19(X10,X11,X12),X10)|(v2_struct_0(X10)|~l1_struct_0(X10)|(v1_xboole_0(X11)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v1_xboole_0(X12)|(~v2_waybel_0(X12,k3_yellow_1(X11))|(~v13_waybel_0(X12,k3_yellow_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X11))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_yellow19])])])])).
cnf(c_0_31, plain, (l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_32, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_33, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_35, negated_conjecture, (v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_36, plain, ![X25]:(~l1_struct_0(X25)|k2_struct_0(X25)=u1_struct_0(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_37, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_39, negated_conjecture, (v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_40, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(k1_connsp_2(esk1_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_42, plain, (v2_struct_0(X1)|r2_hidden(X2,k1_connsp_2(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_43, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|~l1_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_45, plain, (v4_orders_2(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_46, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_47, plain, ![X26, X27, X28]:((~v2_struct_0(k3_yellow19(X26,X27,X28))|(v2_struct_0(X26)|~l1_struct_0(X26)|(v1_xboole_0(X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))|(v1_xboole_0(X28)|(~v1_subset_1(X28,u1_struct_0(k3_yellow_1(X27)))|(~v2_waybel_0(X28,k3_yellow_1(X27))|(~v13_waybel_0(X28,k3_yellow_1(X27))|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X27)))))))))))&((v6_waybel_0(k3_yellow19(X26,X27,X28),X26)|(v2_struct_0(X26)|~l1_struct_0(X26)|(v1_xboole_0(X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))|(v1_xboole_0(X28)|(~v1_subset_1(X28,u1_struct_0(k3_yellow_1(X27)))|(~v2_waybel_0(X28,k3_yellow_1(X27))|(~v13_waybel_0(X28,k3_yellow_1(X27))|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X27)))))))))))&(v7_waybel_0(k3_yellow19(X26,X27,X28))|(v2_struct_0(X26)|~l1_struct_0(X26)|(v1_xboole_0(X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26)))|(v1_xboole_0(X28)|(~v1_subset_1(X28,u1_struct_0(k3_yellow_1(X27)))|(~v2_waybel_0(X28,k3_yellow_1(X27))|(~v13_waybel_0(X28,k3_yellow_1(X27))|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X27))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_yellow19])])])])).
fof(c_0_48, plain, ![X29, X30, X31]:((~r2_hidden(X31,k10_yellow_6(X29,X30))|r2_waybel_7(X29,k2_yellow19(X29,X30),X31)|~m1_subset_1(X31,u1_struct_0(X29))|(v2_struct_0(X30)|(~v4_orders_2(X30)|(~v7_waybel_0(X30)|~l1_waybel_0(X30,X29))))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))))&(~r2_waybel_7(X29,k2_yellow19(X29,X30),X31)|r2_hidden(X31,k10_yellow_6(X29,X30))|~m1_subset_1(X31,u1_struct_0(X29))|(v2_struct_0(X30)|(~v4_orders_2(X30)|(~v7_waybel_0(X30)|~l1_waybel_0(X30,X29))))|(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t13_yellow19])])])])])).
cnf(c_0_49, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32]), c_0_32]), c_0_32])).
cnf(c_0_50, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_51, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_35, c_0_32])).
cnf(c_0_52, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_53, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_25])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))))), inference(rw,[status(thm)],[c_0_38, c_0_32])).
cnf(c_0_55, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_39, c_0_32])).
cnf(c_0_56, negated_conjecture, (~r2_hidden(X1,k1_connsp_2(esk1_0,esk3_0))|~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk3_0,k1_connsp_2(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_58, plain, (X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|v2_struct_0(X1)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_32]), c_0_32]), c_0_32]), c_0_32])).
cnf(c_0_59, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0)))))), inference(rw,[status(thm)],[c_0_44, c_0_32])).
cnf(c_0_60, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_61, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_32]), c_0_32]), c_0_32])).
cnf(c_0_62, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_32]), c_0_32]), c_0_32])).
cnf(c_0_63, plain, (v7_waybel_0(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_64, plain, (r2_waybel_7(X2,k2_yellow19(X2,X3),X1)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,k10_yellow_6(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_65, plain, (l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),X2),X1)|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_66, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])])).
cnf(c_0_67, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_52]), c_0_53])])).
cnf(c_0_68, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_52]), c_0_53])])).
cnf(c_0_69, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_70, negated_conjecture, (k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))=esk2_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_51]), c_0_53]), c_0_54]), c_0_55]), c_0_59])]), c_0_60]), c_0_27])).
cnf(c_0_71, plain, (v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),X2))|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))), inference(spm,[status(thm)],[c_0_61, c_0_50])).
cnf(c_0_72, plain, (v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),X2))), inference(spm,[status(thm)],[c_0_62, c_0_50])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk3_0,k10_yellow_6(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0)))|r2_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_74, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v7_waybel_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_32]), c_0_32]), c_0_32]), c_0_32])).
cnf(c_0_75, negated_conjecture, (r2_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)|v2_struct_0(X1)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_waybel_0(X1,esk1_0)|~r2_hidden(esk3_0,k10_yellow_6(esk1_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_76, negated_conjecture, (l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_53]), c_0_67]), c_0_68])]), c_0_69]), c_0_60]), c_0_27])).
cnf(c_0_77, negated_conjecture, (k2_yellow19(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_52]), c_0_53])])).
cnf(c_0_78, negated_conjecture, (v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_66]), c_0_53]), c_0_67]), c_0_68])]), c_0_69]), c_0_60]), c_0_27])).
cnf(c_0_79, negated_conjecture, (~v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_66]), c_0_53]), c_0_67]), c_0_68])]), c_0_69]), c_0_60]), c_0_27])).
cnf(c_0_80, negated_conjecture, (r2_waybel_7(esk1_0,esk2_0,esk3_0)|r2_hidden(esk3_0,k10_yellow_6(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_52]), c_0_53])])).
cnf(c_0_81, plain, (v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),X2))|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1)))))), inference(spm,[status(thm)],[c_0_74, c_0_50])).
cnf(c_0_82, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(esk1_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_52]), c_0_53])])).
cnf(c_0_83, plain, (r2_hidden(X3,k10_yellow_6(X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_7(X1,k2_yellow19(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_84, negated_conjecture, (~r2_hidden(esk3_0,k10_yellow_6(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0)))|~r2_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_85, negated_conjecture, (r2_waybel_7(esk1_0,esk2_0,esk3_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77]), c_0_78])]), c_0_79]), c_0_80])).
cnf(c_0_86, negated_conjecture, (v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_66]), c_0_53]), c_0_67]), c_0_68]), c_0_82])]), c_0_69]), c_0_60]), c_0_27])).
cnf(c_0_87, negated_conjecture, (r2_hidden(esk3_0,k10_yellow_6(esk1_0,X1))|v2_struct_0(X1)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~l1_waybel_0(X1,esk1_0)|~r2_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_88, negated_conjecture, (~r2_waybel_7(esk1_0,esk2_0,esk3_0)|~r2_hidden(esk3_0,k10_yellow_6(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_52]), c_0_53])])).
cnf(c_0_89, negated_conjecture, (r2_waybel_7(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_86])])).
cnf(c_0_90, negated_conjecture, (r2_hidden(esk3_0,k10_yellow_6(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0)))|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_76]), c_0_78]), c_0_77])]), c_0_79]), c_0_85])).
cnf(c_0_91, negated_conjecture, (~r2_hidden(esk3_0,k10_yellow_6(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_89])])).
cnf(c_0_92, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_86])]), c_0_91]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 93
# Proof object clause steps            : 64
# Proof object formula steps           : 29
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 105
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 30
# Processed clauses                    : 122
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 108
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 8
# Generated clauses                    : 67
# ...of the previous two non-trivial   : 69
# Contextual simplify-reflections      : 2
# Paramodulations                      : 67
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 74
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 44
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 10
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 1967
# Rec. Clause-clause subsumption calls : 72
# Non-unit clause-clause subsumptions  : 14
# Unit Clause-clause subsumption calls : 64
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7257

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.004 s
# Total time               : 0.014 s
# Maximum resident set size: 26276 pages
