# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', t3_subset)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', t63_subset_1)).
fof(t44_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tex_2(X2,X1)<=>v1_zfmisc_1(X2)))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', t44_tex_2)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', d1_xboole_0)).
fof(t1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', t1_tex_2)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', fc2_xboole_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', redefinition_k6_domain_1)).
fof(t36_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', t36_tex_2)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', d3_tarski)).
fof(t34_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>~((v2_tex_2(X2,X1)&![X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&(v1_tdlat_3(X3)&m1_pre_topc(X3,X1))))=>~(X2=u1_struct_0(X3))))))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', t34_tex_2)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', dt_m1_pre_topc)).
fof(fc7_struct_0, axiom, ![X1]:((v7_struct_0(X1)&l1_struct_0(X1))=>v1_zfmisc_1(u1_struct_0(X1))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', fc7_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', dt_l1_pre_topc)).
fof(cc32_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:(m1_pre_topc(X2,X1)=>((~(v2_struct_0(X2))&~(v7_struct_0(X2)))=>(~(v2_struct_0(X2))&~(v1_tdlat_3(X2)))))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', cc32_tex_2)).
fof(cc2_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/842e88be5d574bf9ac59e85db2f89887.E_input', cc2_tdlat_3)).
fof(c_0_16, plain, ![X34, X35]:((~m1_subset_1(X34,k1_zfmisc_1(X35))|r1_tarski(X34,X35))&(~r1_tarski(X34,X35)|m1_subset_1(X34,k1_zfmisc_1(X35)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_17, plain, ![X11, X12]:(~r2_hidden(X11,X12)|m1_subset_1(k1_tarski(X11),k1_zfmisc_1(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v2_tdlat_3(X1)&l1_pre_topc(X1))))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tex_2(X2,X1)<=>v1_zfmisc_1(X2))))), inference(assume_negation,[status(cth)],[t44_tex_2])).
fof(c_0_19, plain, ![X16, X17]:(((~m1_subset_1(X17,X16)|r2_hidden(X17,X16)|v1_xboole_0(X16))&(~r2_hidden(X17,X16)|m1_subset_1(X17,X16)|v1_xboole_0(X16)))&((~m1_subset_1(X17,X16)|v1_xboole_0(X17)|~v1_xboole_0(X16))&(~v1_xboole_0(X17)|m1_subset_1(X17,X16)|~v1_xboole_0(X16)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_20, plain, ![X30, X31, X32]:((~v1_xboole_0(X30)|~r2_hidden(X31,X30))&(r2_hidden(esk5_1(X32),X32)|v1_xboole_0(X32))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_21, plain, ![X37, X38]:(v1_xboole_0(X37)|(v1_xboole_0(X38)|~v1_zfmisc_1(X38)|(~r1_tarski(X37,X38)|X37=X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t1_tex_2])])])])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X29]:~v1_xboole_0(k1_tarski(X29)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&(v2_tdlat_3(esk1_0)&l1_pre_topc(esk1_0))))&((~v1_xboole_0(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))&((~v2_tex_2(esk2_0,esk1_0)|~v1_zfmisc_1(esk2_0))&(v2_tex_2(esk2_0,esk1_0)|v1_zfmisc_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
fof(c_0_26, plain, ![X20, X21]:(v1_xboole_0(X20)|~m1_subset_1(X21,X20)|k6_domain_1(X20,X21)=k1_tarski(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_27, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|X1=X2|~v1_zfmisc_1(X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_34, plain, ![X9, X10]:(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|(~m1_subset_1(X10,u1_struct_0(X9))|v2_tex_2(k6_domain_1(u1_struct_0(X9),X10),X9))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_tex_2])])])])).
cnf(c_0_35, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_37, plain, ![X23, X24, X25, X26, X27]:((~r1_tarski(X23,X24)|(~r2_hidden(X25,X23)|r2_hidden(X25,X24)))&((r2_hidden(esk4_2(X26,X27),X26)|r1_tarski(X26,X27))&(~r2_hidden(esk4_2(X26,X27),X27)|r1_tarski(X26,X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_39, plain, ![X6, X7]:(((~v2_struct_0(esk3_2(X6,X7))|~v2_tex_2(X7,X6)|(v1_xboole_0(X7)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((v1_pre_topc(esk3_2(X6,X7))|~v2_tex_2(X7,X6)|(v1_xboole_0(X7)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&((v1_tdlat_3(esk3_2(X6,X7))|~v2_tex_2(X7,X6)|(v1_xboole_0(X7)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(m1_pre_topc(esk3_2(X6,X7),X6)|~v2_tex_2(X7,X6)|(v1_xboole_0(X7)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6)))))))&(X7=u1_struct_0(esk3_2(X6,X7))|~v2_tex_2(X7,X6)|(v1_xboole_0(X7)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6))))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_tex_2])])])])])])).
cnf(c_0_40, plain, (k1_tarski(X1)=X2|~r2_hidden(X1,X2)|~v1_zfmisc_1(X2)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_28])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk5_1(esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, plain, (v2_struct_0(X1)|v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (k6_domain_1(X1,X2)=k1_tarski(X2)|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_28])).
cnf(c_0_44, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_22, c_0_38])).
cnf(c_0_46, plain, (m1_pre_topc(esk3_2(X1,X2),X1)|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_48, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_49, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_50, negated_conjecture, (k1_tarski(esk5_1(esk2_0))=esk2_0|~v1_zfmisc_1(esk2_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51, negated_conjecture, (v2_tex_2(esk2_0,esk1_0)|v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_52, plain, (v2_tex_2(k1_tarski(X1),X2)|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_36])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54, negated_conjecture, (m1_pre_topc(esk3_2(esk1_0,esk2_0),esk1_0)|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_38]), c_0_47]), c_0_48])]), c_0_32]), c_0_49])).
cnf(c_0_55, negated_conjecture, (k1_tarski(esk5_1(esk2_0))=esk2_0|v2_tex_2(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_56, plain, ![X18, X19]:(~l1_pre_topc(X18)|(~m1_pre_topc(X19,X18)|l1_pre_topc(X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_57, negated_conjecture, (v2_tex_2(k1_tarski(X1),esk1_0)|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_47]), c_0_48])]), c_0_49])).
cnf(c_0_58, negated_conjecture, (k1_tarski(esk5_1(esk2_0))=esk2_0|m1_pre_topc(esk3_2(esk1_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59, plain, (v1_tdlat_3(esk3_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_60, plain, ![X15]:(~v7_struct_0(X15)|~l1_struct_0(X15)|v1_zfmisc_1(u1_struct_0(X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_struct_0])])).
cnf(c_0_61, plain, (X1=u1_struct_0(esk3_2(X2,X1))|v1_xboole_0(X1)|v2_struct_0(X2)|~v2_tex_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_62, plain, ![X22]:(~l1_pre_topc(X22)|l1_struct_0(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_63, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, negated_conjecture, (m1_pre_topc(esk3_2(esk1_0,esk2_0),esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_41])]), c_0_54])).
fof(c_0_65, plain, ![X13, X14]:((~v2_struct_0(X14)|(v2_struct_0(X14)|v7_struct_0(X14))|~m1_pre_topc(X14,X13)|(v2_struct_0(X13)|(~v2_pre_topc(X13)|(~v2_tdlat_3(X13)|~l1_pre_topc(X13)))))&(~v1_tdlat_3(X14)|(v2_struct_0(X14)|v7_struct_0(X14))|~m1_pre_topc(X14,X13)|(v2_struct_0(X13)|(~v2_pre_topc(X13)|(~v2_tdlat_3(X13)|~l1_pre_topc(X13)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc32_tex_2])])])])])).
fof(c_0_66, plain, ![X36]:(~l1_pre_topc(X36)|(~v2_tdlat_3(X36)|v2_pre_topc(X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_tdlat_3])])).
cnf(c_0_67, negated_conjecture, (v1_tdlat_3(esk3_2(esk1_0,esk2_0))|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_38]), c_0_47]), c_0_48])]), c_0_32]), c_0_49])).
cnf(c_0_68, plain, (v1_zfmisc_1(u1_struct_0(X1))|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_69, negated_conjecture, (u1_struct_0(esk3_2(esk1_0,esk2_0))=esk2_0|~v2_tex_2(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_38]), c_0_47]), c_0_48])]), c_0_32]), c_0_49])).
cnf(c_0_70, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_71, negated_conjecture, (l1_pre_topc(esk3_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_47])])).
cnf(c_0_72, plain, (v2_struct_0(X1)|v7_struct_0(X1)|v2_struct_0(X2)|~v1_tdlat_3(X1)|~m1_pre_topc(X1,X2)|~v2_pre_topc(X2)|~v2_tdlat_3(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_73, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, negated_conjecture, (k1_tarski(esk5_1(esk2_0))=esk2_0|v1_tdlat_3(esk3_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_67, c_0_55])).
cnf(c_0_75, negated_conjecture, (v1_zfmisc_1(esk2_0)|~l1_struct_0(esk3_2(esk1_0,esk2_0))|~v7_struct_0(esk3_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_51])).
cnf(c_0_76, negated_conjecture, (l1_struct_0(esk3_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_77, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|~v2_struct_0(esk3_2(X1,X2))|~v2_tex_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_78, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_79, plain, (v7_struct_0(X1)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X1,X2)|~v1_tdlat_3(X1)|~l1_pre_topc(X2)|~v2_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_80, negated_conjecture, (v1_tdlat_3(esk3_2(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_74]), c_0_41])]), c_0_67])).
cnf(c_0_81, negated_conjecture, (v2_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_82, negated_conjecture, (~v2_tex_2(esk2_0,esk1_0)|~v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_83, negated_conjecture, (v1_zfmisc_1(esk2_0)|~v7_struct_0(esk3_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_76])])).
cnf(c_0_84, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|~r1_tarski(X1,u1_struct_0(X2))|~v2_tex_2(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~v2_struct_0(esk3_2(X2,X1))), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_85, negated_conjecture, (v7_struct_0(esk3_2(esk1_0,esk2_0))|v2_struct_0(esk3_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_64]), c_0_80]), c_0_47]), c_0_81])]), c_0_49])).
cnf(c_0_86, negated_conjecture, (~v7_struct_0(esk3_2(esk1_0,esk2_0))|~v2_tex_2(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_87, negated_conjecture, (~v2_tex_2(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_45]), c_0_47]), c_0_48])]), c_0_32]), c_0_49]), c_0_86])).
cnf(c_0_88, negated_conjecture, (k1_tarski(esk5_1(esk2_0))=esk2_0), inference(sr,[status(thm)],[c_0_55, c_0_87])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_88]), c_0_41])]), c_0_87]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 57
# Proof object formula steps           : 33
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 34
# Processed clauses                    : 361
# ...of these trivial                  : 5
# ...subsumed                          : 105
# ...remaining for further processing  : 251
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 17
# Generated clauses                    : 920
# ...of the previous two non-trivial   : 793
# Contextual simplify-reflections      : 19
# Paramodulations                      : 918
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 189
#    Positive orientable unit clauses  : 38
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 16
#    Non-unit-clauses                  : 135
# Current number of unprocessed clauses: 484
# ...number of literals in the above   : 2414
# Current number of archived formulas  : 0
# Current number of archived clauses   : 62
# Clause-clause subsumption calls (NU) : 6412
# Rec. Clause-clause subsumption calls : 953
# Non-unit clause-clause subsumptions  : 102
# Unit Clause-clause subsumption calls : 574
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 24
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17879

# -------------------------------------------------
# User time                : 0.022 s
# System time              : 0.000 s
# Total time               : 0.022 s
# Maximum resident set size: 28728 pages
