# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t26_waybel_9, conjecture, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_waybel_9(X1))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(v5_pre_topc(k4_waybel_1(X1,X3),X1,X1)=>r2_hidden(k12_lattice3(X1,X3,k11_yellow_6(X1,X2)),k10_yellow_6(X1,k3_waybel_2(X1,X3,X2))))))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', t26_waybel_9)).
fof(dt_l1_waybel_9, axiom, ![X1]:(l1_waybel_9(X1)=>(l1_pre_topc(X1)&l1_orders_2(X1))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', dt_l1_waybel_9)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', cc1_lattice3)).
fof(dt_k11_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&(v8_pre_topc(X1)&l1_pre_topc(X1))))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1))))))=>m1_subset_1(k11_yellow_6(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', dt_k11_yellow_6)).
fof(dt_k4_waybel_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(v1_funct_1(k4_waybel_1(X1,X2))&(v1_funct_2(k4_waybel_1(X1,X2),u1_struct_0(X1),u1_struct_0(X1))&m1_subset_1(k4_waybel_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', dt_k4_waybel_1)).
fof(t18_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k6_waybel_9(X1,X1,k4_waybel_1(X1,X3),X2)=k3_waybel_2(X1,X3,X2)))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', t18_waybel_9)).
fof(d18_waybel_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))))=>(X3=k4_waybel_1(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k3_funct_2(u1_struct_0(X1),u1_struct_0(X1),X3,X4)=k11_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', d18_waybel_1)).
fof(t25_waybel_9, axiom, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_waybel_9(X1))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))))=>(v5_pre_topc(X3,X1,X1)=>r2_hidden(k2_yellow_6(u1_struct_0(X1),X1,X3,k11_yellow_6(X1,X2)),k10_yellow_6(X1,k6_waybel_9(X1,X1,X3,X2))))))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', t25_waybel_9)).
fof(redefinition_k2_yellow_6, axiom, ![X1, X2, X3, X4]:((~(v1_xboole_0(X1))&((~(v2_struct_0(X2))&l1_orders_2(X2))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))))&m1_subset_1(X4,X1))))=>k2_yellow_6(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', redefinition_k2_yellow_6)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', redefinition_k3_funct_2)).
fof(redefinition_k12_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', redefinition_k12_lattice3)).
fof(fc2_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>~(v1_xboole_0(k2_struct_0(X1)))), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', fc2_yellow_0)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/8450c417f7b54df6b5d15599d711e5e3.E_input', dt_l1_pre_topc)).
fof(c_0_14, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&l1_waybel_9(X1))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&(v3_yellow_6(X2,X1)&l1_waybel_0(X2,X1)))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(v5_pre_topc(k4_waybel_1(X1,X3),X1,X1)=>r2_hidden(k12_lattice3(X1,X3,k11_yellow_6(X1,X2)),k10_yellow_6(X1,k3_waybel_2(X1,X3,X2)))))))), inference(assume_negation,[status(cth)],[t26_waybel_9])).
fof(c_0_15, plain, ![X8]:((l1_pre_topc(X8)|~l1_waybel_9(X8))&(l1_orders_2(X8)|~l1_waybel_9(X8))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_9])])])).
fof(c_0_16, negated_conjecture, ((v2_pre_topc(esk1_0)&(v8_pre_topc(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&l1_waybel_9(esk1_0))))))))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&(v3_yellow_6(esk2_0,esk1_0)&l1_waybel_0(esk2_0,esk1_0)))))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(v5_pre_topc(k4_waybel_1(esk1_0,esk3_0),esk1_0,esk1_0)&~r2_hidden(k12_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_17, plain, ![X9]:(~l1_orders_2(X9)|(~v1_lattice3(X9)|~v2_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
cnf(c_0_18, plain, (l1_orders_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (l1_waybel_9(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, plain, ![X10, X11]:(v2_struct_0(X10)|(~v2_pre_topc(X10)|(~v8_pre_topc(X10)|~l1_pre_topc(X10)))|(v2_struct_0(X11)|(~v4_orders_2(X11)|(~v7_waybel_0(X11)|(~v3_yellow_6(X11,X10)|~l1_waybel_0(X11,X10)))))|m1_subset_1(k11_yellow_6(X10,X11),u1_struct_0(X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k11_yellow_6])])])).
cnf(c_0_21, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (l1_orders_2(esk1_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_24, plain, ![X12, X13]:((v1_funct_1(k4_waybel_1(X12,X13))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,u1_struct_0(X12))))&((v1_funct_2(k4_waybel_1(X12,X13),u1_struct_0(X12),u1_struct_0(X12))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,u1_struct_0(X12))))&(m1_subset_1(k4_waybel_1(X12,X13),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X12))))|(v2_struct_0(X12)|~l1_orders_2(X12)|~m1_subset_1(X13,u1_struct_0(X12)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_waybel_1])])])])).
fof(c_0_25, plain, ![X14, X15, X16]:(v2_struct_0(X14)|~l1_orders_2(X14)|(v2_struct_0(X15)|~l1_waybel_0(X15,X14)|(~m1_subset_1(X16,u1_struct_0(X14))|k6_waybel_9(X14,X14,k4_waybel_1(X14,X16),X15)=k3_waybel_2(X14,X16,X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_9])])])])).
fof(c_0_26, plain, ![X17, X18, X19, X20]:((X19!=k4_waybel_1(X17,X18)|(~m1_subset_1(X20,u1_struct_0(X17))|k3_funct_2(u1_struct_0(X17),u1_struct_0(X17),X19,X20)=k11_lattice3(X17,X18,X20))|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X17))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X17))))))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l1_orders_2(X17)))&((m1_subset_1(esk4_3(X17,X18,X19),u1_struct_0(X17))|X19=k4_waybel_1(X17,X18)|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X17))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X17))))))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l1_orders_2(X17)))&(k3_funct_2(u1_struct_0(X17),u1_struct_0(X17),X19,esk4_3(X17,X18,X19))!=k11_lattice3(X17,X18,esk4_3(X17,X18,X19))|X19=k4_waybel_1(X17,X18)|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X17))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X17))))))|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|~l1_orders_2(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_waybel_1])])])])])])).
cnf(c_0_27, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(k11_yellow_6(X1,X2),u1_struct_0(X1))|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~v3_yellow_6(X2,X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (v3_yellow_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_30, negated_conjecture, (v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_31, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_32, negated_conjecture, (v8_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_33, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_34, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_35, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])])).
fof(c_0_36, plain, ![X22, X23, X24]:(~v2_pre_topc(X22)|(~v8_pre_topc(X22)|(~v3_orders_2(X22)|(~v4_orders_2(X22)|(~v5_orders_2(X22)|(~v1_lattice3(X22)|(~v2_lattice3(X22)|~l1_waybel_9(X22)))))))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|(~v3_yellow_6(X23,X22)|~l1_waybel_0(X23,X22))))|(~v1_funct_1(X24)|(~v1_funct_2(X24,u1_struct_0(X22),u1_struct_0(X22))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X22),u1_struct_0(X22)))))|(~v5_pre_topc(X24,X22,X22)|r2_hidden(k2_yellow_6(u1_struct_0(X22),X22,X24,k11_yellow_6(X22,X23)),k10_yellow_6(X22,k6_waybel_9(X22,X22,X24,X23))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_waybel_9])])])])).
cnf(c_0_37, plain, (v1_funct_2(k4_waybel_1(X1,X2),u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_39, plain, (v1_funct_1(k4_waybel_1(X1,X2))|v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_40, plain, (m1_subset_1(k4_waybel_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_41, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k6_waybel_9(X1,X1,k4_waybel_1(X1,X3),X2)=k3_waybel_2(X1,X3,X2)|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_42, plain, ![X25, X26, X27, X28]:(v1_xboole_0(X25)|(v2_struct_0(X26)|~l1_orders_2(X26)|(~v1_funct_1(X27)|(~v1_funct_2(X27,X25,u1_struct_0(X26))|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,u1_struct_0(X26)))))|~m1_subset_1(X28,X25)))|k2_yellow_6(X25,X26,X27,X28)=k1_funct_1(X27,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k2_yellow_6])])])).
fof(c_0_43, plain, ![X29, X30, X31, X32]:(v1_xboole_0(X29)|(~v1_funct_1(X31)|(~v1_funct_2(X31,X29,X30)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30))))|~m1_subset_1(X32,X29))|k3_funct_2(X29,X30,X31,X32)=k1_funct_1(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_44, plain, (k3_funct_2(u1_struct_0(X2),u1_struct_0(X2),X1,X4)=k11_lattice3(X2,X3,X4)|v2_struct_0(X2)|X1!=k4_waybel_1(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2))))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(k11_yellow_6(esk1_0,esk2_0),u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32]), c_0_33])]), c_0_34]), c_0_35])).
cnf(c_0_46, plain, (l1_pre_topc(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_47, plain, ![X33, X34, X35]:(~v5_orders_2(X33)|(~v2_lattice3(X33)|~l1_orders_2(X33))|(~m1_subset_1(X34,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33)))|k12_lattice3(X33,X34,X35)=k11_lattice3(X33,X34,X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_48, plain, (v2_struct_0(X2)|r2_hidden(k2_yellow_6(u1_struct_0(X1),X1,X3,k11_yellow_6(X1,X2)),k10_yellow_6(X1,k6_waybel_9(X1,X1,X3,X2)))|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~l1_waybel_9(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~v3_yellow_6(X2,X1)|~l1_waybel_0(X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v5_pre_topc(X3,X1,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, negated_conjecture, (v5_pre_topc(k4_waybel_1(esk1_0,esk3_0),esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_50, negated_conjecture, (v1_funct_2(k4_waybel_1(esk1_0,esk3_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_23])]), c_0_35])).
cnf(c_0_51, negated_conjecture, (v1_funct_1(k4_waybel_1(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_38]), c_0_23])]), c_0_35])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(k4_waybel_1(esk1_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_38]), c_0_23])]), c_0_35])).
cnf(c_0_53, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_54, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_55, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_56, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_57, negated_conjecture, (k6_waybel_9(esk1_0,esk1_0,k4_waybel_1(esk1_0,X1),esk2_0)=k3_waybel_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_29]), c_0_23])]), c_0_34]), c_0_35])).
cnf(c_0_58, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|k2_yellow_6(X1,X2,X3,X4)=k1_funct_1(X3,X4)|~l1_orders_2(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_59, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_60, plain, (k3_funct_2(u1_struct_0(X1),u1_struct_0(X1),k4_waybel_1(X1,X2),X3)=k11_lattice3(X1,X2,X3)|v2_struct_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_44]), c_0_40]), c_0_39]), c_0_37])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(k11_yellow_6(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_19])])).
cnf(c_0_62, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_63, negated_conjecture, (r2_hidden(k2_yellow_6(u1_struct_0(esk1_0),esk1_0,k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,X1)),k10_yellow_6(esk1_0,k6_waybel_9(esk1_0,esk1_0,k4_waybel_1(esk1_0,esk3_0),X1)))|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v3_yellow_6(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_51]), c_0_52]), c_0_19]), c_0_53]), c_0_22]), c_0_54]), c_0_55]), c_0_56]), c_0_32]), c_0_33])])).
cnf(c_0_64, negated_conjecture, (k6_waybel_9(esk1_0,esk1_0,k4_waybel_1(esk1_0,esk3_0),esk2_0)=k3_waybel_2(esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_38])).
cnf(c_0_65, negated_conjecture, (k2_yellow_6(u1_struct_0(esk1_0),esk1_0,k4_waybel_1(esk1_0,esk3_0),X1)=k1_funct_1(k4_waybel_1(esk1_0,esk3_0),X1)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_50]), c_0_51]), c_0_23]), c_0_52])]), c_0_35])).
cnf(c_0_66, negated_conjecture, (k1_funct_1(k4_waybel_1(esk1_0,esk3_0),X1)=k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k4_waybel_1(esk1_0,esk3_0),X1)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_50]), c_0_51]), c_0_52])])).
cnf(c_0_67, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k4_waybel_1(esk1_0,X1),k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,X1,k11_yellow_6(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_23])]), c_0_35])).
cnf(c_0_68, negated_conjecture, (k12_lattice3(esk1_0,X1,k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,X1,k11_yellow_6(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_61]), c_0_23]), c_0_53]), c_0_54])])).
fof(c_0_69, plain, ![X36]:(v2_struct_0(X36)|~l1_orders_2(X36)|~v1_xboole_0(k2_struct_0(X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_yellow_0])])])).
fof(c_0_70, plain, ![X37]:(~l1_struct_0(X37)|k2_struct_0(X37)=u1_struct_0(X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_71, negated_conjecture, (r2_hidden(k2_yellow_6(u1_struct_0(esk1_0),esk1_0,k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_28]), c_0_64]), c_0_29]), c_0_30]), c_0_31])]), c_0_34])).
cnf(c_0_72, negated_conjecture, (k2_yellow_6(u1_struct_0(esk1_0),esk1_0,k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))=k1_funct_1(k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_61])).
cnf(c_0_73, negated_conjecture, (k1_funct_1(k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))=k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_66, c_0_61])).
cnf(c_0_74, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_67, c_0_38])).
cnf(c_0_75, negated_conjecture, (~r2_hidden(k12_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_76, negated_conjecture, (k12_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_68, c_0_38])).
cnf(c_0_77, plain, (v2_struct_0(X1)|~l1_orders_2(X1)|~v1_xboole_0(k2_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_78, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_79, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k1_funct_1(k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_80, negated_conjecture, (k1_funct_1(k4_waybel_1(esk1_0,esk3_0),k11_yellow_6(esk1_0,esk2_0))=k11_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_81, negated_conjecture, (~r2_hidden(k11_lattice3(esk1_0,esk3_0,k11_yellow_6(esk1_0,esk2_0)),k10_yellow_6(esk1_0,k3_waybel_2(esk1_0,esk3_0,esk2_0)))), inference(rw,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_82, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_83, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])).
fof(c_0_84, plain, ![X38]:(~l1_pre_topc(X38)|l1_struct_0(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_85, negated_conjecture, (~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_23])]), c_0_35])).
cnf(c_0_86, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_87, negated_conjecture, (~l1_pre_topc(esk1_0)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_88, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_46]), c_0_19])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 89
# Proof object clause steps            : 60
# Proof object formula steps           : 29
# Proof object conjectures             : 45
# Proof object clause conjectures      : 42
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 70
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 34
# Processed clauses                    : 126
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 126
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 17
# Generated clauses                    : 92
# ...of the previous two non-trivial   : 94
# Contextual simplify-reflections      : 3
# Paramodulations                      : 91
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 74
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 36
# ...number of literals in the above   : 111
# Current number of archived formulas  : 0
# Current number of archived clauses   : 51
# Clause-clause subsumption calls (NU) : 1059
# Rec. Clause-clause subsumption calls : 286
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 80
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 31
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7050

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.007 s
# Total time               : 0.013 s
# Maximum resident set size: 28736 pages
