# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2))))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', t8_waybel_7)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', l3_subset_1)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', t2_tarski)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', d1_xboole_0)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', d20_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', t4_subset)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', d7_subset_1)).
fof(t44_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,k3_yellow_0(X1),X2))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', t44_yellow_0)).
fof(fc14_subset_1, axiom, ![X1]:~(v1_subset_1(k2_subset_1(X1),X1)), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', fc14_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', d4_subset_1)).
fof(dt_k3_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', dt_k3_yellow_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/846894ba1d0b41a5a404ae35bdf91841.E_input', t2_subset)).
fof(c_0_13, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t8_waybel_7])).
fof(c_0_14, plain, ![X24, X25, X26]:(~m1_subset_1(X25,k1_zfmisc_1(X24))|(~r2_hidden(X26,X25)|r2_hidden(X26,X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_15, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_yellow_0(esk1_0)&l1_orders_2(esk1_0))))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v1_subset_1(esk2_0,u1_struct_0(esk1_0))|r2_hidden(k3_yellow_0(esk1_0),esk2_0))&(v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
cnf(c_0_16, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, plain, ![X34, X35]:((~r2_hidden(esk6_2(X34,X35),X34)|~r2_hidden(esk6_2(X34,X35),X35)|X34=X35)&(r2_hidden(esk6_2(X34,X35),X34)|r2_hidden(esk6_2(X34,X35),X35)|X34=X35)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
fof(c_0_19, plain, ![X32, X33]:(((~m1_subset_1(X33,X32)|r2_hidden(X33,X32)|v1_xboole_0(X32))&(~r2_hidden(X33,X32)|m1_subset_1(X33,X32)|v1_xboole_0(X32)))&((~m1_subset_1(X33,X32)|v1_xboole_0(X33)|~v1_xboole_0(X32))&(~v1_xboole_0(X33)|m1_subset_1(X33,X32)|~v1_xboole_0(X32)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_20, plain, ![X20, X21, X22]:((~v1_xboole_0(X20)|~r2_hidden(X21,X20))&(r2_hidden(esk5_1(X22),X22)|v1_xboole_0(X22))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_21, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, plain, (r2_hidden(esk6_2(X1,X2),X1)|r2_hidden(esk6_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X7, X8, X9, X10]:((~v13_waybel_0(X8,X7)|(~m1_subset_1(X9,u1_struct_0(X7))|(~m1_subset_1(X10,u1_struct_0(X7))|(~r2_hidden(X9,X8)|~r1_orders_2(X7,X9,X10)|r2_hidden(X10,X8))))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|~l1_orders_2(X7))&((m1_subset_1(esk3_2(X7,X8),u1_struct_0(X7))|v13_waybel_0(X8,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|~l1_orders_2(X7))&((m1_subset_1(esk4_2(X7,X8),u1_struct_0(X7))|v13_waybel_0(X8,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|~l1_orders_2(X7))&(((r2_hidden(esk3_2(X7,X8),X8)|v13_waybel_0(X8,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|~l1_orders_2(X7))&(r1_orders_2(X7,esk3_2(X7,X8),esk4_2(X7,X8))|v13_waybel_0(X8,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|~l1_orders_2(X7)))&(~r2_hidden(esk4_2(X7,X8),X8)|v13_waybel_0(X8,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|~l1_orders_2(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
fof(c_0_24, plain, ![X27, X28, X29]:(~r2_hidden(X27,X28)|~m1_subset_1(X28,k1_zfmisc_1(X29))|m1_subset_1(X27,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_25, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, negated_conjecture, (esk2_0=X1|r2_hidden(esk6_2(esk2_0,X1),u1_struct_0(esk1_0))|r2_hidden(esk6_2(esk2_0,X1),X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(esk6_2(esk2_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(ef,[status(thm)],[c_0_27])).
fof(c_0_32, plain, ![X17, X18]:((~v1_subset_1(X18,X17)|X18!=X17|~m1_subset_1(X18,k1_zfmisc_1(X17)))&(X18=X17|v1_subset_1(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(X17)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_33, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|m1_subset_1(esk6_2(esk2_0,u1_struct_0(esk1_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_36, plain, (X1=X2|~r2_hidden(esk6_2(X1,X2),X1)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_37, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_38, plain, ![X30, X31]:(v2_struct_0(X30)|(~v5_orders_2(X30)|(~v1_yellow_0(X30)|~l1_orders_2(X30)))|(~m1_subset_1(X31,u1_struct_0(X30))|r1_orders_2(X30,k3_yellow_0(X30),X31))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t44_yellow_0])])])])).
cnf(c_0_39, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(esk6_2(esk2_0,u1_struct_0(esk1_0)),X1)|~r1_orders_2(esk1_0,X2,esk6_2(esk2_0,u1_struct_0(esk1_0)))|~r2_hidden(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v13_waybel_0(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_40, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_41, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r2_hidden(esk6_2(esk2_0,u1_struct_0(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_31])).
cnf(c_0_42, negated_conjecture, (r2_hidden(k3_yellow_0(esk1_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_43, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_37, c_0_17])).
cnf(c_0_44, plain, (v2_struct_0(X1)|r1_orders_2(X1,k3_yellow_0(X1),X2)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_46, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_47, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_48, plain, ![X13]:~v1_subset_1(k2_subset_1(X13),X13), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_subset_1])])).
fof(c_0_49, plain, ![X16]:k2_subset_1(X16)=X16, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_50, plain, ![X19]:(~l1_orders_2(X19)|m1_subset_1(k3_yellow_0(X19),u1_struct_0(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_yellow_0])])).
cnf(c_0_51, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r1_orders_2(esk1_0,X1,esk6_2(esk2_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_17])]), c_0_41])).
cnf(c_0_52, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_53, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_orders_2(esk1_0,k3_yellow_0(esk1_0),esk6_2(esk2_0,u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_34]), c_0_35]), c_0_45]), c_0_46])]), c_0_47])).
cnf(c_0_54, plain, (~v1_subset_1(k2_subset_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_56, plain, ![X14, X15]:(~m1_subset_1(X14,X15)|(v1_xboole_0(X15)|r2_hidden(X14,X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_57, plain, (m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53])).
cnf(c_0_59, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_60, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_61, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(k3_yellow_0(esk1_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_35])])).
cnf(c_0_63, negated_conjecture, (~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_58]), c_0_60])).
cnf(c_0_64, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_65, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_64]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 39
# Proof object formula steps           : 27
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 34
# Processed clauses                    : 3535
# ...of these trivial                  : 1
# ...subsumed                          : 2682
# ...remaining for further processing  : 852
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 349
# Generated clauses                    : 15369
# ...of the previous two non-trivial   : 14640
# Contextual simplify-reflections      : 37
# Paramodulations                      : 15332
# Factorizations                       : 32
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 431
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 414
# Current number of unprocessed clauses: 11081
# ...number of literals in the above   : 85256
# Current number of archived formulas  : 0
# Current number of archived clauses   : 421
# Clause-clause subsumption calls (NU) : 81031
# Rec. Clause-clause subsumption calls : 14135
# Non-unit clause-clause subsumptions  : 2022
# Unit Clause-clause subsumption calls : 365
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 348452

# -------------------------------------------------
# User time                : 0.295 s
# System time              : 0.000 s
# Total time               : 0.295 s
# Maximum resident set size: 28752 pages
