# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t46_setfam_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>~((~(X2=k1_xboole_0)&k7_setfam_1(X1,X2)=k1_xboole_0))), file('data/MPTP/proofs/84d486c9ab744068bd372880c92b23e1.E_input', t46_setfam_1)).
fof(d8_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(X3=k7_setfam_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(r2_hidden(X4,X3)<=>r2_hidden(k3_subset_1(X1,X4),X2)))))), file('data/MPTP/proofs/84d486c9ab744068bd372880c92b23e1.E_input', d8_setfam_1)).
fof(involutiveness_k7_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k7_setfam_1(X1,k7_setfam_1(X1,X2))=X2), file('data/MPTP/proofs/84d486c9ab744068bd372880c92b23e1.E_input', involutiveness_k7_setfam_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/84d486c9ab744068bd372880c92b23e1.E_input', t4_subset_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/84d486c9ab744068bd372880c92b23e1.E_input', l3_subset_1)).
fof(c_0_5, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>~((~(X2=k1_xboole_0)&k7_setfam_1(X1,X2)=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t46_setfam_1])).
fof(c_0_6, plain, ![X7, X8, X9, X10]:(((~r2_hidden(X10,X9)|r2_hidden(k3_subset_1(X7,X10),X8)|~m1_subset_1(X10,k1_zfmisc_1(X7))|X9!=k7_setfam_1(X7,X8)|~m1_subset_1(X9,k1_zfmisc_1(k1_zfmisc_1(X7)))|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(X7))))&(~r2_hidden(k3_subset_1(X7,X10),X8)|r2_hidden(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(X7))|X9!=k7_setfam_1(X7,X8)|~m1_subset_1(X9,k1_zfmisc_1(k1_zfmisc_1(X7)))|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(X7)))))&((m1_subset_1(esk3_3(X7,X8,X9),k1_zfmisc_1(X7))|X9=k7_setfam_1(X7,X8)|~m1_subset_1(X9,k1_zfmisc_1(k1_zfmisc_1(X7)))|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(X7))))&((~r2_hidden(esk3_3(X7,X8,X9),X9)|~r2_hidden(k3_subset_1(X7,esk3_3(X7,X8,X9)),X8)|X9=k7_setfam_1(X7,X8)|~m1_subset_1(X9,k1_zfmisc_1(k1_zfmisc_1(X7)))|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(X7))))&(r2_hidden(esk3_3(X7,X8,X9),X9)|r2_hidden(k3_subset_1(X7,esk3_3(X7,X8,X9)),X8)|X9=k7_setfam_1(X7,X8)|~m1_subset_1(X9,k1_zfmisc_1(k1_zfmisc_1(X7)))|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_setfam_1])])])])])).
fof(c_0_7, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))&(esk2_0!=k1_xboole_0&k7_setfam_1(esk1_0,esk2_0)=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])).
fof(c_0_8, plain, ![X12, X13]:(~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(X12)))|k7_setfam_1(X12,k7_setfam_1(X12,X13))=X13), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k7_setfam_1])])).
cnf(c_0_9, plain, (r2_hidden(X2,X4)|~r2_hidden(k3_subset_1(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(X1))|X4!=k7_setfam_1(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (k7_setfam_1(esk1_0,esk2_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_12, plain, ![X14]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X14)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_13, plain, (m1_subset_1(esk3_3(X1,X2,X3),k1_zfmisc_1(X1))|X3=k7_setfam_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, plain, (k7_setfam_1(X2,k7_setfam_1(X2,X1))=X1|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, negated_conjecture, (r2_hidden(X1,X2)|X2!=k1_xboole_0|~r2_hidden(k3_subset_1(esk1_0,X1),esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])])).
cnf(c_0_16, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, negated_conjecture, (k7_setfam_1(esk1_0,X1)=esk2_0|m1_subset_1(esk3_3(esk1_0,X1,esk2_0),k1_zfmisc_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(spm,[status(thm)],[c_0_13, c_0_11])).
cnf(c_0_18, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_19, plain, (r2_hidden(esk3_3(X1,X2,X3),X3)|r2_hidden(k3_subset_1(X1,esk3_3(X1,X2,X3)),X2)|X3=k7_setfam_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_20, plain, ![X15, X16, X17]:(~m1_subset_1(X16,k1_zfmisc_1(X15))|(~r2_hidden(X17,X16)|r2_hidden(X17,X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_21, plain, (r2_hidden(k3_subset_1(X3,X1),X4)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X3))|X2!=k7_setfam_1(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22, negated_conjecture, (k7_setfam_1(esk1_0,k1_xboole_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_11]), c_0_10])).
cnf(c_0_23, negated_conjecture, (r2_hidden(X1,k1_xboole_0)|~r2_hidden(k3_subset_1(esk1_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_16])])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk3_3(esk1_0,esk2_0,esk2_0),k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_11]), c_0_10]), c_0_18])).
cnf(c_0_25, negated_conjecture, (k7_setfam_1(esk1_0,X1)=esk2_0|r2_hidden(k3_subset_1(esk1_0,esk3_3(esk1_0,X1,esk2_0)),X1)|r2_hidden(esk3_3(esk1_0,X1,esk2_0),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(spm,[status(thm)],[c_0_19, c_0_11])).
cnf(c_0_26, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (X3=k7_setfam_1(X1,X2)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~r2_hidden(k3_subset_1(X1,esk3_3(X1,X2,X3)),X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_28, negated_conjecture, (r2_hidden(k3_subset_1(esk1_0,X1),k1_xboole_0)|X2!=esk2_0|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_16])])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk3_3(esk1_0,esk2_0,esk2_0),k1_xboole_0)|~r2_hidden(k3_subset_1(esk1_0,esk3_3(esk1_0,esk2_0,esk2_0)),esk2_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k3_subset_1(esk1_0,esk3_3(esk1_0,esk2_0,esk2_0)),esk2_0)|r2_hidden(esk3_3(esk1_0,esk2_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_11]), c_0_10]), c_0_18])).
cnf(c_0_31, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_26, c_0_16])).
cnf(c_0_32, negated_conjecture, (k7_setfam_1(esk1_0,X1)=esk2_0|~r2_hidden(k3_subset_1(esk1_0,esk3_3(esk1_0,X1,esk2_0)),X1)|~r2_hidden(esk3_3(esk1_0,X1,esk2_0),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(spm,[status(thm)],[c_0_27, c_0_11])).
cnf(c_0_33, negated_conjecture, (r2_hidden(k3_subset_1(esk1_0,X1),k1_xboole_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_28]), c_0_11])])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk3_3(esk1_0,esk2_0,esk2_0),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
cnf(c_0_35, negated_conjecture, (~r2_hidden(k3_subset_1(esk1_0,esk3_3(esk1_0,esk2_0,esk2_0)),esk2_0)|~r2_hidden(esk3_3(esk1_0,esk2_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_11]), c_0_10]), c_0_18])).
cnf(c_0_36, negated_conjecture, (r2_hidden(k3_subset_1(esk1_0,esk3_3(esk1_0,esk2_0,esk2_0)),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_24])])).
cnf(c_0_37, negated_conjecture, (~r2_hidden(k3_subset_1(esk1_0,esk3_3(esk1_0,esk2_0,esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_34])])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k3_subset_1(esk1_0,esk3_3(esk1_0,esk2_0,esk2_0)),X1)), inference(spm,[status(thm)],[c_0_31, c_0_36])).
cnf(c_0_39, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 29
# Proof object formula steps           : 11
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 11
# Processed clauses                    : 63
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 60
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 8
# Generated clauses                    : 88
# ...of the previous two non-trivial   : 77
# Contextual simplify-reflections      : 1
# Paramodulations                      : 81
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 52
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 42
# Current number of unprocessed clauses: 24
# ...number of literals in the above   : 106
# Current number of archived formulas  : 0
# Current number of archived clauses   : 8
# Clause-clause subsumption calls (NU) : 370
# Rec. Clause-clause subsumption calls : 229
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3309

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28744 pages
