# No SInE strategy applied
# Auto-Mode selected heuristic G_____0013_C18_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t52_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4))))))), file('data/MPTP/proofs/84ea79e234df4567bb737c8a5d02cb7d.E_input', t52_waybel_0)).
fof(t53_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))=>(r1_yellow_0(X1,X2)<=>r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/84ea79e234df4567bb737c8a5d02cb7d.E_input', t53_waybel_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/84ea79e234df4567bb737c8a5d02cb7d.E_input', d9_yellow_0)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/84ea79e234df4567bb737c8a5d02cb7d.E_input', d7_yellow_0)).
fof(c_0_4, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>(![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r1_yellow_0(X1,X4)))&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&![X5]:((v1_finset_1(X5)&m1_subset_1(X5,k1_zfmisc_1(X2)))=>~((r1_yellow_0(X1,X5)&X4=k1_yellow_0(X1,X5)))))))&![X4]:((v1_finset_1(X4)&m1_subset_1(X4,k1_zfmisc_1(X2)))=>(~(X4=k1_xboole_0)=>r2_hidden(k1_yellow_0(X1,X4),X3)))))), introduced(definition)).
fof(c_0_5, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)<=>r2_lattice3(X1,X3,X4))))))), inference(apply_def,[status(thm)],[t52_waybel_0, c_0_4])).
fof(c_0_6, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(epred1_3(X1,X2,X3)=>(r1_yellow_0(X1,X2)<=>r1_yellow_0(X1,X3))))))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[t53_waybel_0]), c_0_4])).
fof(c_0_7, plain, ![X9, X10, X11, X12]:(((r2_lattice3(X9,X10,X11)|X11!=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&(~m1_subset_1(X12,u1_struct_0(X9))|(~r2_lattice3(X9,X10,X12)|r1_orders_2(X9,X11,X12))|X11!=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9)))&((m1_subset_1(esk4_3(X9,X10,X11),u1_struct_0(X9))|~r2_lattice3(X9,X10,X11)|X11=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&((r2_lattice3(X9,X10,esk4_3(X9,X10,X11))|~r2_lattice3(X9,X10,X11)|X11=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))&(~r1_orders_2(X9,X11,esk4_3(X9,X10,X11))|~r2_lattice3(X9,X10,X11)|X11=k1_yellow_0(X9,X10)|~r1_yellow_0(X9,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~l1_orders_2(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_8, plain, ![X14, X15, X17, X18, X20, X21, X24]:(((m1_subset_1(esk5_2(X14,X15),u1_struct_0(X14))|~r1_yellow_0(X14,X15)|~l1_orders_2(X14))&((r2_lattice3(X14,X15,esk5_2(X14,X15))|~r1_yellow_0(X14,X15)|~l1_orders_2(X14))&((~m1_subset_1(X17,u1_struct_0(X14))|(~r2_lattice3(X14,X15,X17)|r1_orders_2(X14,esk5_2(X14,X15),X17))|~r1_yellow_0(X14,X15)|~l1_orders_2(X14))&((m1_subset_1(esk6_3(X14,X15,X18),u1_struct_0(X14))|~r2_lattice3(X14,X15,X18)|X18=esk5_2(X14,X15)|~m1_subset_1(X18,u1_struct_0(X14))|~r1_yellow_0(X14,X15)|~l1_orders_2(X14))&((r2_lattice3(X14,X15,esk6_3(X14,X15,X18))|~r2_lattice3(X14,X15,X18)|X18=esk5_2(X14,X15)|~m1_subset_1(X18,u1_struct_0(X14))|~r1_yellow_0(X14,X15)|~l1_orders_2(X14))&(~r1_orders_2(X14,X18,esk6_3(X14,X15,X18))|~r2_lattice3(X14,X15,X18)|X18=esk5_2(X14,X15)|~m1_subset_1(X18,u1_struct_0(X14))|~r1_yellow_0(X14,X15)|~l1_orders_2(X14)))))))&(((m1_subset_1(esk8_3(X14,X20,X21),u1_struct_0(X14))|m1_subset_1(esk7_3(X14,X20,X21),u1_struct_0(X14))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))&(((r2_lattice3(X14,X20,esk8_3(X14,X20,X21))|m1_subset_1(esk7_3(X14,X20,X21),u1_struct_0(X14))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))&(~m1_subset_1(X24,u1_struct_0(X14))|(~r2_lattice3(X14,X20,X24)|r1_orders_2(X14,esk8_3(X14,X20,X21),X24))|m1_subset_1(esk7_3(X14,X20,X21),u1_struct_0(X14))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14)))&(esk8_3(X14,X20,X21)!=X21|m1_subset_1(esk7_3(X14,X20,X21),u1_struct_0(X14))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))))&(((m1_subset_1(esk8_3(X14,X20,X21),u1_struct_0(X14))|r2_lattice3(X14,X20,esk7_3(X14,X20,X21))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))&(((r2_lattice3(X14,X20,esk8_3(X14,X20,X21))|r2_lattice3(X14,X20,esk7_3(X14,X20,X21))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))&(~m1_subset_1(X24,u1_struct_0(X14))|(~r2_lattice3(X14,X20,X24)|r1_orders_2(X14,esk8_3(X14,X20,X21),X24))|r2_lattice3(X14,X20,esk7_3(X14,X20,X21))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14)))&(esk8_3(X14,X20,X21)!=X21|r2_lattice3(X14,X20,esk7_3(X14,X20,X21))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))))&((m1_subset_1(esk8_3(X14,X20,X21),u1_struct_0(X14))|~r1_orders_2(X14,X21,esk7_3(X14,X20,X21))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))&(((r2_lattice3(X14,X20,esk8_3(X14,X20,X21))|~r1_orders_2(X14,X21,esk7_3(X14,X20,X21))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))&(~m1_subset_1(X24,u1_struct_0(X14))|(~r2_lattice3(X14,X20,X24)|r1_orders_2(X14,esk8_3(X14,X20,X21),X24))|~r1_orders_2(X14,X21,esk7_3(X14,X20,X21))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14)))&(esk8_3(X14,X20,X21)!=X21|~r1_orders_2(X14,X21,esk7_3(X14,X20,X21))|~r2_lattice3(X14,X20,X21)|~m1_subset_1(X21,u1_struct_0(X14))|r1_yellow_0(X14,X20)|~l1_orders_2(X14))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
fof(c_0_9, plain, ![X25, X26, X27, X28]:((~r2_lattice3(X25,X26,X28)|r2_lattice3(X25,X27,X28)|~m1_subset_1(X28,u1_struct_0(X25))|~epred1_3(X25,X26,X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v2_struct_0(X25)|(~v3_orders_2(X25)|(~v4_orders_2(X25)|~l1_orders_2(X25)))))&(~r2_lattice3(X25,X27,X28)|r2_lattice3(X25,X26,X28)|~m1_subset_1(X28,u1_struct_0(X25))|~epred1_3(X25,X26,X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X25)))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v2_struct_0(X25)|(~v3_orders_2(X25)|(~v4_orders_2(X25)|~l1_orders_2(X25)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])).
fof(c_0_10, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(epred1_3(esk1_0,esk2_0,esk3_0)&((~r1_yellow_0(esk1_0,esk2_0)|~r1_yellow_0(esk1_0,esk3_0))&(r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])).
cnf(c_0_11, plain, (X2=k1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~r1_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r1_orders_2(X2,esk5_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r2_lattice3(X1,X4,X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~epred1_3(X1,X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (epred1_3(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_21, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22, plain, (r2_lattice3(X1,X4,X3)|v2_struct_0(X1)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~epred1_3(X1,X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, plain, (esk5_2(X1,X2)=k1_yellow_0(X1,X3)|~r2_lattice3(X1,X2,esk4_3(X1,X3,esk5_2(X1,X2)))|~r2_lattice3(X1,X3,esk5_2(X1,X2))|~r1_yellow_0(X1,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(esk4_3(X1,X3,esk5_2(X1,X2)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
cnf(c_0_24, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,X1)|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_25, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,X1)|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_15]), c_0_17]), c_0_16]), c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_26, plain, (r2_lattice3(X1,X2,esk4_3(X1,X2,X3))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, negated_conjecture, (esk5_2(esk1_0,esk3_0)=k1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X1,esk5_2(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk4_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_18])])).
cnf(c_0_28, negated_conjecture, (X1=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,esk3_0,X1))|~r2_lattice3(esk1_0,esk3_0,X1)|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk4_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_18])])).
cnf(c_0_29, negated_conjecture, (esk5_2(esk1_0,esk3_0)=k1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk3_0,esk5_2(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk4_3(esk1_0,esk3_0,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_2(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_30, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31, negated_conjecture, (esk5_2(esk1_0,esk3_0)=k1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk3_0,esk5_2(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk5_2(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_18])])).
cnf(c_0_32, plain, (r2_lattice3(X1,X2,esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_33, negated_conjecture, (esk5_2(esk1_0,esk3_0)=k1_yellow_0(esk1_0,esk3_0)|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk5_2(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_18])])).
cnf(c_0_34, negated_conjecture, (esk5_2(esk1_0,esk3_0)=k1_yellow_0(esk1_0,esk3_0)|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_13]), c_0_18])])).
cnf(c_0_35, plain, (r1_orders_2(X2,esk8_3(X2,X3,X4),X1)|r2_lattice3(X2,X3,esk7_3(X2,X3,X4))|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_34]), c_0_18])])).
cnf(c_0_37, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)),X2)|r2_lattice3(esk1_0,X1,esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_18])])).
cnf(c_0_39, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)),X2)|r2_lattice3(esk1_0,X1,esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_42, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_40]), c_0_18])])).
cnf(c_0_43, plain, (r1_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_44, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_45, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),X1)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_36]), c_0_39])).
cnf(c_0_46, plain, (r1_yellow_0(X1,X2)|esk8_3(X1,X2,esk5_2(X1,X3))!=esk5_2(X1,X3)|~r2_lattice3(X1,X3,esk7_3(X1,X2,esk5_2(X1,X3)))|~r2_lattice3(X1,X2,esk5_2(X1,X3))|~r1_yellow_0(X1,X3)|~m1_subset_1(esk7_3(X1,X2,esk5_2(X1,X3)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_12]), c_0_13])).
cnf(c_0_47, plain, (r1_orders_2(X2,esk8_3(X2,X3,X4),X1)|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_orders_2(X2,X4,esk7_3(X2,X3,X4))|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_48, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk8_3(X1,X2,esk5_2(X1,X3)),u1_struct_0(X1))|~r2_lattice3(X1,X3,esk7_3(X1,X2,esk5_2(X1,X3)))|~r2_lattice3(X1,X2,esk5_2(X1,X3))|~r1_yellow_0(X1,X3)|~m1_subset_1(esk7_3(X1,X2,esk5_2(X1,X3)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_12]), c_0_13])).
cnf(c_0_49, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=k1_yellow_0(esk1_0,X1)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,esk4_3(esk1_0,X1,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))))|~r2_lattice3(esk1_0,X1,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk4_3(esk1_0,X1,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_45]), c_0_18])])).
cnf(c_0_50, negated_conjecture, (r1_yellow_0(esk1_0,X1)|esk8_3(esk1_0,X1,esk5_2(esk1_0,esk3_0))!=esk5_2(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X1,esk5_2(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_24]), c_0_18])])).
cnf(c_0_51, plain, (r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_52, plain, (r1_orders_2(X1,esk8_3(X1,X2,esk5_2(X1,X3)),X4)|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X3,esk7_3(X1,X2,esk5_2(X1,X3)))|~r2_lattice3(X1,X2,esk5_2(X1,X3))|~r2_lattice3(X1,X2,X4)|~r1_yellow_0(X1,X3)|~m1_subset_1(esk7_3(X1,X2,esk5_2(X1,X3)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_12]), c_0_13])).
cnf(c_0_53, negated_conjecture, (r1_yellow_0(esk1_0,X1)|m1_subset_1(esk8_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X1,esk5_2(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_24]), c_0_18])])).
cnf(c_0_54, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_55, plain, (esk5_2(X1,X2)=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(esk4_3(X1,X2,esk5_2(X1,X2)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_26]), c_0_13]), c_0_32])).
cnf(c_0_56, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_57, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~m1_subset_1(esk4_3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_28]), c_0_39])).
cnf(c_0_58, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|esk8_3(esk1_0,esk2_0,esk5_2(esk1_0,esk3_0))!=esk5_2(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,esk3_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_2(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_18])]), c_0_39])).
cnf(c_0_59, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_60, plain, (X2=esk5_2(X1,X3)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_61, plain, (r1_orders_2(X2,esk8_3(X2,X3,X4),X1)|m1_subset_1(esk7_3(X2,X3,X4),u1_struct_0(X2))|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_62, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_63, plain, (r2_lattice3(X1,X2,esk6_3(X1,X2,X3))|X3=esk5_2(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_64, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)),X2)|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_34]), c_0_18])])).
cnf(c_0_65, negated_conjecture, (r1_yellow_0(esk1_0,X1)|m1_subset_1(esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_34]), c_0_18])])).
cnf(c_0_66, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk8_3(esk1_0,esk2_0,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,esk3_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_2(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_18])]), c_0_39])).
cnf(c_0_67, plain, (esk5_2(X1,X2)=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_30]), c_0_13]), c_0_32])).
cnf(c_0_68, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,esk5_2(X1,X3)))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X3,esk7_3(X1,X2,esk5_2(X1,X3)))|~r2_lattice3(X1,X2,esk5_2(X1,X3))|~r1_yellow_0(X1,X3)|~m1_subset_1(esk7_3(X1,X2,esk5_2(X1,X3)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_12]), c_0_13])).
cnf(c_0_69, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~m1_subset_1(esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_30]), c_0_18])]), c_0_39])).
cnf(c_0_70, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|esk8_3(esk1_0,esk2_0,esk5_2(esk1_0,esk3_0))!=esk5_2(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,esk3_0))|~m1_subset_1(esk5_2(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_18])])).
cnf(c_0_71, plain, (esk8_3(X1,X2,X3)=esk5_2(X1,X4)|r1_yellow_0(X1,X2)|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|~r2_lattice3(X1,X2,esk6_3(X1,X4,esk8_3(X1,X2,X3)))|~r2_lattice3(X1,X4,esk8_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X4)|~m1_subset_1(esk6_3(X1,X4,esk8_3(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])).
cnf(c_0_72, negated_conjecture, (X1=esk5_2(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk6_3(esk1_0,esk3_0,X1))|~r2_lattice3(esk1_0,esk3_0,X1)|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk6_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_63]), c_0_18])])).
cnf(c_0_73, negated_conjecture, (esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))=esk5_2(esk1_0,X2)|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,esk6_3(esk1_0,X2,esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))))|~r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X2,esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk1_0,X2,esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_64]), c_0_18])]), c_0_65])).
cnf(c_0_74, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_18])]), c_0_36]), c_0_39])).
cnf(c_0_75, negated_conjecture, (r2_lattice3(esk1_0,X1,esk8_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X1,esk5_2(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk5_2(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_24]), c_0_18])])).
cnf(c_0_76, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=k1_yellow_0(esk1_0,esk3_0)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~m1_subset_1(esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_69, c_0_24])).
cnf(c_0_77, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_78, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))!=k1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_67]), c_0_18])]), c_0_36]), c_0_39])).
cnf(c_0_79, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,X1))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(esk6_3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_18])]), c_0_39])).
cnf(c_0_80, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|X3=esk5_2(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_81, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(esk6_3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_72]), c_0_74]), c_0_39])).
cnf(c_0_82, negated_conjecture, (r2_lattice3(esk1_0,X1,esk8_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0))|~r1_yellow_0(esk1_0,esk3_0)|~m1_subset_1(esk7_3(esk1_0,X1,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_67]), c_0_18])])).
cnf(c_0_83, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_18])]), c_0_78])).
cnf(c_0_84, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,X1))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_18])]), c_0_39])).
cnf(c_0_85, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,esk3_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_80]), c_0_18])]), c_0_74]), c_0_39])).
cnf(c_0_86, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_36]), c_0_74]), c_0_39])).
cnf(c_0_87, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,X1))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_84, c_0_24])).
cnf(c_0_88, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_89, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_24]), c_0_74]), c_0_86])).
cnf(c_0_90, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)))|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_86]), c_0_74]), c_0_78])).
cnf(c_0_91, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_18])])).
cnf(c_0_92, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_24]), c_0_90])).
cnf(c_0_93, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_62]), c_0_18])])).
cnf(c_0_94, negated_conjecture, (esk8_3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))=esk5_2(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_95, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|esk5_2(esk1_0,esk3_0)!=k1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_78, c_0_94])).
cnf(c_0_96, plain, (r1_orders_2(X1,esk8_3(X1,X2,esk5_2(X1,X3)),X4)|r2_lattice3(X1,X2,esk7_3(X1,X2,esk5_2(X1,X3)))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,esk5_2(X1,X3))|~r2_lattice3(X1,X2,X4)|~r1_yellow_0(X1,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_13])).
cnf(c_0_97, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)|esk5_2(esk1_0,esk3_0)!=k1_yellow_0(esk1_0,esk3_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_42]), c_0_36]), c_0_39])).
cnf(c_0_98, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,X1)),X2)|r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk3_0,esk5_2(esk1_0,X1)))|r1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,X1))|~r2_lattice3(esk1_0,esk3_0,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_24]), c_0_18])])).
cnf(c_0_99, negated_conjecture, (~r1_yellow_0(esk1_0,esk2_0)|~r1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_100, negated_conjecture, (r1_yellow_0(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_67]), c_0_18])]), c_0_39])).
cnf(c_0_101, negated_conjecture, (esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,X1))=esk5_2(esk1_0,X2)|r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk3_0,esk5_2(esk1_0,X1)))|r1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk3_0,esk6_3(esk1_0,X2,esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,X1))))|~r2_lattice3(esk1_0,X2,esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,X1))|~r1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk1_0,X2,esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,X1))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_98]), c_0_18])])).
cnf(c_0_102, negated_conjecture, (~r1_yellow_0(esk1_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_103, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))=esk5_2(esk1_0,X2)|r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk3_0,esk6_3(esk1_0,X2,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))))|~r2_lattice3(esk1_0,X2,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk1_0,X2,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_67]), c_0_18])]), c_0_102])).
cnf(c_0_104, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))=esk5_2(esk1_0,X2)|r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,esk6_3(esk1_0,X2,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))))|~r2_lattice3(esk1_0,X2,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk1_0,X2,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_103, c_0_24])).
cnf(c_0_105, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))=esk5_2(esk1_0,esk2_0)|r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_63]), c_0_100]), c_0_18])])).
cnf(c_0_106, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))=esk5_2(esk1_0,esk2_0)|r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_80]), c_0_100]), c_0_18])])).
cnf(c_0_107, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk7_3(esk1_0,esk3_0,X1))|r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,X1))|r1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_77]), c_0_18])])).
cnf(c_0_108, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk3_0,X1))|r1_yellow_0(esk1_0,esk3_0)|m1_subset_1(esk8_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(esk7_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_54]), c_0_18])])).
cnf(c_0_109, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_12, c_0_67])).
cnf(c_0_110, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_13, c_0_67])).
cnf(c_0_111, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))=esk5_2(esk1_0,esk2_0)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_25, c_0_106])).
cnf(c_0_112, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,X1))|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk3_0,X1))|r1_yellow_0(esk1_0,esk3_0)|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(esk7_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_107]), c_0_108])).
cnf(c_0_113, plain, (r1_orders_2(X1,esk8_3(X1,X2,k1_yellow_0(X1,X3)),X4)|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X3,esk7_3(X1,X2,k1_yellow_0(X1,X3)))|~r2_lattice3(X1,X2,k1_yellow_0(X1,X3))|~r2_lattice3(X1,X2,X4)|~r1_yellow_0(X1,X3)|~m1_subset_1(esk7_3(X1,X2,k1_yellow_0(X1,X3)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_109]), c_0_110])).
cnf(c_0_114, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1))=esk5_2(esk1_0,esk2_0)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_102]), c_0_24])).
cnf(c_0_115, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|r1_orders_2(esk1_0,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),X1)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_100]), c_0_18])]), c_0_102]), c_0_25])).
cnf(c_0_116, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,k1_yellow_0(X1,X3)))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X3,esk7_3(X1,X2,k1_yellow_0(X1,X3)))|~r2_lattice3(X1,X2,k1_yellow_0(X1,X3))|~r1_yellow_0(X1,X3)|~m1_subset_1(esk7_3(X1,X2,k1_yellow_0(X1,X3)),u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_68, c_0_67])).
cnf(c_0_117, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,X1)|~r2_lattice3(esk1_0,esk3_0,esk6_3(esk1_0,X1,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))))|~r2_lattice3(esk1_0,X1,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk1_0,X1,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_115]), c_0_18])])).
cnf(c_0_118, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|r2_lattice3(esk1_0,esk3_0,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_114]), c_0_100]), c_0_18])]), c_0_102]), c_0_25])).
cnf(c_0_119, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk6_3(esk1_0,X1,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))))|~r2_lattice3(esk1_0,X1,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk1_0,X1,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_117, c_0_24])).
cnf(c_0_120, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_25, c_0_118])).
cnf(c_0_121, negated_conjecture, (esk8_3(esk1_0,esk3_0,X1)=esk5_2(esk1_0,X2)|r1_yellow_0(esk1_0,esk3_0)|m1_subset_1(esk7_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,esk6_3(esk1_0,X2,esk8_3(esk1_0,esk3_0,X1)))|~r2_lattice3(esk1_0,X2,esk8_3(esk1_0,esk3_0,X1))|~r2_lattice3(esk1_0,esk3_0,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk1_0,X2,esk8_3(esk1_0,esk3_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_24]), c_0_18])])).
cnf(c_0_122, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,X1))|r1_yellow_0(esk1_0,esk3_0)|m1_subset_1(esk7_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_88]), c_0_18])])).
cnf(c_0_123, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk3_0,X1))|r1_yellow_0(esk1_0,esk3_0)|esk8_3(esk1_0,esk3_0,X1)!=X1|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(esk7_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_51]), c_0_18])])).
cnf(c_0_124, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(esk6_3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_63]), c_0_100]), c_0_18])]), c_0_120])).
cnf(c_0_125, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)|m1_subset_1(esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,esk5_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_108]), c_0_18])]), c_0_39])).
cnf(c_0_126, negated_conjecture, (esk8_3(esk1_0,esk3_0,X1)=esk5_2(esk1_0,esk2_0)|r1_yellow_0(esk1_0,esk3_0)|m1_subset_1(esk7_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(esk6_3(esk1_0,esk2_0,esk8_3(esk1_0,esk3_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_63]), c_0_18])]), c_0_39]), c_0_122])).
cnf(c_0_127, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)|esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,esk2_0))!=esk5_2(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,esk5_2(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_123]), c_0_18])]), c_0_39])).
cnf(c_0_128, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_80]), c_0_100]), c_0_18])]), c_0_120])).
cnf(c_0_129, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)|m1_subset_1(esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_67]), c_0_18])]), c_0_39])).
cnf(c_0_130, negated_conjecture, (esk8_3(esk1_0,esk3_0,X1)=esk5_2(esk1_0,esk2_0)|r1_yellow_0(esk1_0,esk3_0)|m1_subset_1(esk7_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_80]), c_0_18])]), c_0_39]), c_0_122])).
cnf(c_0_131, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)|esk8_3(esk1_0,esk3_0,esk5_2(esk1_0,esk2_0))!=esk5_2(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_59]), c_0_18])])).
cnf(c_0_132, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_102])).
cnf(c_0_133, negated_conjecture, (esk8_3(esk1_0,esk3_0,X1)=esk5_2(esk1_0,esk2_0)|r1_yellow_0(esk1_0,esk3_0)|m1_subset_1(esk7_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_62]), c_0_18])])).
cnf(c_0_134, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)|esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_67]), c_0_18])]), c_0_39])).
cnf(c_0_135, negated_conjecture, (esk8_3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_102])).
cnf(c_0_136, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk3_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_102])).
cnf(c_0_137, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,esk2_0)|~r2_lattice3(esk1_0,esk2_0,k1_yellow_0(esk1_0,esk2_0))|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_136, c_0_24])).
cnf(c_0_138, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_67])).
cnf(c_0_139, negated_conjecture, (esk5_2(esk1_0,esk2_0)!=k1_yellow_0(esk1_0,esk2_0)|~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_100]), c_0_18])])).
cnf(c_0_140, negated_conjecture, (~m1_subset_1(k1_yellow_0(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_67]), c_0_100]), c_0_18])])).
cnf(c_0_141, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_110]), c_0_100]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 142
# Proof object clause steps            : 131
# Proof object formula steps           : 11
# Proof object conjectures             : 95
# Proof object clause conjectures      : 92
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 96
# Proof object simplifying inferences  : 199
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 88
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 88
# Processed clauses                    : 6059
# ...of these trivial                  : 110
# ...subsumed                          : 3745
# ...remaining for further processing  : 2204
# Other redundant clauses eliminated   : 14
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 327
# Backward-rewritten                   : 115
# Generated clauses                    : 11114
# ...of the previous two non-trivial   : 8623
# Contextual simplify-reflections      : 792
# Paramodulations                      : 11089
# Factorizations                       : 0
# Equation resolutions                 : 25
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1760
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 1750
# Current number of unprocessed clauses: 2145
# ...number of literals in the above   : 30962
# Current number of archived formulas  : 0
# Current number of archived clauses   : 442
# Clause-clause subsumption calls (NU) : 1688222
# Rec. Clause-clause subsumption calls : 19377
# Non-unit clause-clause subsumptions  : 4817
# Unit Clause-clause subsumption calls : 2650
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 680880

# -------------------------------------------------
# User time                : 1.032 s
# System time              : 0.007 s
# Total time               : 1.038 s
# Maximum resident set size: 28832 pages
