# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#
# Preprocessing time       : 0.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t139_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(X1=k1_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>r1_funct_2(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),X5,k10_tmap_1(X1,X2,X3,X4,k2_tmap_1(X1,X2,X5,X3),k2_tmap_1(X1,X2,X5,X4))))))))), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', t139_tmap_1)).
fof(dt_k2_tmap_1, axiom, ![X1, X2, X3, X4]:((l1_struct_0(X1)&(l1_struct_0(X2)&((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))&l1_struct_0(X4))))=>(v1_funct_1(k2_tmap_1(X1,X2,X3,X4))&(v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', dt_k2_tmap_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', dt_l1_pre_topc)).
fof(dt_k10_tmap_1, axiom, ![X1, X2, X3, X4, X5, X6]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))&((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))&((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))&(v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', dt_k10_tmap_1)).
fof(d4_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_pre_topc(X4,X1)=>k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4)))))), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', d4_tmap_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', dt_m1_pre_topc)).
fof(d5_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:(m1_pre_topc(X3,X1)=>![X4]:(m1_pre_topc(X4,X1)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>(m1_pre_topc(X4,X3)=>k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4)))))))), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', d5_tmap_1)).
fof(t138_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))=>r2_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),X5,k10_tmap_1(X1,X2,X3,X4,k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5)))))))), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', t138_tmap_1)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', t2_tsep_1)).
fof(redefinition_r1_funct_2, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X2))&(~(v1_xboole_0(X4))&((v1_funct_1(X5)&(v1_funct_2(X5,X1,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X3,X4)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))))=>(r1_funct_2(X1,X2,X3,X4,X5,X6)<=>X5=X6)), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', redefinition_r1_funct_2)).
fof(redefinition_r2_funct_2, axiom, ![X1, X2, X3, X4]:(((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))=>(r2_funct_2(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', redefinition_r2_funct_2)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/8526bf124f3e4816a414b51bf1915b58.E_input', fc2_struct_0)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(X1=k1_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>r1_funct_2(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),X5,k10_tmap_1(X1,X2,X3,X4,k2_tmap_1(X1,X2,X5,X3),k2_tmap_1(X1,X2,X5,X4)))))))))), inference(assume_negation,[status(cth)],[t139_tmap_1])).
fof(c_0_13, plain, ![X17, X18, X19, X20]:((v1_funct_1(k2_tmap_1(X17,X18,X19,X20))|(~l1_struct_0(X17)|(~l1_struct_0(X18)|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18)))))|~l1_struct_0(X20)))))&((v1_funct_2(k2_tmap_1(X17,X18,X19,X20),u1_struct_0(X20),u1_struct_0(X18))|(~l1_struct_0(X17)|(~l1_struct_0(X18)|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18)))))|~l1_struct_0(X20)))))&(m1_subset_1(k2_tmap_1(X17,X18,X19,X20),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(X18))))|(~l1_struct_0(X17)|(~l1_struct_0(X18)|(~v1_funct_1(X19)|(~v1_funct_2(X19,u1_struct_0(X17),u1_struct_0(X18))|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X17),u1_struct_0(X18)))))|~l1_struct_0(X20))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_tmap_1])])])).
fof(c_0_14, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&(esk1_0=k1_tsep_1(esk1_0,esk3_0,esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
cnf(c_0_15, plain, (m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_struct_0(X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (v1_funct_2(esk5_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, ![X12]:(~l1_pre_topc(X12)|l1_struct_0(X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_20, plain, (v1_funct_1(k2_tmap_1(X1,X2,X3,X4))|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_struct_0(X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_struct_0(X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_23, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25, negated_conjecture, (v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1))|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_26, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_struct_0(esk2_0)|~l1_struct_0(esk1_0)|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_16]), c_0_17]), c_0_18])])).
fof(c_0_27, plain, ![X39, X40, X41, X42, X43, X44]:((v1_funct_1(k10_tmap_1(X39,X40,X41,X42,X43,X44))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))|(v2_struct_0(X41)|~m1_pre_topc(X41,X39)|(v2_struct_0(X42)|~m1_pre_topc(X42,X39)|(~v1_funct_1(X43)|(~v1_funct_2(X43,u1_struct_0(X41),u1_struct_0(X40))|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X40)))))|(~v1_funct_1(X44)|(~v1_funct_2(X44,u1_struct_0(X42),u1_struct_0(X40))|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X42),u1_struct_0(X40))))))))))))&((v1_funct_2(k10_tmap_1(X39,X40,X41,X42,X43,X44),u1_struct_0(k1_tsep_1(X39,X41,X42)),u1_struct_0(X40))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))|(v2_struct_0(X41)|~m1_pre_topc(X41,X39)|(v2_struct_0(X42)|~m1_pre_topc(X42,X39)|(~v1_funct_1(X43)|(~v1_funct_2(X43,u1_struct_0(X41),u1_struct_0(X40))|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X40)))))|(~v1_funct_1(X44)|(~v1_funct_2(X44,u1_struct_0(X42),u1_struct_0(X40))|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X42),u1_struct_0(X40))))))))))))&(m1_subset_1(k10_tmap_1(X39,X40,X41,X42,X43,X44),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X39,X41,X42)),u1_struct_0(X40))))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))|(v2_struct_0(X41)|~m1_pre_topc(X41,X39)|(v2_struct_0(X42)|~m1_pre_topc(X42,X39)|(~v1_funct_1(X43)|(~v1_funct_2(X43,u1_struct_0(X41),u1_struct_0(X40))|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X40)))))|(~v1_funct_1(X44)|(~v1_funct_2(X44,u1_struct_0(X42),u1_struct_0(X40))|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X42),u1_struct_0(X40)))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_tmap_1])])])])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~l1_struct_0(esk1_0)|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])])).
cnf(c_0_29, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_30, negated_conjecture, (v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1))|~l1_struct_0(esk1_0)|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_23]), c_0_24])])).
cnf(c_0_31, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_struct_0(esk1_0)|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_23]), c_0_24])])).
cnf(c_0_32, plain, (m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_23]), c_0_29])])).
cnf(c_0_34, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36, negated_conjecture, (v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1))|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_23]), c_0_29])])).
cnf(c_0_37, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),u1_struct_0(X1),u1_struct_0(esk2_0))|~l1_struct_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_23]), c_0_29])])).
cnf(c_0_38, plain, (v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39, plain, (v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_1(X6)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(k10_tmap_1(X1,esk2_0,X2,X3,X4,k2_tmap_1(esk1_0,esk2_0,esk5_0,X3)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(esk2_0))))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_24]), c_0_34])]), c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_41, negated_conjecture, (v1_funct_2(k10_tmap_1(X1,esk2_0,X2,X3,X4,k2_tmap_1(esk1_0,esk2_0,esk5_0,X3)),u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(esk2_0))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_33]), c_0_24]), c_0_34])]), c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_42, negated_conjecture, (v1_funct_1(k10_tmap_1(X1,esk2_0,X2,X3,X4,k2_tmap_1(esk1_0,esk2_0,esk5_0,X3)))|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_struct_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X4)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_33]), c_0_24]), c_0_34])]), c_0_35]), c_0_36]), c_0_37])).
fof(c_0_43, plain, ![X30, X31, X32, X33]:(v2_struct_0(X30)|(~v2_pre_topc(X30)|~l1_pre_topc(X30))|(v2_struct_0(X31)|(~v2_pre_topc(X31)|~l1_pre_topc(X31))|(~v1_funct_1(X32)|(~v1_funct_2(X32,u1_struct_0(X30),u1_struct_0(X31))|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X30),u1_struct_0(X31)))))|(~m1_pre_topc(X33,X30)|k2_tmap_1(X30,X31,X32,X33)=k2_partfun1(u1_struct_0(X30),u1_struct_0(X31),X32,u1_struct_0(X33)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tmap_1])])])])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(k10_tmap_1(X1,esk2_0,X2,X3,k2_tmap_1(esk1_0,esk2_0,esk5_0,X2),k2_tmap_1(esk1_0,esk2_0,esk5_0,X3)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(esk2_0))))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(X3)|~l1_struct_0(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_33]), c_0_36]), c_0_37])).
cnf(c_0_45, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_46, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_47, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_48, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_49, plain, ![X28, X29]:(~l1_pre_topc(X28)|(~m1_pre_topc(X29,X28)|l1_pre_topc(X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_50, negated_conjecture, (v1_funct_2(k10_tmap_1(X1,esk2_0,X2,X3,k2_tmap_1(esk1_0,esk2_0,esk5_0,X2),k2_tmap_1(esk1_0,esk2_0,esk5_0,X3)),u1_struct_0(k1_tsep_1(X1,X2,X3)),u1_struct_0(esk2_0))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(X3)|~l1_struct_0(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_33]), c_0_36]), c_0_37])).
cnf(c_0_51, negated_conjecture, (v1_funct_1(k10_tmap_1(X1,esk2_0,X2,X3,k2_tmap_1(esk1_0,esk2_0,esk5_0,X2),k2_tmap_1(esk1_0,esk2_0,esk5_0,X3)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_struct_0(X3)|~l1_struct_0(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_33]), c_0_36]), c_0_37])).
fof(c_0_52, plain, ![X34, X35, X36, X37, X38]:(v2_struct_0(X34)|(~v2_pre_topc(X34)|~l1_pre_topc(X34))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|(~m1_pre_topc(X36,X34)|(~m1_pre_topc(X37,X34)|(~v1_funct_1(X38)|(~v1_funct_2(X38,u1_struct_0(X36),u1_struct_0(X35))|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X36),u1_struct_0(X35)))))|(~m1_pre_topc(X37,X36)|k3_tmap_1(X34,X35,X36,X37,X38)=k2_partfun1(u1_struct_0(X36),u1_struct_0(X35),X38,u1_struct_0(X37)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_tmap_1])])])])).
cnf(c_0_53, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(k10_tmap_1(esk1_0,esk2_0,X1,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,X1,esk4_0)),u1_struct_0(esk2_0))))|v2_struct_0(X1)|~l1_struct_0(esk4_0)|~l1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_29]), c_0_46])]), c_0_47]), c_0_48])).
cnf(c_0_55, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_56, negated_conjecture, (esk1_0=k1_tsep_1(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_57, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_58, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, negated_conjecture, (v1_funct_2(k10_tmap_1(esk1_0,esk2_0,X1,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)),u1_struct_0(k1_tsep_1(esk1_0,X1,esk4_0)),u1_struct_0(esk2_0))|v2_struct_0(X1)|~l1_struct_0(esk4_0)|~l1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_45]), c_0_29]), c_0_46])]), c_0_47]), c_0_48])).
cnf(c_0_60, negated_conjecture, (v1_funct_1(k10_tmap_1(esk1_0,esk2_0,X1,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,X1),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)))|v2_struct_0(X1)|~l1_struct_0(esk4_0)|~l1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_45]), c_0_29]), c_0_46])]), c_0_47]), c_0_48])).
cnf(c_0_61, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k3_tmap_1(X1,X2,X3,X4,X5)=k2_partfun1(u1_struct_0(X3),u1_struct_0(X2),X5,u1_struct_0(X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_pre_topc(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, negated_conjecture, (k2_partfun1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,u1_struct_0(X1))=k2_tmap_1(esk1_0,esk2_0,esk5_0,X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_16]), c_0_17]), c_0_18]), c_0_24]), c_0_29]), c_0_34]), c_0_46])]), c_0_47]), c_0_35])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~l1_struct_0(esk4_0)|~l1_struct_0(esk3_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_57])).
cnf(c_0_64, negated_conjecture, (l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_45]), c_0_29])])).
cnf(c_0_65, negated_conjecture, (v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~l1_struct_0(esk4_0)|~l1_struct_0(esk3_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_55]), c_0_56]), c_0_57])).
cnf(c_0_66, negated_conjecture, (v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)))|~l1_struct_0(esk4_0)|~l1_struct_0(esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_55]), c_0_57])).
fof(c_0_67, plain, ![X45, X46, X47, X48, X49]:(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(v2_struct_0(X46)|(~v2_pre_topc(X46)|~l1_pre_topc(X46))|(v2_struct_0(X47)|~m1_pre_topc(X47,X45)|(v2_struct_0(X48)|~m1_pre_topc(X48,X45)|(~v1_funct_1(X49)|(~v1_funct_2(X49,u1_struct_0(k1_tsep_1(X45,X47,X48)),u1_struct_0(X46))|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X45,X47,X48)),u1_struct_0(X46)))))|r2_funct_2(u1_struct_0(k1_tsep_1(X45,X47,X48)),u1_struct_0(X46),X49,k10_tmap_1(X45,X46,X47,X48,k3_tmap_1(X45,X46,k1_tsep_1(X45,X47,X48),X47,X49),k3_tmap_1(X45,X46,k1_tsep_1(X45,X47,X48),X48,X49)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t138_tmap_1])])])])).
cnf(c_0_68, negated_conjecture, (k3_tmap_1(X1,esk2_0,esk1_0,X2,esk5_0)=k2_partfun1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,u1_struct_0(X2))|v2_struct_0(X1)|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(esk1_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_16]), c_0_17]), c_0_18]), c_0_24]), c_0_34])]), c_0_35])).
cnf(c_0_69, negated_conjecture, (k2_partfun1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,u1_struct_0(esk3_0))=k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0)), inference(spm,[status(thm)],[c_0_62, c_0_55])).
fof(c_0_70, plain, ![X27]:(~l1_pre_topc(X27)|m1_pre_topc(X27,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_71, negated_conjecture, (k2_partfun1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,u1_struct_0(esk4_0))=k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)), inference(spm,[status(thm)],[c_0_62, c_0_45])).
fof(c_0_72, plain, ![X21, X22, X23, X24, X25, X26]:((~r1_funct_2(X21,X22,X23,X24,X25,X26)|X25=X26|(v1_xboole_0(X22)|(v1_xboole_0(X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,X21,X22)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(X21,X22))))|(~v1_funct_1(X26)|(~v1_funct_2(X26,X23,X24)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X23,X24)))))))))&(X25!=X26|r1_funct_2(X21,X22,X23,X24,X25,X26)|(v1_xboole_0(X22)|(v1_xboole_0(X24)|(~v1_funct_1(X25)|(~v1_funct_2(X25,X21,X22)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(X21,X22))))|(~v1_funct_1(X26)|(~v1_funct_2(X26,X23,X24)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X23,X24)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r1_funct_2])])])])).
fof(c_0_73, plain, ![X13, X14, X15, X16]:((~r2_funct_2(X13,X14,X15,X16)|X15=X16|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))))&(X15!=X16|r2_funct_2(X13,X14,X15,X16)|(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,X13,X14)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_funct_2])])])).
cnf(c_0_74, negated_conjecture, (m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~l1_struct_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_23]), c_0_64])])).
cnf(c_0_75, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_55]), c_0_29])])).
cnf(c_0_76, negated_conjecture, (v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)),u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~l1_struct_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_23]), c_0_64])])).
cnf(c_0_77, negated_conjecture, (v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)))|~l1_struct_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_23]), c_0_64])])).
cnf(c_0_78, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r2_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),X5,k10_tmap_1(X1,X2,X3,X4,k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X5),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X5)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_79, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,esk5_0)=k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0)|~m1_pre_topc(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_55]), c_0_69]), c_0_55]), c_0_29]), c_0_46])]), c_0_47])).
cnf(c_0_80, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_81, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,esk5_0)=k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)|~m1_pre_topc(esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_45]), c_0_71]), c_0_45]), c_0_29]), c_0_46])]), c_0_47])).
cnf(c_0_82, plain, (r1_funct_2(X3,X4,X5,X6,X1,X2)|v1_xboole_0(X4)|v1_xboole_0(X6)|X1!=X2|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X2)|~v1_funct_2(X2,X5,X6)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_83, plain, (X3=X4|~r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_84, negated_conjecture, (m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_23]), c_0_75])])).
cnf(c_0_85, negated_conjecture, (v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)),u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_23]), c_0_75])])).
cnf(c_0_86, negated_conjecture, (v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_23]), c_0_75])])).
cnf(c_0_87, negated_conjecture, (r2_funct_2(u1_struct_0(esk1_0),u1_struct_0(X1),X2,k10_tmap_1(esk1_0,X1,esk3_0,esk4_0,k3_tmap_1(esk1_0,X1,esk1_0,esk3_0,X2),k3_tmap_1(esk1_0,X1,esk1_0,esk4_0,X2)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_56]), c_0_45]), c_0_55]), c_0_29]), c_0_46])]), c_0_47]), c_0_57]), c_0_48])).
cnf(c_0_88, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk3_0,esk5_0)=k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_29])])).
cnf(c_0_89, negated_conjecture, (k3_tmap_1(esk1_0,esk2_0,esk1_0,esk4_0,esk5_0)=k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_80]), c_0_29])])).
cnf(c_0_90, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_funct_2(X3,X2,X4,X1,X5,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_2(X5,X4,X1)|~v1_funct_2(X5,X3,X2)|~v1_funct_1(X5)), inference(er,[status(thm)],[c_0_82])).
cnf(c_0_91, negated_conjecture, (~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_92, negated_conjecture, (X1=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0))|~r2_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),X1,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85]), c_0_86])])).
cnf(c_0_93, negated_conjecture, (r2_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_16]), c_0_88]), c_0_89]), c_0_17]), c_0_18]), c_0_24]), c_0_34])]), c_0_35])).
cnf(c_0_94, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|v1_xboole_0(X1)|r1_funct_2(X2,X1,u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,esk5_0)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_2(esk5_0,X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_95, negated_conjecture, (~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0)))), inference(rw,[status(thm)],[c_0_91, c_0_56])).
cnf(c_0_96, negated_conjecture, (k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,k2_tmap_1(esk1_0,esk2_0,esk5_0,esk3_0),k2_tmap_1(esk1_0,esk2_0,esk5_0,esk4_0))=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_16]), c_0_93]), c_0_17]), c_0_18])])).
fof(c_0_97, plain, ![X50]:(v2_struct_0(X50)|~l1_struct_0(X50)|~v1_xboole_0(u1_struct_0(X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_98, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_16]), c_0_17])])).
cnf(c_0_99, negated_conjecture, (~r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk5_0,esk5_0)), inference(rw,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_100, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_101, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))), inference(sr,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_102, negated_conjecture, (~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_35])).
cnf(c_0_103, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_23]), c_0_24])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 104
# Proof object clause steps            : 79
# Proof object formula steps           : 25
# Proof object conjectures             : 66
# Proof object clause conjectures      : 63
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 147
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 729
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 727
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 126
# Backward-rewritten                   : 119
# Generated clauses                    : 1242
# ...of the previous two non-trivial   : 1251
# Contextual simplify-reflections      : 186
# Paramodulations                      : 1239
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 479
#    Positive orientable unit clauses  : 70
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 403
# Current number of unprocessed clauses: 554
# ...number of literals in the above   : 4473
# Current number of archived formulas  : 0
# Current number of archived clauses   : 246
# Clause-clause subsumption calls (NU) : 47493
# Rec. Clause-clause subsumption calls : 13046
# Non-unit clause-clause subsumptions  : 314
# Unit Clause-clause subsumption calls : 2205
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 709
# BW rewrite match successes           : 40
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 107102

# -------------------------------------------------
# User time                : 0.068 s
# System time              : 0.000 s
# Total time               : 0.068 s
# Maximum resident set size: 28744 pages
