# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_zfmisc_1, conjecture, ![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4)))), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t28_zfmisc_1)).
fof(l45_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_tarski(X2,X3))<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', l45_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t70_enumset1)).
fof(t94_enumset1, axiom, ![X1]:k5_enumset1(X1,X1,X1,X1,X1,X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t94_enumset1)).
fof(t89_enumset1, axiom, ![X1, X2, X3]:k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t89_enumset1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', d2_tarski)).
fof(t12_zfmisc_1, axiom, ![X1, X2]:r1_tarski(k1_tarski(X1),k2_tarski(X1,X2)), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t12_zfmisc_1)).
fof(t25_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))&(~(X1=X2)&~(X1=X3)))), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t25_zfmisc_1)).
fof(t89_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2)), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t89_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t100_xboole_1)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t4_boole)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t1_xboole_0)).
fof(t74_xboole_1, axiom, ![X1, X2, X3]:~((~(r1_xboole_0(X1,k3_xboole_0(X2,X3)))&r1_xboole_0(X1,X2))), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t74_xboole_1)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/865c02cd038a4973852d46b8fe794809.E_input', t66_xboole_1)).
fof(c_0_14, negated_conjecture, ~(![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4))))), inference(assume_negation,[status(cth)],[t28_zfmisc_1])).
fof(c_0_15, plain, ![X9, X10, X11]:((~r1_tarski(X9,k2_tarski(X10,X11))|(X9=k1_xboole_0|(X9=k1_tarski(X10)|(X9=k1_tarski(X11)|X9=k2_tarski(X10,X11)))))&((X9!=k1_xboole_0|r1_tarski(X9,k2_tarski(X10,X11)))&((X9!=k1_tarski(X10)|r1_tarski(X9,k2_tarski(X10,X11)))&((X9!=k1_tarski(X11)|r1_tarski(X9,k2_tarski(X10,X11)))&(X9!=k2_tarski(X10,X11)|r1_tarski(X9,k2_tarski(X10,X11))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l45_zfmisc_1])])])])).
fof(c_0_16, plain, ![X15, X16]:k1_enumset1(X15,X15,X16)=k2_tarski(X15,X16), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_17, plain, ![X12]:k5_enumset1(X12,X12,X12,X12,X12,X12,X12)=k1_tarski(X12), inference(variable_rename,[status(thm)],[t94_enumset1])).
fof(c_0_18, plain, ![X35, X36, X37]:k5_enumset1(X35,X35,X35,X35,X35,X36,X37)=k1_enumset1(X35,X36,X37), inference(variable_rename,[status(thm)],[t89_enumset1])).
fof(c_0_19, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))&(esk1_0!=esk3_0&esk1_0!=esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_20, plain, ![X20, X21, X22, X23, X24, X25, X26, X27]:(((~r2_hidden(X23,X22)|(X23=X20|X23=X21)|X22!=k2_tarski(X20,X21))&((X24!=X20|r2_hidden(X24,X22)|X22!=k2_tarski(X20,X21))&(X24!=X21|r2_hidden(X24,X22)|X22!=k2_tarski(X20,X21))))&(((esk5_3(X25,X26,X27)!=X25|~r2_hidden(esk5_3(X25,X26,X27),X27)|X27=k2_tarski(X25,X26))&(esk5_3(X25,X26,X27)!=X26|~r2_hidden(esk5_3(X25,X26,X27),X27)|X27=k2_tarski(X25,X26)))&(r2_hidden(esk5_3(X25,X26,X27),X27)|(esk5_3(X25,X26,X27)=X25|esk5_3(X25,X26,X27)=X26)|X27=k2_tarski(X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_21, plain, ![X17, X18]:r1_tarski(k1_tarski(X17),k2_tarski(X17,X18)), inference(variable_rename,[status(thm)],[t12_zfmisc_1])).
cnf(c_0_22, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|X1=k1_tarski(X3)|X1=k2_tarski(X2,X3)|~r1_tarski(X1,k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k5_enumset1(X1,X1,X1,X1,X1,X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_28, plain, ![X29, X30, X31]:(~r1_tarski(k1_tarski(X29),k2_tarski(X30,X31))|(X29=X30|X29=X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_zfmisc_1])])])).
cnf(c_0_29, plain, (r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (X1=k1_xboole_0|X1=k5_enumset1(X3,X3,X3,X3,X3,X3,X3)|X1=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)|X1=k5_enumset1(X2,X2,X2,X2,X2,X2,X2)|~r1_tarski(X1,k5_enumset1(X2,X2,X2,X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_23]), c_0_24]), c_0_24]), c_0_25]), c_0_25])).
cnf(c_0_31, negated_conjecture, (r1_tarski(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_23]), c_0_23]), c_0_25]), c_0_25])).
cnf(c_0_32, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k5_enumset1(X2,X2,X2,X2,X2,X2,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_23]), c_0_25])).
cnf(c_0_33, plain, (X1=X2|X1=X3|~r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (r1_tarski(k5_enumset1(X1,X1,X1,X1,X1,X1,X1),k5_enumset1(X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_35, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_37, plain, (r2_hidden(X1,X2)|X2!=k5_enumset1(X1,X1,X1,X1,X1,X1,X3)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (X1=X3|X1=X2|~r1_tarski(k5_enumset1(X1,X1,X1,X1,X1,X1,X1),k5_enumset1(X2,X2,X2,X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_23]), c_0_24]), c_0_25])).
cnf(c_0_39, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|r1_tarski(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, negated_conjecture, (esk1_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_41, plain, ![X13, X14]:r1_xboole_0(k3_xboole_0(X13,X14),k4_xboole_0(X13,X14)), inference(variable_rename,[status(thm)],[t89_xboole_1])).
fof(c_0_42, plain, ![X41, X42]:k4_xboole_0(X41,X42)=k5_xboole_0(X41,k3_xboole_0(X41,X42)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_43, plain, ![X19]:k4_xboole_0(k1_xboole_0,X19)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_44, plain, (X1=X4|X1=X3|X2!=k5_enumset1(X3,X3,X3,X3,X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_23]), c_0_25])).
cnf(c_0_45, plain, (r2_hidden(X1,k5_enumset1(X1,X1,X1,X1,X1,X1,X2))), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_46, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_47, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_50, plain, ![X32, X33, X34]:(((~r2_hidden(X32,X33)|~r2_hidden(X32,X34)|~r2_hidden(X32,k5_xboole_0(X33,X34)))&(r2_hidden(X32,X33)|r2_hidden(X32,X34)|~r2_hidden(X32,k5_xboole_0(X33,X34))))&((~r2_hidden(X32,X33)|r2_hidden(X32,X34)|r2_hidden(X32,k5_xboole_0(X33,X34)))&(~r2_hidden(X32,X34)|r2_hidden(X32,X33)|r2_hidden(X32,k5_xboole_0(X33,X34))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_51, plain, (X1=X2|X1=X3|~r2_hidden(X1,k5_enumset1(X3,X3,X3,X3,X3,X3,X2))), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_52, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|r2_hidden(esk1_0,k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, negated_conjecture, (esk1_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_54, plain, ![X38, X39, X40]:(r1_xboole_0(X38,k3_xboole_0(X39,X40))|~r1_xboole_0(X38,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t74_xboole_1])])])).
cnf(c_0_55, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_56, plain, (k5_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[c_0_49, c_0_48])).
cnf(c_0_57, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_58, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_40])).
fof(c_0_60, plain, ![X43]:((r1_xboole_0(X43,X43)|X43!=k1_xboole_0)&(X43=k1_xboole_0|~r1_xboole_0(X43,X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_61, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (r1_xboole_0(k3_xboole_0(k1_xboole_0,X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_63, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k5_enumset1(X4,X4,X4,X4,X4,X4,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_23]), c_0_25])).
cnf(c_0_64, plain, (~r2_hidden(X1,k5_xboole_0(X2,k5_enumset1(X1,X1,X1,X1,X1,X1,X3)))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_58, c_0_45])).
cnf(c_0_65, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|r2_hidden(esk1_0,k5_enumset1(esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0,esk4_0))), inference(spm,[status(thm)],[c_0_45, c_0_59])).
cnf(c_0_66, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_67, plain, (r1_xboole_0(k3_xboole_0(k1_xboole_0,X1),k3_xboole_0(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_68, plain, (r2_hidden(X1,X2)|X2!=k5_enumset1(X3,X3,X3,X3,X3,X3,X1)), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_69, plain, (~r2_hidden(X1,k5_xboole_0(k5_enumset1(X1,X1,X1,X1,X1,X1,X2),k5_enumset1(X1,X1,X1,X1,X1,X1,X3)))), inference(spm,[status(thm)],[c_0_64, c_0_45])).
cnf(c_0_70, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_65]), c_0_53])).
cnf(c_0_71, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_72, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_73, plain, (r2_hidden(X1,k5_enumset1(X2,X2,X2,X2,X2,X2,X1))), inference(er,[status(thm)],[c_0_68])).
cnf(c_0_74, negated_conjecture, (~r2_hidden(esk1_0,k5_xboole_0(k1_xboole_0,k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1)))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_75, plain, (k5_xboole_0(k1_xboole_0,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_56, c_0_71])).
cnf(c_0_76, negated_conjecture, (~r2_hidden(esk1_0,k5_xboole_0(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1),k1_xboole_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_77, plain, (r2_hidden(X1,k5_xboole_0(k5_enumset1(X2,X2,X2,X2,X2,X2,X1),X3))|r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_78, negated_conjecture, (~r2_hidden(esk1_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_70]), c_0_75])).
cnf(c_0_79, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 80
# Proof object clause steps            : 51
# Proof object formula steps           : 29
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 25
# Processed clauses                    : 2401
# ...of these trivial                  : 1
# ...subsumed                          : 1332
# ...remaining for further processing  : 1068
# Other redundant clauses eliminated   : 28
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 7
# Generated clauses                    : 32043
# ...of the previous two non-trivial   : 21985
# Contextual simplify-reflections      : 2
# Paramodulations                      : 31963
# Factorizations                       : 44
# Equation resolutions                 : 36
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1048
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 1019
# Current number of unprocessed clauses: 19602
# ...number of literals in the above   : 147564
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 1662040
# Rec. Clause-clause subsumption calls : 48244
# Non-unit clause-clause subsumptions  : 1345
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 35
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 846337

# -------------------------------------------------
# User time                : 0.995 s
# System time              : 0.007 s
# Total time               : 1.001 s
# Maximum resident set size: 26188 pages
