# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t136_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:((v6_orders_2(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>r3_orders_1(u1_orders_2(X1),X2))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', t136_orders_2)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', l3_subset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', d3_tarski)).
fof(t94_orders_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r4_relat_2(X3,X1)&r1_tarski(X2,X1))=>r4_relat_2(X3,X2))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', t94_orders_1)).
fof(t95_orders_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r8_relat_2(X3,X1)&r1_tarski(X2,X1))=>r8_relat_2(X3,X2))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', t95_orders_1)).
fof(d8_orders_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r3_orders_1(X1,X2)<=>(r1_relat_2(X1,X2)&(r8_relat_2(X1,X2)&(r4_relat_2(X1,X2)&r6_relat_2(X1,X2)))))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', d8_orders_1)).
fof(t92_orders_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r7_relat_2(X2,X1)<=>(r1_relat_2(X2,X1)&r6_relat_2(X2,X1)))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', t92_orders_1)).
fof(d11_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v6_orders_2(X2,X1)<=>r7_relat_2(u1_orders_2(X1),X2)))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', d11_orders_2)).
fof(d6_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>(v5_orders_2(X1)<=>r4_relat_2(u1_orders_2(X1),u1_struct_0(X1)))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', d6_orders_2)).
fof(d5_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>(v4_orders_2(X1)<=>r8_relat_2(u1_orders_2(X1),u1_struct_0(X1)))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', d5_orders_2)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', cc2_relat_1)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', dt_u1_orders_2)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/865fc1e8e9844a66969eb3fa40510568.E_input', fc6_relat_1)).
fof(c_0_13, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:((v6_orders_2(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>r3_orders_1(u1_orders_2(X1),X2)))), inference(assume_negation,[status(cth)],[t136_orders_2])).
fof(c_0_14, plain, ![X6, X7, X8]:(~m1_subset_1(X7,k1_zfmisc_1(X6))|(~r2_hidden(X8,X7)|r2_hidden(X8,X6))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_15, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&((v6_orders_2(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))&~r3_orders_1(u1_orders_2(esk1_0),esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
fof(c_0_16, plain, ![X9, X10, X11, X12, X13]:((~r1_tarski(X9,X10)|(~r2_hidden(X11,X9)|r2_hidden(X11,X10)))&((r2_hidden(esk3_2(X12,X13),X12)|r1_tarski(X12,X13))&(~r2_hidden(esk3_2(X12,X13),X13)|r1_tarski(X12,X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_17, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_21, plain, ![X15, X16, X17]:(~v1_relat_1(X17)|(~r4_relat_2(X17,X15)|~r1_tarski(X16,X15)|r4_relat_2(X17,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_orders_1])])).
cnf(c_0_22, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r2_hidden(esk3_2(X1,u1_struct_0(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X18, X19, X20]:(~v1_relat_1(X20)|(~r8_relat_2(X20,X18)|~r1_tarski(X19,X18)|r8_relat_2(X20,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_orders_1])])).
fof(c_0_25, plain, ![X21, X22]:(((r1_relat_2(X21,X22)|~r3_orders_1(X21,X22)|~v1_relat_1(X21))&((r8_relat_2(X21,X22)|~r3_orders_1(X21,X22)|~v1_relat_1(X21))&((r4_relat_2(X21,X22)|~r3_orders_1(X21,X22)|~v1_relat_1(X21))&(r6_relat_2(X21,X22)|~r3_orders_1(X21,X22)|~v1_relat_1(X21)))))&(~r1_relat_2(X21,X22)|(~r8_relat_2(X21,X22)|(~r4_relat_2(X21,X22)|~r6_relat_2(X21,X22)))|r3_orders_1(X21,X22)|~v1_relat_1(X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_orders_1])])])])).
fof(c_0_26, plain, ![X23, X24]:(((r1_relat_2(X24,X23)|~r7_relat_2(X24,X23)|~v1_relat_1(X24))&(r6_relat_2(X24,X23)|~r7_relat_2(X24,X23)|~v1_relat_1(X24)))&(~r1_relat_2(X24,X23)|~r6_relat_2(X24,X23)|r7_relat_2(X24,X23)|~v1_relat_1(X24))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t92_orders_1])])])).
fof(c_0_27, plain, ![X25, X26]:((~v6_orders_2(X26,X25)|r7_relat_2(u1_orders_2(X25),X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))&(~r7_relat_2(u1_orders_2(X25),X26)|v6_orders_2(X26,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_orders_2(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_orders_2])])])])).
cnf(c_0_28, plain, (r4_relat_2(X1,X3)|~v1_relat_1(X1)|~r4_relat_2(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_30, plain, ![X27]:((~v5_orders_2(X27)|r4_relat_2(u1_orders_2(X27),u1_struct_0(X27))|~l1_orders_2(X27))&(~r4_relat_2(u1_orders_2(X27),u1_struct_0(X27))|v5_orders_2(X27)|~l1_orders_2(X27))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_orders_2])])])).
cnf(c_0_31, plain, (r8_relat_2(X1,X3)|~v1_relat_1(X1)|~r8_relat_2(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, ![X28]:((~v4_orders_2(X28)|r8_relat_2(u1_orders_2(X28),u1_struct_0(X28))|~l1_orders_2(X28))&(~r8_relat_2(u1_orders_2(X28),u1_struct_0(X28))|v4_orders_2(X28)|~l1_orders_2(X28))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_orders_2])])])).
cnf(c_0_33, plain, (r3_orders_1(X1,X2)|~r1_relat_2(X1,X2)|~r8_relat_2(X1,X2)|~r4_relat_2(X1,X2)|~r6_relat_2(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (r6_relat_2(X1,X2)|~r7_relat_2(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (r1_relat_2(X1,X2)|~r7_relat_2(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (r7_relat_2(u1_orders_2(X2),X1)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, negated_conjecture, (v6_orders_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_38, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_39, negated_conjecture, (r4_relat_2(X1,esk2_0)|~r4_relat_2(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_40, plain, (r4_relat_2(u1_orders_2(X1),u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_42, negated_conjecture, (r8_relat_2(X1,esk2_0)|~r8_relat_2(X1,u1_struct_0(esk1_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_31, c_0_29])).
cnf(c_0_43, plain, (r8_relat_2(u1_orders_2(X1),u1_struct_0(X1))|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_45, plain, ![X29, X30]:(~v1_relat_1(X29)|(~m1_subset_1(X30,k1_zfmisc_1(X29))|v1_relat_1(X30))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_46, plain, ![X31]:(~l1_orders_2(X31)|m1_subset_1(u1_orders_2(X31),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X31),u1_struct_0(X31))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
fof(c_0_47, plain, ![X32, X33]:v1_relat_1(k2_zfmisc_1(X32,X33)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_48, plain, (r3_orders_1(X1,X2)|~r7_relat_2(X1,X2)|~r8_relat_2(X1,X2)|~r4_relat_2(X1,X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])).
cnf(c_0_49, negated_conjecture, (r7_relat_2(u1_orders_2(esk1_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_18]), c_0_37]), c_0_38])])).
cnf(c_0_50, negated_conjecture, (~r3_orders_1(u1_orders_2(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_51, negated_conjecture, (r4_relat_2(u1_orders_2(esk1_0),esk2_0)|~v1_relat_1(u1_orders_2(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_38]), c_0_41])])).
cnf(c_0_52, negated_conjecture, (r8_relat_2(u1_orders_2(esk1_0),esk2_0)|~v1_relat_1(u1_orders_2(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_38]), c_0_44])])).
cnf(c_0_53, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, negated_conjecture, (~v1_relat_1(u1_orders_2(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_57, plain, (v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])])).
cnf(c_0_58, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_38])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 32
# Proof object formula steps           : 27
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 80
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 79
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 29
# ...of the previous two non-trivial   : 19
# Contextual simplify-reflections      : 4
# Paramodulations                      : 29
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 48
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 3
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 420
# Rec. Clause-clause subsumption calls : 231
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 19
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2563

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28584 pages
