# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t70_enumset1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', idempotence_k2_xboole_0)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t72_enumset1)).
fof(t94_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t94_xboole_1)).
fof(d6_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1)), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', d6_xboole_0)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t100_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', d5_xboole_0)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t1_boole)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t12_xboole_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t83_xboole_1)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t70_xboole_1)).
fof(t51_xboole_1, axiom, ![X1, X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t51_xboole_1)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t1_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t3_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', d3_tarski)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t36_xboole_1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', l1_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t69_enumset1)).
fof(t140_zfmisc_1, conjecture, ![X1, X2]:(r2_hidden(X1,X2)=>k2_xboole_0(k4_xboole_0(X2,k1_tarski(X1)),k1_tarski(X1))=X2), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t140_zfmisc_1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t98_xboole_1)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/8670267a5a5d45a8968b0146d07b015c.E_input', t2_tarski)).
fof(c_0_23, plain, ![X5, X6]:k3_tarski(k2_tarski(X5,X6))=k2_xboole_0(X5,X6), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_24, plain, ![X7, X8]:k1_enumset1(X7,X7,X8)=k2_tarski(X7,X8), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_25, plain, ![X9]:k2_xboole_0(X9,X9)=X9, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_26, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_28, plain, ![X10, X11, X12]:k2_enumset1(X10,X10,X11,X12)=k1_enumset1(X10,X11,X12), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_29, plain, ![X13, X14, X15, X16]:k3_enumset1(X13,X13,X14,X15,X16)=k2_enumset1(X13,X14,X15,X16), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_30, plain, ![X17, X18]:k2_xboole_0(X17,X18)=k5_xboole_0(k5_xboole_0(X17,X18),k3_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t94_xboole_1])).
fof(c_0_31, plain, ![X30, X31]:k5_xboole_0(X30,X31)=k2_xboole_0(k4_xboole_0(X30,X31),k4_xboole_0(X31,X30)), inference(variable_rename,[status(thm)],[d6_xboole_0])).
fof(c_0_32, plain, ![X28, X29]:k4_xboole_0(X28,X29)=k5_xboole_0(X28,k3_xboole_0(X28,X29)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_33, plain, ![X19, X20, X21, X22, X23, X24, X25, X26]:((((r2_hidden(X22,X19)|~r2_hidden(X22,X21)|X21!=k4_xboole_0(X19,X20))&(~r2_hidden(X22,X20)|~r2_hidden(X22,X21)|X21!=k4_xboole_0(X19,X20)))&(~r2_hidden(X23,X19)|r2_hidden(X23,X20)|r2_hidden(X23,X21)|X21!=k4_xboole_0(X19,X20)))&((~r2_hidden(esk1_3(X24,X25,X26),X26)|(~r2_hidden(esk1_3(X24,X25,X26),X24)|r2_hidden(esk1_3(X24,X25,X26),X25))|X26=k4_xboole_0(X24,X25))&((r2_hidden(esk1_3(X24,X25,X26),X24)|r2_hidden(esk1_3(X24,X25,X26),X26)|X26=k4_xboole_0(X24,X25))&(~r2_hidden(esk1_3(X24,X25,X26),X25)|r2_hidden(esk1_3(X24,X25,X26),X26)|X26=k4_xboole_0(X24,X25))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_34, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_36, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_43, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X2))=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_44, plain, (k5_xboole_0(X1,X2)=k3_tarski(k3_enumset1(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X2,k3_xboole_0(X2,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_35]), c_0_40]), c_0_40]), c_0_40]), c_0_36]), c_0_37])).
cnf(c_0_45, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_41, c_0_40])).
cnf(c_0_47, plain, (k5_xboole_0(k5_xboole_0(X1,X1),k3_xboole_0(X1,X1))=X1), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, plain, (k5_xboole_0(k5_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X2,k3_xboole_0(X2,X1))),k3_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X2,k3_xboole_0(X2,X1))))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_44, c_0_43])).
fof(c_0_49, plain, ![X43]:k2_xboole_0(X43,k1_xboole_0)=X43, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_50, plain, ![X47, X48]:(~r1_tarski(X47,X48)|k2_xboole_0(X47,X48)=X48), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_51, plain, ![X35, X36]:((~r1_xboole_0(X35,X36)|k4_xboole_0(X35,X36)=X35)&(k4_xboole_0(X35,X36)!=X35|r1_xboole_0(X35,X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
cnf(c_0_52, plain, (r2_hidden(X1,X2)|X3!=k5_xboole_0(X2,k3_xboole_0(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_45, c_0_40])).
cnf(c_0_53, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_54, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X1))=k5_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
fof(c_0_55, plain, ![X44, X45, X46]:((r1_xboole_0(X44,k2_xboole_0(X45,X46))|(~r1_xboole_0(X44,X45)|~r1_xboole_0(X44,X46)))&((r1_xboole_0(X44,X45)|~r1_xboole_0(X44,k2_xboole_0(X45,X46)))&(r1_xboole_0(X44,X46)|~r1_xboole_0(X44,k2_xboole_0(X45,X46))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
fof(c_0_56, plain, ![X55, X56]:k2_xboole_0(k3_xboole_0(X55,X56),k4_xboole_0(X55,X56))=X55, inference(variable_rename,[status(thm)],[t51_xboole_1])).
cnf(c_0_57, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_58, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_59, plain, ![X32, X33, X34]:(((~r2_hidden(X32,X33)|~r2_hidden(X32,X34)|~r2_hidden(X32,k5_xboole_0(X33,X34)))&(r2_hidden(X32,X33)|r2_hidden(X32,X34)|~r2_hidden(X32,k5_xboole_0(X33,X34))))&((~r2_hidden(X32,X33)|r2_hidden(X32,X34)|r2_hidden(X32,k5_xboole_0(X33,X34)))&(~r2_hidden(X32,X34)|r2_hidden(X32,X33)|r2_hidden(X32,k5_xboole_0(X33,X34))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_60, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (~r2_hidden(X1,k5_xboole_0(X2,X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_63, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64, plain, (k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_65, plain, ![X37, X38, X40, X41, X42]:(((r2_hidden(esk2_2(X37,X38),X37)|r1_xboole_0(X37,X38))&(r2_hidden(esk2_2(X37,X38),X38)|r1_xboole_0(X37,X38)))&(~r2_hidden(X42,X40)|~r2_hidden(X42,X41)|~r1_xboole_0(X40,X41))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_66, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_67, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_35]), c_0_36]), c_0_37])).
fof(c_0_68, plain, ![X49, X50, X51, X52, X53]:((~r1_tarski(X49,X50)|(~r2_hidden(X51,X49)|r2_hidden(X51,X50)))&((r2_hidden(esk3_2(X52,X53),X52)|r1_tarski(X52,X53))&(~r2_hidden(esk3_2(X52,X53),X53)|r1_tarski(X52,X53)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_69, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_60, c_0_40])).
cnf(c_0_71, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_72, plain, (~r2_hidden(X1,k5_xboole_0(X2,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_54]), c_0_62])).
cnf(c_0_73, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k3_tarski(k3_enumset1(X2,X2,X2,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_74, plain, (k3_tarski(k3_enumset1(k3_xboole_0(X1,X2),k3_xboole_0(X1,X2),k3_xboole_0(X1,X2),k3_xboole_0(X1,X2),k5_xboole_0(X1,k3_xboole_0(X1,X2))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_35]), c_0_40]), c_0_36]), c_0_37])).
cnf(c_0_75, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_76, plain, (k5_xboole_0(k5_xboole_0(X1,k1_xboole_0),k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_66, c_0_43])).
cnf(c_0_77, plain, (k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))=X2|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_67])).
cnf(c_0_78, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_79, plain, ![X60, X61]:r1_tarski(k4_xboole_0(X60,X61),X60), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_80, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_81, plain, (r2_hidden(X1,k3_xboole_0(X2,X2))|~r2_hidden(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_54]), c_0_72])).
cnf(c_0_82, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k5_xboole_0(k5_xboole_0(X2,X3),k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_73, c_0_43])).
cnf(c_0_83, plain, (k5_xboole_0(k5_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,k3_xboole_0(X1,X2))),k3_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,k3_xboole_0(X1,X2))))=X1), inference(rw,[status(thm)],[c_0_74, c_0_43])).
cnf(c_0_84, plain, (r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(esk2_2(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_53, c_0_75])).
cnf(c_0_85, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_86, plain, ![X62, X63]:((~r1_tarski(k1_tarski(X62),X63)|r2_hidden(X62,X63))&(~r2_hidden(X62,X63)|r1_tarski(k1_tarski(X62),X63))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_87, plain, ![X64]:k2_tarski(X64,X64)=k1_tarski(X64), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_88, plain, (k1_xboole_0=X1|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_89, plain, (r1_tarski(k5_xboole_0(X1,X1),X2)), inference(spm,[status(thm)],[c_0_72, c_0_78])).
cnf(c_0_90, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_91, plain, (~r1_xboole_0(X1,X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_92, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_93, plain, (r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_94, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_95, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
fof(c_0_96, negated_conjecture, ~(![X1, X2]:(r2_hidden(X1,X2)=>k2_xboole_0(k4_xboole_0(X2,k1_tarski(X1)),k1_tarski(X1))=X2)), inference(assume_negation,[status(cth)],[t140_zfmisc_1])).
cnf(c_0_97, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_98, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_90, c_0_40])).
cnf(c_0_99, plain, (r1_xboole_0(X1,k5_xboole_0(X2,X2))|~r2_hidden(esk2_2(X1,k5_xboole_0(X2,X2)),X2)), inference(spm,[status(thm)],[c_0_62, c_0_75])).
cnf(c_0_100, plain, (~r1_xboole_0(k3_xboole_0(X1,X2),X1)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_101, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_93, c_0_70])).
cnf(c_0_102, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X2,X3)|~r1_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_73, c_0_67])).
cnf(c_0_103, plain, (r1_tarski(k3_enumset1(X1,X1,X1,X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95]), c_0_27]), c_0_36]), c_0_37])).
fof(c_0_104, negated_conjecture, (r2_hidden(esk5_0,esk6_0)&k2_xboole_0(k4_xboole_0(esk6_0,k1_tarski(esk5_0)),k1_tarski(esk5_0))!=esk6_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_96])])])).
cnf(c_0_105, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X1))=k1_xboole_0), inference(rw,[status(thm)],[c_0_54, c_0_97])).
cnf(c_0_106, plain, (r1_tarski(X1,X1)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_98, c_0_70])).
cnf(c_0_107, plain, (r1_xboole_0(X1,k5_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_99, c_0_85])).
cnf(c_0_108, plain, (~r1_xboole_0(k3_xboole_0(k3_xboole_0(X1,X2),X3),X1)|~r2_hidden(X4,k3_xboole_0(k3_xboole_0(X1,X2),X3))), inference(spm,[status(thm)],[c_0_100, c_0_92])).
cnf(c_0_109, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_101, c_0_92])).
fof(c_0_110, plain, ![X67, X68]:k2_xboole_0(X67,X68)=k5_xboole_0(X67,k4_xboole_0(X68,X67)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
cnf(c_0_111, plain, (r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_102, c_0_98])).
cnf(c_0_112, plain, (r1_xboole_0(X1,k3_enumset1(X2,X2,X2,X2,X2))|~r1_xboole_0(X1,X3)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_113, negated_conjecture, (r2_hidden(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_114, plain, (k3_xboole_0(X1,X1)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_105]), c_0_76])).
cnf(c_0_115, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_116, plain, (k5_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_98])).
cnf(c_0_117, plain, (r1_tarski(k1_xboole_0,X1)), inference(rw,[status(thm)],[c_0_89, c_0_97])).
cnf(c_0_118, plain, (~r1_xboole_0(X1,k3_xboole_0(X1,X2))|~r2_hidden(X3,k3_xboole_0(k3_xboole_0(X1,X2),X4))), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_119, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)|~r2_hidden(esk2_2(X1,k3_xboole_0(X2,X3)),X2)), inference(spm,[status(thm)],[c_0_80, c_0_75])).
fof(c_0_120, plain, ![X57, X58]:((~r2_hidden(esk4_2(X57,X58),X57)|~r2_hidden(esk4_2(X57,X58),X58)|X57=X58)&(r2_hidden(esk4_2(X57,X58),X57)|r2_hidden(esk4_2(X57,X58),X58)|X57=X58)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
cnf(c_0_121, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_122, plain, (~r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)|~r2_hidden(X3,k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_91, c_0_111])).
cnf(c_0_123, negated_conjecture, (r1_xboole_0(X1,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))|~r1_xboole_0(X1,esk6_0)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_124, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_114]), c_0_97]), c_0_115])])).
cnf(c_0_125, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_116]), c_0_116]), c_0_117])])).
cnf(c_0_126, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X3,X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_127, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_128, plain, (~r1_xboole_0(X1,k3_xboole_0(X1,X2))|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_118, c_0_114])).
cnf(c_0_129, plain, (r1_xboole_0(X1,k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_119, c_0_85])).
cnf(c_0_130, plain, (~r2_hidden(X1,k1_xboole_0)), inference(rw,[status(thm)],[c_0_72, c_0_97])).
cnf(c_0_131, plain, (r2_hidden(esk4_2(X1,X2),X1)|r2_hidden(esk4_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_132, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X2))=k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_35]), c_0_40]), c_0_36]), c_0_37])).
cnf(c_0_133, negated_conjecture, (~r1_xboole_0(k5_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),k3_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),X1)),esk6_0)|~r2_hidden(X2,k5_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),k3_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),X1)))), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_134, plain, (r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_101, c_0_93])).
cnf(c_0_135, plain, (~r2_hidden(X1,k3_xboole_0(X2,k1_xboole_0))|~r2_hidden(X1,k5_xboole_0(X2,k1_xboole_0))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_69, c_0_76])).
cnf(c_0_136, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_124]), c_0_125]), c_0_117])])).
cnf(c_0_137, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_61, c_0_126])).
cnf(c_0_138, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X4!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_127, c_0_40])).
cnf(c_0_139, negated_conjecture, (k2_xboole_0(k4_xboole_0(esk6_0,k1_tarski(esk5_0)),k1_tarski(esk5_0))!=esk6_0), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_140, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_128, c_0_129])).
cnf(c_0_141, plain, (k1_xboole_0=X1|r2_hidden(esk4_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_130, c_0_131])).
cnf(c_0_142, plain, (k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))=k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_43, c_0_132])).
cnf(c_0_143, negated_conjecture, (~r2_hidden(X1,k5_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),k3_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),esk6_0)))), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_144, plain, (~r2_hidden(X1,k3_xboole_0(X2,k1_xboole_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_135, c_0_136])]), c_0_137])).
cnf(c_0_145, plain, (r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_138])).
cnf(c_0_146, negated_conjecture, (k3_tarski(k3_enumset1(k5_xboole_0(esk6_0,k3_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))),k5_xboole_0(esk6_0,k3_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))),k5_xboole_0(esk6_0,k3_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))),k5_xboole_0(esk6_0,k3_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))),k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0)))!=esk6_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_139, c_0_95]), c_0_95]), c_0_27]), c_0_27]), c_0_35]), c_0_40]), c_0_40]), c_0_36]), c_0_36]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_37]), c_0_37]), c_0_37])).
cnf(c_0_147, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_148, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k3_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k3_xboole_0(X1,X2))))=X1), inference(rw,[status(thm)],[c_0_83, c_0_142])).
cnf(c_0_149, negated_conjecture, (k5_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),k3_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),esk6_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_143, c_0_141])).
cnf(c_0_150, plain, (r1_tarski(k3_xboole_0(X1,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_144, c_0_78])).
cnf(c_0_151, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_145]), c_0_137])).
cnf(c_0_152, negated_conjecture, (k5_xboole_0(k5_xboole_0(k5_xboole_0(esk6_0,k3_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))),k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0)),k3_xboole_0(k5_xboole_0(esk6_0,k3_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))),k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0)))!=esk6_0), inference(rw,[status(thm)],[c_0_146, c_0_43])).
cnf(c_0_153, plain, (k3_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_147, c_0_134])).
cnf(c_0_154, plain, (k5_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(k5_xboole_0(X2,k3_xboole_0(X2,X1)),k3_xboole_0(k5_xboole_0(X2,k3_xboole_0(X2,X1)),k5_xboole_0(X1,k3_xboole_0(X1,X2)))))=k5_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_48, c_0_142])).
cnf(c_0_155, negated_conjecture, (k3_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),esk6_0)=k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_125]), c_0_97]), c_0_136])).
cnf(c_0_156, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_150])).
cnf(c_0_157, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_158, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|r2_hidden(esk3_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_151, c_0_78])).
cnf(c_0_159, negated_conjecture, (k5_xboole_0(k5_xboole_0(esk6_0,k3_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))),k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))!=esk6_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_153]), c_0_136])).
cnf(c_0_160, negated_conjecture, (k5_xboole_0(esk6_0,k3_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0)))=k5_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_97]), c_0_97]), c_0_156]), c_0_136]), c_0_124])).
cnf(c_0_161, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_162, negated_conjecture, (k5_xboole_0(k5_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),esk6_0),k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))!=esk6_0), inference(rw,[status(thm)],[c_0_159, c_0_160])).
cnf(c_0_163, negated_conjecture, (k5_xboole_0(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),esk6_0)=k5_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_160]), c_0_155]), c_0_97]), c_0_155]), c_0_97]), c_0_125]), c_0_97]), c_0_136])).
cnf(c_0_164, negated_conjecture, (r1_tarski(k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0),esk6_0)), inference(spm,[status(thm)],[c_0_161, c_0_155])).
cnf(c_0_165, negated_conjecture, (k5_xboole_0(k5_xboole_0(esk6_0,k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0)),k3_enumset1(esk5_0,esk5_0,esk5_0,esk5_0,esk5_0))!=esk6_0), inference(rw,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_166, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_163]), c_0_155]), c_0_164])]), c_0_165]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 167
# Proof object clause steps            : 120
# Proof object formula steps           : 47
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 57
# Proof object simplifying inferences  : 107
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 35
# Processed clauses                    : 1994
# ...of these trivial                  : 17
# ...subsumed                          : 1510
# ...remaining for further processing  : 466
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 48
# Generated clauses                    : 13913
# ...of the previous two non-trivial   : 12189
# Contextual simplify-reflections      : 9
# Paramodulations                      : 13834
# Factorizations                       : 64
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 364
#    Positive orientable unit clauses  : 30
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 324
# Current number of unprocessed clauses: 10072
# ...number of literals in the above   : 38221
# Current number of archived formulas  : 0
# Current number of archived clauses   : 105
# Clause-clause subsumption calls (NU) : 18658
# Rec. Clause-clause subsumption calls : 13780
# Non-unit clause-clause subsumptions  : 765
# Unit Clause-clause subsumption calls : 458
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 134
# BW rewrite match successes           : 48
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 241657

# -------------------------------------------------
# User time                : 0.153 s
# System time              : 0.007 s
# Total time               : 0.160 s
# Maximum resident set size: 28648 pages
