# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_waybel_9, conjecture, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X4),X1,X1))&(v11_waybel_0(X3,X1)&r3_waybel_9(X1,X3,X2)))=>X2=k1_waybel_9(X1,X3))))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', t36_waybel_9)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', cc1_lattice3)).
fof(dt_l1_waybel_9, axiom, ![X1]:(l1_waybel_9(X1)=>(l1_pre_topc(X1)&l1_orders_2(X1))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', dt_l1_waybel_9)).
fof(l51_waybel_9, axiom, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&(v11_waybel_0(X2,X1)&r3_waybel_9(X1,X2,X3))))=>r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)))))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', l51_waybel_9)).
fof(d2_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(l1_waybel_0(X2,X1)=>k1_waybel_9(X1,X2)=k5_yellow_2(X1,u1_waybel_0(X1,X2)))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', d2_waybel_9)).
fof(dt_u1_waybel_0, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_waybel_0(X2,X1))=>(v1_funct_1(u1_waybel_0(X1,X2))&(v1_funct_2(u1_waybel_0(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', dt_u1_waybel_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', dt_l1_orders_2)).
fof(d6_yellow_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(v1_relat_1(X2)=>k5_yellow_2(X1,X2)=k2_yellow_0(X1,k2_relat_1(X2)))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', d6_yellow_2)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', cc1_relset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', redefinition_k2_relset_1)).
fof(l52_waybel_9, axiom, ![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((X3=X4&(![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X5),X1,X1))&r3_waybel_9(X1,X2,X3)))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)=>r1_orders_2(X1,X5,X4)))))))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', l52_waybel_9)).
fof(t31_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/86b0ce91cacd42b3891bba57f0576586.E_input', t31_yellow_0)).
fof(c_0_12, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&(v8_pre_topc(X1)&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&(v2_lattice3(X1)&(v1_compts_1(X1)&l1_waybel_9(X1)))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1))))=>((![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>v5_pre_topc(k4_waybel_1(X1,X4),X1,X1))&(v11_waybel_0(X3,X1)&r3_waybel_9(X1,X3,X2)))=>X2=k1_waybel_9(X1,X3)))))), inference(assume_negation,[status(cth)],[t36_waybel_9])).
fof(c_0_13, plain, ![X30]:(~l1_orders_2(X30)|(~v1_lattice3(X30)|~v2_struct_0(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
fof(c_0_14, negated_conjecture, ![X9]:((v2_pre_topc(esk1_0)&(v8_pre_topc(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&(v2_lattice3(esk1_0)&(v1_compts_1(esk1_0)&l1_waybel_9(esk1_0)))))))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((~v2_struct_0(esk3_0)&(v4_orders_2(esk3_0)&(v7_waybel_0(esk3_0)&l1_waybel_0(esk3_0,esk1_0))))&(((~m1_subset_1(X9,u1_struct_0(esk1_0))|v5_pre_topc(k4_waybel_1(esk1_0,X9),esk1_0,esk1_0))&(v11_waybel_0(esk3_0,esk1_0)&r3_waybel_9(esk1_0,esk3_0,esk2_0)))&esk2_0!=k1_waybel_9(esk1_0,esk3_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])).
fof(c_0_15, plain, ![X10]:((l1_pre_topc(X10)|~l1_waybel_9(X10))&(l1_orders_2(X10)|~l1_waybel_9(X10))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_9])])])).
fof(c_0_16, plain, ![X11, X12, X13, X14]:((m1_subset_1(esk4_4(X11,X12,X13,X14),u1_struct_0(X11))|(~v11_waybel_0(X12,X11)|~r3_waybel_9(X11,X12,X13))|X13!=X14|r1_lattice3(X11,k2_relset_1(u1_struct_0(X12),u1_struct_0(X11),u1_waybel_0(X11,X12)),X14)|~m1_subset_1(X14,u1_struct_0(X11))|~m1_subset_1(X13,u1_struct_0(X11))|(v2_struct_0(X12)|(~v4_orders_2(X12)|(~v7_waybel_0(X12)|~l1_waybel_0(X12,X11))))|(~v2_pre_topc(X11)|(~v8_pre_topc(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|(~v1_lattice3(X11)|(~v2_lattice3(X11)|(~v1_compts_1(X11)|~l1_waybel_9(X11))))))))))&(~v5_pre_topc(k4_waybel_1(X11,esk4_4(X11,X12,X13,X14)),X11,X11)|(~v11_waybel_0(X12,X11)|~r3_waybel_9(X11,X12,X13))|X13!=X14|r1_lattice3(X11,k2_relset_1(u1_struct_0(X12),u1_struct_0(X11),u1_waybel_0(X11,X12)),X14)|~m1_subset_1(X14,u1_struct_0(X11))|~m1_subset_1(X13,u1_struct_0(X11))|(v2_struct_0(X12)|(~v4_orders_2(X12)|(~v7_waybel_0(X12)|~l1_waybel_0(X12,X11))))|(~v2_pre_topc(X11)|(~v8_pre_topc(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|(~v1_lattice3(X11)|(~v2_lattice3(X11)|(~v1_compts_1(X11)|~l1_waybel_9(X11))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l51_waybel_9])])])])])])).
cnf(c_0_17, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (l1_orders_2(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (l1_waybel_9(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (m1_subset_1(esk4_4(X1,X2,X3,X4),u1_struct_0(X1))|r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|v2_struct_0(X2)|~v11_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X34, X35]:(v2_struct_0(X34)|~l1_orders_2(X34)|(~l1_waybel_0(X35,X34)|k1_waybel_9(X34,X35)=k5_yellow_2(X34,u1_waybel_0(X34,X35)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_waybel_9])])])])).
cnf(c_0_23, negated_conjecture, (~l1_orders_2(esk1_0)|~v2_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, negated_conjecture, (l1_orders_2(esk1_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_25, plain, ![X28, X29]:((v1_funct_1(u1_waybel_0(X28,X29))|(~l1_struct_0(X28)|~l1_waybel_0(X29,X28)))&((v1_funct_2(u1_waybel_0(X28,X29),u1_struct_0(X29),u1_struct_0(X28))|(~l1_struct_0(X28)|~l1_waybel_0(X29,X28)))&(m1_subset_1(u1_waybel_0(X28,X29),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X28))))|(~l1_struct_0(X28)|~l1_waybel_0(X29,X28))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_waybel_0])])])).
cnf(c_0_26, plain, (r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|m1_subset_1(esk4_4(X1,X2,X3,X3),u1_struct_0(X1))|~r3_waybel_9(X1,X2,X3)|~v11_waybel_0(X2,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (v11_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_28, negated_conjecture, (l1_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_29, negated_conjecture, (v7_waybel_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_30, negated_conjecture, (v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_31, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_32, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_33, negated_conjecture, (v4_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_34, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_35, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36, negated_conjecture, (v8_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_39, plain, (v2_struct_0(X1)|k1_waybel_9(X1,X2)=k5_yellow_2(X1,u1_waybel_0(X1,X2))|~l1_orders_2(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_40, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24])])).
cnf(c_0_41, plain, (m1_subset_1(u1_waybel_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_42, plain, ![X25]:(~l1_orders_2(X25)|l1_struct_0(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_43, negated_conjecture, (r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),X1)|m1_subset_1(esk4_4(esk1_0,esk3_0,X1,X1),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), c_0_20]), c_0_30]), c_0_31]), c_0_18]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_44, negated_conjecture, (r3_waybel_9(esk1_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_46, negated_conjecture, (esk2_0!=k1_waybel_9(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_47, negated_conjecture, (k1_waybel_9(esk1_0,esk3_0)=k5_yellow_2(esk1_0,u1_waybel_0(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_28]), c_0_24])]), c_0_40])).
fof(c_0_48, plain, ![X26, X27]:(v2_struct_0(X26)|~l1_orders_2(X26)|(~v1_relat_1(X27)|k5_yellow_2(X26,X27)=k2_yellow_0(X26,k2_relat_1(X27)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_yellow_2])])])])).
fof(c_0_49, plain, ![X31, X32, X33]:(~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))|v1_relat_1(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(u1_waybel_0(esk1_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0))))|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_41, c_0_28])).
cnf(c_0_51, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_52, plain, ![X22, X23, X24]:(~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X22,X23)))|k2_relset_1(X22,X23,X24)=k2_relat_1(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_53, plain, ![X16, X17, X18, X19, X21]:((m1_subset_1(esk5_4(X16,X17,X18,X19),u1_struct_0(X16))|~r3_waybel_9(X16,X17,X18)|X18!=X19|(~m1_subset_1(X21,u1_struct_0(X16))|(~r1_lattice3(X16,k2_relset_1(u1_struct_0(X17),u1_struct_0(X16),u1_waybel_0(X16,X17)),X21)|r1_orders_2(X16,X21,X19)))|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16))))|(~v2_pre_topc(X16)|(~v8_pre_topc(X16)|(~v3_orders_2(X16)|(~v4_orders_2(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|(~v2_lattice3(X16)|(~v1_compts_1(X16)|~l1_waybel_9(X16))))))))))&(~v5_pre_topc(k4_waybel_1(X16,esk5_4(X16,X17,X18,X19)),X16,X16)|~r3_waybel_9(X16,X17,X18)|X18!=X19|(~m1_subset_1(X21,u1_struct_0(X16))|(~r1_lattice3(X16,k2_relset_1(u1_struct_0(X17),u1_struct_0(X16),u1_waybel_0(X16,X17)),X21)|r1_orders_2(X16,X21,X19)))|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X16))|(v2_struct_0(X17)|(~v4_orders_2(X17)|(~v7_waybel_0(X17)|~l1_waybel_0(X17,X16))))|(~v2_pre_topc(X16)|(~v8_pre_topc(X16)|(~v3_orders_2(X16)|(~v4_orders_2(X16)|(~v5_orders_2(X16)|(~v1_lattice3(X16)|(~v2_lattice3(X16)|(~v1_compts_1(X16)|~l1_waybel_9(X16))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l52_waybel_9])])])])])])).
cnf(c_0_54, plain, (r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X4)|v2_struct_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk4_4(X1,X2,X3,X4)),X1,X1)|~v11_waybel_0(X2,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_55, negated_conjecture, (v5_pre_topc(k4_waybel_1(esk1_0,X1),esk1_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_56, negated_conjecture, (r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),esk2_0)|m1_subset_1(esk4_4(esk1_0,esk3_0,esk2_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_57, negated_conjecture, (k5_yellow_2(esk1_0,u1_waybel_0(esk1_0,esk3_0))!=esk2_0), inference(rw,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_58, plain, (v2_struct_0(X1)|k5_yellow_2(X1,X2)=k2_yellow_0(X1,k2_relat_1(X2))|~l1_orders_2(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(u1_waybel_0(esk1_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_24])])).
cnf(c_0_61, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),u1_struct_0(X1))|r1_orders_2(X1,X5,X4)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X5,u1_struct_0(X1))|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_63, plain, ![X36, X37, X38, X39, X40]:(((r1_lattice3(X36,X38,X37)|(X37!=k2_yellow_0(X36,X38)|~r2_yellow_0(X36,X38))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36)))&(~m1_subset_1(X39,u1_struct_0(X36))|(~r1_lattice3(X36,X38,X39)|r1_orders_2(X36,X39,X37))|(X37!=k2_yellow_0(X36,X38)|~r2_yellow_0(X36,X38))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36))))&(((X37=k2_yellow_0(X36,X40)|(m1_subset_1(esk6_3(X36,X37,X40),u1_struct_0(X36))|~r1_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36)))&(r2_yellow_0(X36,X40)|(m1_subset_1(esk6_3(X36,X37,X40),u1_struct_0(X36))|~r1_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36))))&(((X37=k2_yellow_0(X36,X40)|(r1_lattice3(X36,X40,esk6_3(X36,X37,X40))|~r1_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36)))&(r2_yellow_0(X36,X40)|(r1_lattice3(X36,X40,esk6_3(X36,X37,X40))|~r1_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36))))&((X37=k2_yellow_0(X36,X40)|(~r1_orders_2(X36,esk6_3(X36,X37,X40),X37)|~r1_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36)))&(r2_yellow_0(X36,X40)|(~r1_orders_2(X36,esk6_3(X36,X37,X40),X37)|~r1_lattice3(X36,X40,X37))|~m1_subset_1(X37,u1_struct_0(X36))|(~v5_orders_2(X36)|~l1_orders_2(X36))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])).
cnf(c_0_64, plain, (r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X3)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~v11_waybel_0(X2,X1)|~v5_pre_topc(k4_waybel_1(X1,esk4_4(X1,X2,X3,X3)),X1,X1)|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X2)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_54])).
cnf(c_0_65, negated_conjecture, (r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),esk2_0)|v5_pre_topc(k4_waybel_1(esk1_0,esk4_4(esk1_0,esk3_0,esk2_0,esk2_0)),esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_66, negated_conjecture, (k2_yellow_0(esk1_0,k2_relat_1(u1_waybel_0(esk1_0,esk3_0)))!=esk2_0|~v1_relat_1(u1_waybel_0(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_24])]), c_0_40])).
cnf(c_0_67, negated_conjecture, (v1_relat_1(u1_waybel_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_68, negated_conjecture, (k2_relat_1(u1_waybel_0(esk1_0,esk3_0))=k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_61, c_0_60])).
cnf(c_0_69, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X4)|m1_subset_1(esk5_4(X1,X4,X3,X3),u1_struct_0(X1))|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X4),u1_struct_0(X1),u1_waybel_0(X1,X4)),X2)|~r3_waybel_9(X1,X4,X3)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_70, plain, (X1=k2_yellow_0(X2,X3)|r1_lattice3(X2,X3,esk6_3(X2,X1,X3))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_71, negated_conjecture, (r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_44]), c_0_27]), c_0_28]), c_0_29]), c_0_45]), c_0_20]), c_0_30]), c_0_31]), c_0_18]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_72, negated_conjecture, (esk2_0!=k2_yellow_0(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67])]), c_0_68])).
cnf(c_0_73, plain, (X1=k2_yellow_0(X2,X3)|m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_74, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|m1_subset_1(esk5_4(esk1_0,esk3_0,esk2_0,esk2_0),u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_44]), c_0_28]), c_0_29]), c_0_45]), c_0_20]), c_0_30]), c_0_31]), c_0_18]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_75, negated_conjecture, (r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_24]), c_0_45]), c_0_32])]), c_0_72])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_71]), c_0_24]), c_0_45]), c_0_32])]), c_0_72])).
cnf(c_0_77, plain, (r1_orders_2(X1,X5,X4)|v2_struct_0(X2)|~v5_pre_topc(k4_waybel_1(X1,esk5_4(X1,X2,X3,X4)),X1,X1)|~r3_waybel_9(X1,X2,X3)|X3!=X4|~m1_subset_1(X5,u1_struct_0(X1))|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X2),u1_struct_0(X1),u1_waybel_0(X1,X2)),X5)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~v8_pre_topc(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v1_compts_1(X1)|~l1_waybel_9(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_78, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))),esk2_0)|m1_subset_1(esk5_4(esk1_0,esk3_0,esk2_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])])).
cnf(c_0_79, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X4)|~r1_lattice3(X1,k2_relset_1(u1_struct_0(X4),u1_struct_0(X1),u1_waybel_0(X1,X4)),X2)|~r3_waybel_9(X1,X4,X3)|~v5_pre_topc(k4_waybel_1(X1,esk5_4(X1,X4,X3,X3)),X1,X1)|~l1_waybel_0(X4,X1)|~v7_waybel_0(X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_9(X1)|~v1_compts_1(X1)|~v2_lattice3(X1)|~v1_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X4)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v8_pre_topc(X1)|~v2_pre_topc(X1)), inference(er,[status(thm)],[c_0_77])).
cnf(c_0_80, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))),esk2_0)|v5_pre_topc(k4_waybel_1(esk1_0,esk5_4(esk1_0,esk3_0,esk2_0,esk2_0)),esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_55, c_0_78])).
cnf(c_0_81, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))),esk2_0)|r1_orders_2(esk1_0,X1,esk2_0)|~r1_lattice3(esk1_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_44]), c_0_28]), c_0_29]), c_0_45]), c_0_20]), c_0_30]), c_0_31]), c_0_18]), c_0_32]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_82, plain, (X1=k2_yellow_0(X2,X3)|~r1_orders_2(X2,esk6_3(X2,X1,X3),X1)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_83, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk2_0,k2_relset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0),u1_waybel_0(esk1_0,esk3_0))),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_75]), c_0_76])])).
cnf(c_0_84, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_71]), c_0_24]), c_0_45]), c_0_32])]), c_0_72]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 85
# Proof object clause steps            : 60
# Proof object formula steps           : 25
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 102
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 41
# Processed clauses                    : 125
# ...of these trivial                  : 2
# ...subsumed                          : 0
# ...remaining for further processing  : 123
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 14
# Generated clauses                    : 47
# ...of the previous two non-trivial   : 48
# Contextual simplify-reflections      : 0
# Paramodulations                      : 41
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 62
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 29
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 6
# Current number of archived formulas  : 0
# Current number of archived clauses   : 55
# Clause-clause subsumption calls (NU) : 2646
# Rec. Clause-clause subsumption calls : 72
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 26
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5958

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28740 pages
