# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', l32_xboole_1)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', d2_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t36_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t3_boole)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t17_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t28_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t41_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', commutativity_k3_xboole_0)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t39_xboole_1)).
fof(t40_xboole_1, axiom, ![X1, X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t40_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t4_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t3_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', d7_xboole_0)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t16_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', idempotence_k3_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', d5_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', commutativity_k2_xboole_0)).
fof(t47_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/86c381bfe7e84eaf8950a5f0996cb0aa.E_input', t47_xboole_1)).
fof(c_0_18, plain, ![X5, X6]:((k4_xboole_0(X5,X6)!=k1_xboole_0|r1_tarski(X5,X6))&(~r1_tarski(X5,X6)|k4_xboole_0(X5,X6)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_19, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_20, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
fof(c_0_21, plain, ![X10, X11]:r1_tarski(k4_xboole_0(X10,X11),X10), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_22, plain, ![X19]:k4_xboole_0(X19,k1_xboole_0)=X19, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_23, plain, ![X7, X8]:r1_tarski(k3_xboole_0(X7,X8),X7), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_24, plain, ![X14, X15]:(~r1_tarski(X14,X15)|k3_xboole_0(X14,X15)=X14), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (k4_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X16, X17, X18]:k4_xboole_0(k4_xboole_0(X16,X17),X18)=k4_xboole_0(X16,k2_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_29, plain, ![X12, X13]:k3_xboole_0(X12,X13)=k3_xboole_0(X13,X12), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_30, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_25, c_0_20])).
cnf(c_0_34, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_36, plain, ![X23, X24]:k2_xboole_0(X23,k4_xboole_0(X24,X23))=k2_xboole_0(X23,X24), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_37, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_38, plain, ![X25, X26]:k4_xboole_0(k2_xboole_0(X25,X26),X26)=k4_xboole_0(X25,X26), inference(variable_rename,[status(thm)],[t40_xboole_1])).
cnf(c_0_39, plain, (k4_xboole_0(X1,o_0_0_xboole_0)=X1), inference(rw,[status(thm)],[c_0_30, c_0_20])).
fof(c_0_40, plain, ![X31, X32, X34, X35, X36]:((r1_xboole_0(X31,X32)|r2_hidden(esk1_2(X31,X32),k3_xboole_0(X31,X32)))&(~r2_hidden(X36,k3_xboole_0(X34,X35))|~r1_xboole_0(X34,X35))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_41, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_31])).
cnf(c_0_42, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_43, plain, (k4_xboole_0(X1,k2_xboole_0(X2,X1))=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_31, c_0_37])).
cnf(c_0_46, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_27, c_0_39])).
cnf(c_0_48, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, plain, (k3_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_41])).
fof(c_0_50, plain, ![X37, X38, X40, X41, X42]:(((r2_hidden(esk2_2(X37,X38),X37)|r1_xboole_0(X37,X38))&(r2_hidden(esk2_2(X37,X38),X38)|r1_xboole_0(X37,X38)))&(~r2_hidden(X42,X40)|~r2_hidden(X42,X41)|~r1_xboole_0(X40,X41))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_51, plain, ![X29, X30]:((~r1_xboole_0(X29,X30)|k3_xboole_0(X29,X30)=k1_xboole_0)&(k3_xboole_0(X29,X30)!=k1_xboole_0|r1_xboole_0(X29,X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_52, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_53, plain, (k2_xboole_0(X1,k4_xboole_0(X2,k2_xboole_0(X3,X1)))=k2_xboole_0(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_44, c_0_35])).
cnf(c_0_54, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_45])).
cnf(c_0_55, plain, (k2_xboole_0(X1,o_0_0_xboole_0)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_46]), c_0_39])).
cnf(c_0_56, plain, (k4_xboole_0(X1,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_47])).
cnf(c_0_57, plain, (~r2_hidden(X1,o_0_0_xboole_0)|~r1_xboole_0(o_0_0_xboole_0,X2)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_58, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_60, plain, ![X20, X21, X22]:k3_xboole_0(k3_xboole_0(X20,X21),X22)=k3_xboole_0(X20,k3_xboole_0(X21,X22)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
fof(c_0_61, plain, ![X9]:k3_xboole_0(X9,X9)=X9, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_62, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_63, plain, ![X43, X44, X45, X46, X47, X48, X49, X50]:((((r2_hidden(X46,X43)|~r2_hidden(X46,X45)|X45!=k4_xboole_0(X43,X44))&(~r2_hidden(X46,X44)|~r2_hidden(X46,X45)|X45!=k4_xboole_0(X43,X44)))&(~r2_hidden(X47,X43)|r2_hidden(X47,X44)|r2_hidden(X47,X45)|X45!=k4_xboole_0(X43,X44)))&((~r2_hidden(esk3_3(X48,X49,X50),X50)|(~r2_hidden(esk3_3(X48,X49,X50),X48)|r2_hidden(esk3_3(X48,X49,X50),X49))|X50=k4_xboole_0(X48,X49))&((r2_hidden(esk3_3(X48,X49,X50),X48)|r2_hidden(esk3_3(X48,X49,X50),X50)|X50=k4_xboole_0(X48,X49))&(~r2_hidden(esk3_3(X48,X49,X50),X49)|r2_hidden(esk3_3(X48,X49,X50),X50)|X50=k4_xboole_0(X48,X49))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_64, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_45, c_0_52])).
cnf(c_0_65, plain, (k2_xboole_0(X1,k4_xboole_0(k3_xboole_0(X2,k2_xboole_0(X3,X1)),X3))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
fof(c_0_66, plain, ![X27, X28]:k2_xboole_0(X27,X28)=k2_xboole_0(X28,X27), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_67, plain, (r1_tarski(o_0_0_xboole_0,X1)), inference(spm,[status(thm)],[c_0_27, c_0_56])).
cnf(c_0_68, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_48, c_0_37])).
cnf(c_0_69, plain, (r1_xboole_0(X1,o_0_0_xboole_0)|~r1_xboole_0(o_0_0_xboole_0,X2)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_70, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_59, c_0_20])).
cnf(c_0_71, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_72, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_73, plain, (k3_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_62, c_0_20])).
cnf(c_0_74, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_48, c_0_58])).
cnf(c_0_75, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_76, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_77, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),X3)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_78, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_79, plain, (k4_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k4_xboole_0(X1,X2))))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_43])).
cnf(c_0_80, plain, (k4_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_67])).
cnf(c_0_81, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(k2_xboole_0(X3,X2),X2)), inference(spm,[status(thm)],[c_0_68, c_0_52])).
cnf(c_0_82, plain, (k2_xboole_0(X1,k2_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_46]), c_0_44])).
cnf(c_0_83, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_84, plain, (~r2_hidden(X1,o_0_0_xboole_0)|~r1_xboole_0(X2,o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_68, c_0_49])).
cnf(c_0_85, plain, (r1_xboole_0(X1,o_0_0_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_49])])).
cnf(c_0_86, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_87, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=o_0_0_xboole_0|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_88, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_75])).
cnf(c_0_89, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_90, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_91, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X3),X2)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_92, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k4_xboole_0(X1,X2))))=X1), inference(spm,[status(thm)],[c_0_42, c_0_79])).
cnf(c_0_93, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_56]), c_0_80])).
cnf(c_0_94, plain, (~r2_hidden(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(k2_xboole_0(X3,X2),k2_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_95, plain, (k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_82])).
cnf(c_0_96, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_83])).
cnf(c_0_97, plain, (~r2_hidden(X1,o_0_0_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_85])])).
cnf(c_0_98, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_54]), c_0_55])).
cnf(c_0_99, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_86, c_0_37])).
cnf(c_0_100, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=o_0_0_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_87]), c_0_71])).
cnf(c_0_101, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk2_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_88, c_0_58])).
cnf(c_0_102, plain, (r2_hidden(esk2_2(k4_xboole_0(X1,X2),X3),X1)|r1_xboole_0(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_103, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X3))),X3)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_104, plain, (k2_xboole_0(o_0_0_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_78, c_0_55])).
cnf(c_0_105, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X1)=k4_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_46, c_0_78])).
cnf(c_0_106, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_42, c_0_93])).
cnf(c_0_107, plain, (k2_xboole_0(X1,X2)!=o_0_0_xboole_0|~r2_hidden(X3,k2_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_70]), c_0_95])).
cnf(c_0_108, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_43]), c_0_97])).
cnf(c_0_109, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_98, c_0_37])).
cnf(c_0_110, plain, (k3_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_111, plain, (r1_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_112, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_113, plain, (k2_xboole_0(X1,X2)=X1|k4_xboole_0(X2,X1)!=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_105]), c_0_37]), c_0_106])).
cnf(c_0_114, plain, (k2_xboole_0(X1,X2)!=o_0_0_xboole_0|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_115, plain, (k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,k3_xboole_0(X2,X3)))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_109, c_0_71])).
cnf(c_0_116, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_117, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_118, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_112]), c_0_37])).
cnf(c_0_119, plain, (k2_xboole_0(X1,X2)=X2|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_78, c_0_113])).
cnf(c_0_120, plain, (k3_xboole_0(X1,X2)!=o_0_0_xboole_0|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_121, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk3_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_116])).
cnf(c_0_122, plain, (~r2_hidden(X1,k3_xboole_0(X2,k3_xboole_0(X3,X4)))|~r1_xboole_0(k3_xboole_0(X2,X3),X4)), inference(spm,[status(thm)],[c_0_48, c_0_71])).
cnf(c_0_123, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_124, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_125, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_27]), c_0_37])).
cnf(c_0_126, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k3_xboole_0(X1,X3)))=k3_xboole_0(X2,k3_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_99]), c_0_71]), c_0_71])).
cnf(c_0_127, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,k3_xboole_0(X1,X2))!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_109, c_0_119])).
cnf(c_0_128, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_129, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_130, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_101, c_0_90])).
cnf(c_0_131, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=k4_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_46, c_0_44])).
cnf(c_0_132, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_124]), c_0_125])).
cnf(c_0_133, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X1,X3)))=k3_xboole_0(X2,k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_126, c_0_125])).
cnf(c_0_134, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_71])).
cnf(c_0_135, plain, (r1_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,k3_xboole_0(X1,X2))))), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_136, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_137, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k3_xboole_0(X2,X1)|k3_xboole_0(X2,X1)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_138, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_135, c_0_125])).
cnf(c_0_139, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X2), inference(spm,[status(thm)],[c_0_136, c_0_78])).
cnf(c_0_140, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_137, c_0_118])).
cnf(c_0_141, plain, (k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_109, c_0_125])).
cnf(c_0_142, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_138, c_0_37])).
fof(c_0_143, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t47_xboole_1])).
cnf(c_0_144, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_141])).
cnf(c_0_145, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_142])).
fof(c_0_146, negated_conjecture, k4_xboole_0(esk4_0,k3_xboole_0(esk4_0,esk5_0))!=k4_xboole_0(esk4_0,esk5_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_143])])])).
cnf(c_0_147, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_39])).
cnf(c_0_148, negated_conjecture, (k4_xboole_0(esk4_0,k3_xboole_0(esk4_0,esk5_0))!=k4_xboole_0(esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_146])).
cnf(c_0_149, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_147]), c_0_105]), c_0_35]), c_0_78]), c_0_98])).
cnf(c_0_150, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_149])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 151
# Proof object clause steps            : 115
# Proof object formula steps           : 36
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 78
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 14164
# ...of these trivial                  : 2815
# ...subsumed                          : 9828
# ...remaining for further processing  : 1521
# Other redundant clauses eliminated   : 2087
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 109
# Generated clauses                    : 449667
# ...of the previous two non-trivial   : 209670
# Contextual simplify-reflections      : 0
# Paramodulations                      : 447574
# Factorizations                       : 6
# Equation resolutions                 : 2087
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1398
#    Positive orientable unit clauses  : 932
#    Positive unorientable unit clauses: 6
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 458
# Current number of unprocessed clauses: 195092
# ...number of literals in the above   : 341052
# Current number of archived formulas  : 0
# Current number of archived clauses   : 113
# Clause-clause subsumption calls (NU) : 39913
# Rec. Clause-clause subsumption calls : 35226
# Non-unit clause-clause subsumptions  : 2871
# Unit Clause-clause subsumption calls : 282
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6690
# BW rewrite match successes           : 247
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3565672

# -------------------------------------------------
# User time                : 2.218 s
# System time              : 0.066 s
# Total time               : 2.285 s
# Maximum resident set size: 150032 pages
