# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', d9_lattice3)).
fof(d5_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v2_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r2_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', d5_yellow_0)).
fof(t31_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', t31_yellow_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', t2_subset)).
fof(t15_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', t15_yellow_0)).
fof(t6_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', t6_yellow_0)).
fof(t43_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_yellow_0(X1)&l1_orders_2(X1))))=>(r2_yellow_0(X1,k1_xboole_0)&r1_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', t43_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/872c67f94ec24560a47a9500b6f6f8d2.E_input', dt_l1_orders_2)).
fof(c_0_9, plain, ![X5, X6, X7, X8]:((~r2_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X8,X7)))|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|r2_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&((r2_hidden(esk1_3(X5,X6,X7),X6)|r2_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&(~r1_orders_2(X5,esk1_3(X5,X6,X7),X7)|r2_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_10, plain, ![X10, X12]:(((m1_subset_1(esk2_1(X10),u1_struct_0(X10))|~v2_yellow_0(X10)|~l1_orders_2(X10))&(r2_lattice3(X10,u1_struct_0(X10),esk2_1(X10))|~v2_yellow_0(X10)|~l1_orders_2(X10)))&(~m1_subset_1(X12,u1_struct_0(X10))|~r2_lattice3(X10,u1_struct_0(X10),X12)|v2_yellow_0(X10)|~l1_orders_2(X10))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_0])])])])])).
fof(c_0_11, plain, ![X13, X14, X15, X16, X17]:(((r1_lattice3(X13,X15,X14)|(X14!=k2_yellow_0(X13,X15)|~r2_yellow_0(X13,X15))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(~m1_subset_1(X16,u1_struct_0(X13))|(~r1_lattice3(X13,X15,X16)|r1_orders_2(X13,X16,X14))|(X14!=k2_yellow_0(X13,X15)|~r2_yellow_0(X13,X15))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&(((X14=k2_yellow_0(X13,X17)|(m1_subset_1(esk3_3(X13,X14,X17),u1_struct_0(X13))|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r2_yellow_0(X13,X17)|(m1_subset_1(esk3_3(X13,X14,X17),u1_struct_0(X13))|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&(((X14=k2_yellow_0(X13,X17)|(r1_lattice3(X13,X17,esk3_3(X13,X14,X17))|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r2_yellow_0(X13,X17)|(r1_lattice3(X13,X17,esk3_3(X13,X14,X17))|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))&((X14=k2_yellow_0(X13,X17)|(~r1_orders_2(X13,esk3_3(X13,X14,X17),X14)|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13)))&(r2_yellow_0(X13,X17)|(~r1_orders_2(X13,esk3_3(X13,X14,X17),X14)|~r1_lattice3(X13,X17,X14))|~m1_subset_1(X14,u1_struct_0(X13))|(~v5_orders_2(X13)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])).
cnf(c_0_12, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (r2_lattice3(X1,u1_struct_0(X1),esk2_1(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (m1_subset_1(esk2_1(X1),u1_struct_0(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X19, X20]:(~m1_subset_1(X19,X20)|(v1_xboole_0(X20)|r2_hidden(X19,X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_16, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk3_3(X1,X3,X2),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (r1_orders_2(X1,X2,esk2_1(X1))|~v2_yellow_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
cnf(c_0_18, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (r2_yellow_0(X1,X2)|m1_subset_1(esk3_3(X1,X3,X2),u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, plain, (r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk2_1(X1))|~v5_orders_2(X1)|~v2_yellow_0(X1)|~r2_hidden(esk3_3(X1,esk2_1(X1),X2),u1_struct_0(X1))|~m1_subset_1(esk3_3(X1,esk2_1(X1),X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_14])).
cnf(c_0_21, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,X2)|r2_hidden(esk3_3(X1,X3,X2),u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_22, plain, ![X21, X22, X24, X25, X26]:(((m1_subset_1(esk4_2(X21,X22),u1_struct_0(X21))|~r1_yellow_0(X21,X22)|(~v5_orders_2(X21)|~l1_orders_2(X21)))&((r2_lattice3(X21,X22,esk4_2(X21,X22))|~r1_yellow_0(X21,X22)|(~v5_orders_2(X21)|~l1_orders_2(X21)))&(~m1_subset_1(X24,u1_struct_0(X21))|(~r2_lattice3(X21,X22,X24)|r1_orders_2(X21,esk4_2(X21,X22),X24))|~r1_yellow_0(X21,X22)|(~v5_orders_2(X21)|~l1_orders_2(X21)))))&((m1_subset_1(esk5_3(X21,X25,X26),u1_struct_0(X21))|~r2_lattice3(X21,X25,X26)|~m1_subset_1(X26,u1_struct_0(X21))|r1_yellow_0(X21,X25)|(~v5_orders_2(X21)|~l1_orders_2(X21)))&((r2_lattice3(X21,X25,esk5_3(X21,X25,X26))|~r2_lattice3(X21,X25,X26)|~m1_subset_1(X26,u1_struct_0(X21))|r1_yellow_0(X21,X25)|(~v5_orders_2(X21)|~l1_orders_2(X21)))&(~r1_orders_2(X21,X26,esk5_3(X21,X25,X26))|~r2_lattice3(X21,X25,X26)|~m1_subset_1(X26,u1_struct_0(X21))|r1_yellow_0(X21,X25)|(~v5_orders_2(X21)|~l1_orders_2(X21)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_yellow_0])])])])])])).
cnf(c_0_23, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk2_1(X1))|~v5_orders_2(X1)|~v2_yellow_0(X1)|~m1_subset_1(esk3_3(X1,esk2_1(X1),X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_14])).
fof(c_0_24, plain, ![X28, X29]:((r2_lattice3(X28,k1_xboole_0,X29)|~m1_subset_1(X29,u1_struct_0(X28))|~l1_orders_2(X28))&(r1_lattice3(X28,k1_xboole_0,X29)|~m1_subset_1(X29,u1_struct_0(X28))|~l1_orders_2(X28))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])).
cnf(c_0_25, plain, (r2_lattice3(X1,X2,esk5_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_yellow_0(X1)&l1_orders_2(X1))))=>(r2_yellow_0(X1,k1_xboole_0)&r1_yellow_0(X1,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t43_yellow_0])).
fof(c_0_28, plain, ![X31]:(v2_struct_0(X31)|~l1_struct_0(X31)|~v1_xboole_0(u1_struct_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_29, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk2_1(X1))|~v5_orders_2(X1)|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_19]), c_0_14])).
cnf(c_0_30, plain, (r1_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_31, plain, ![X32]:(~l1_orders_2(X32)|l1_struct_0(X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_32, plain, (r1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk5_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (r1_yellow_0(X1,X2)|r1_orders_2(X1,X3,esk5_3(X1,X2,X4))|~v5_orders_2(X1)|~r2_hidden(X3,X2)|~r2_lattice3(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_25]), c_0_26])).
fof(c_0_34, negated_conjecture, ((~v2_struct_0(esk6_0)&(v5_orders_2(esk6_0)&(v2_yellow_0(esk6_0)&l1_orders_2(esk6_0))))&(~r2_yellow_0(esk6_0,k1_xboole_0)|~r1_yellow_0(esk6_0,u1_struct_0(esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_27])])])])).
cnf(c_0_35, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,k1_xboole_0)|~v5_orders_2(X1)|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_14])).
cnf(c_0_37, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (r1_yellow_0(X1,X2)|~v5_orders_2(X1)|~r2_hidden(X3,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (v2_struct_0(X1)|r2_yellow_0(X1,k1_xboole_0)|~v5_orders_2(X1)|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_41, negated_conjecture, (v5_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, negated_conjecture, (v2_yellow_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (l1_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (r1_yellow_0(X1,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_yellow_0(X1)|~r2_hidden(esk2_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_13]), c_0_14])).
cnf(c_0_45, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk2_1(X1),u1_struct_0(X1))|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_18, c_0_14])).
cnf(c_0_46, negated_conjecture, (~r2_yellow_0(esk6_0,k1_xboole_0)|~r1_yellow_0(esk6_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, negated_conjecture, (r2_yellow_0(esk6_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41]), c_0_42]), c_0_43])])).
cnf(c_0_48, plain, (r1_yellow_0(X1,u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|~v5_orders_2(X1)|~v2_yellow_0(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, negated_conjecture, (~r1_yellow_0(esk6_0,u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47])])).
cnf(c_0_50, negated_conjecture, (v1_xboole_0(u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_41]), c_0_42]), c_0_43])]), c_0_49])).
cnf(c_0_51, negated_conjecture, (~l1_struct_0(esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_50]), c_0_39])).
cnf(c_0_52, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_37]), c_0_43])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 34
# Proof object formula steps           : 19
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 21
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 102
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 99
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 1
# Generated clauses                    : 69
# ...of the previous two non-trivial   : 63
# Contextual simplify-reflections      : 16
# Paramodulations                      : 67
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 62
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 54
# Current number of unprocessed clauses: 22
# ...number of literals in the above   : 183
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 1315
# Rec. Clause-clause subsumption calls : 116
# Non-unit clause-clause subsumptions  : 22
# Unit Clause-clause subsumption calls : 31
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5043

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28856 pages
