# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_enumset1, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t41_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t100_xboole_1)).
fof(t44_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t44_enumset1)).
fof(t42_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_tarski(X2,X3)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t42_enumset1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', commutativity_k2_xboole_0)).
fof(t55_enumset1, conjecture, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t55_enumset1)).
fof(t51_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k3_enumset1(X2,X3,X4,X5,X6)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t51_enumset1)).
fof(t47_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t47_enumset1)).
fof(t50_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t50_enumset1)).
fof(t48_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t48_enumset1)).
fof(t46_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', t46_enumset1)).
fof(l62_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X4,X5,X6)), file('data/MPTP/proofs/879c3cb9ae9f4f9aa93dc35183e1f94a.E_input', l62_enumset1)).
fof(c_0_13, plain, ![X49, X50]:k2_tarski(X49,X50)=k2_xboole_0(k1_tarski(X49),k1_tarski(X50)), inference(variable_rename,[status(thm)],[t41_enumset1])).
fof(c_0_14, plain, ![X13, X14]:k2_xboole_0(X13,X14)=k5_xboole_0(X13,k4_xboole_0(X14,X13)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_15, plain, ![X15, X16]:k4_xboole_0(X15,X16)=k5_xboole_0(X15,k3_xboole_0(X15,X16)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_16, plain, ![X55, X56, X57, X58]:k2_enumset1(X55,X56,X57,X58)=k2_xboole_0(k1_tarski(X55),k1_enumset1(X56,X57,X58)), inference(variable_rename,[status(thm)],[t44_enumset1])).
fof(c_0_17, plain, ![X22, X23, X24]:k1_enumset1(X22,X23,X24)=k2_xboole_0(k1_tarski(X22),k2_tarski(X23,X24)), inference(variable_rename,[status(thm)],[t42_enumset1])).
cnf(c_0_18, plain, (k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_21, plain, ![X36, X37]:k2_xboole_0(X36,X37)=k2_xboole_0(X37,X36), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_22, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k1_tarski(X6))), inference(assume_negation,[status(cth)],[t55_enumset1])).
fof(c_0_23, plain, ![X25, X26, X27, X28, X29, X30]:k4_enumset1(X25,X26,X27,X28,X29,X30)=k2_xboole_0(k1_tarski(X25),k3_enumset1(X26,X27,X28,X29,X30)), inference(variable_rename,[status(thm)],[t51_enumset1])).
fof(c_0_24, plain, ![X31, X32, X33, X34, X35]:k3_enumset1(X31,X32,X33,X34,X35)=k2_xboole_0(k1_tarski(X31),k2_enumset1(X32,X33,X34,X35)), inference(variable_rename,[status(thm)],[t47_enumset1])).
cnf(c_0_25, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_tarski(X1),k1_enumset1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_26, plain, ![X17, X18, X19, X20, X21]:k3_enumset1(X17,X18,X19,X20,X21)=k2_xboole_0(k2_enumset1(X17,X18,X19,X20),k1_tarski(X21)), inference(variable_rename,[status(thm)],[t50_enumset1])).
fof(c_0_27, plain, ![X38, X39, X40, X41, X42]:k3_enumset1(X38,X39,X40,X41,X42)=k2_xboole_0(k2_tarski(X38,X39),k1_enumset1(X40,X41,X42)), inference(variable_rename,[status(thm)],[t48_enumset1])).
cnf(c_0_28, plain, (k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (k2_tarski(X1,X2)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_30, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_31, plain, ![X51, X52, X53, X54]:k2_enumset1(X51,X52,X53,X54)=k2_xboole_0(k1_enumset1(X51,X52,X53),k1_tarski(X54)), inference(variable_rename,[status(thm)],[t46_enumset1])).
fof(c_0_32, negated_conjecture, k4_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)!=k2_xboole_0(k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k1_tarski(esk6_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
cnf(c_0_33, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_tarski(X1),k3_enumset1(X2,X3,X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_34, plain, ![X43, X44, X45, X46, X47, X48]:k4_enumset1(X43,X44,X45,X46,X47,X48)=k2_xboole_0(k1_enumset1(X43,X44,X45),k1_enumset1(X46,X47,X48)), inference(variable_rename,[status(thm)],[l62_enumset1])).
cnf(c_0_35, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_tarski(X1),k2_enumset1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (k2_enumset1(X1,X2,X3,X4)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_19]), c_0_20])).
cnf(c_0_37, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_tarski(X5))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k2_tarski(X1,X2),k1_enumset1(X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39, plain, (k1_enumset1(X1,X2,X3)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k3_xboole_0(k1_tarski(X3),k1_tarski(X2)))),k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_tarski(X3),k3_xboole_0(k1_tarski(X3),k1_tarski(X2)))),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_19]), c_0_20]), c_0_29]), c_0_29])).
cnf(c_0_40, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))=k5_xboole_0(X2,k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_19]), c_0_19]), c_0_20]), c_0_20])).
cnf(c_0_41, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, negated_conjecture, (k4_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)!=k2_xboole_0(k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k1_tarski(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k3_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_19]), c_0_20])).
cnf(c_0_44, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X4,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k5_xboole_0(k1_tarski(X1),k5_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_enumset1(X3,X4,X5),k3_xboole_0(k1_enumset1(X3,X4,X5),k1_tarski(X2)))),k3_xboole_0(k5_xboole_0(k1_tarski(X2),k5_xboole_0(k1_enumset1(X3,X4,X5),k3_xboole_0(k1_enumset1(X3,X4,X5),k1_tarski(X2)))),k1_tarski(X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_19]), c_0_20]), c_0_36]), c_0_36])).
cnf(c_0_46, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k5_xboole_0(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1)))),k5_xboole_0(k1_tarski(X5),k3_xboole_0(k1_tarski(X5),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_19]), c_0_20]), c_0_36]), c_0_36])).
cnf(c_0_47, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k5_xboole_0(k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1)))),k5_xboole_0(k1_enumset1(X3,X4,X5),k3_xboole_0(k1_enumset1(X3,X4,X5),k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_tarski(X2),k3_xboole_0(k1_tarski(X2),k1_tarski(X1)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_19]), c_0_20]), c_0_29]), c_0_29])).
cnf(c_0_48, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_39])).
cnf(c_0_49, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k1_enumset1(X2,X3,X4),k3_xboole_0(k1_enumset1(X2,X3,X4),k1_tarski(X1))))=k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(k1_tarski(X4),k3_xboole_0(k1_tarski(X4),k1_enumset1(X1,X2,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_19]), c_0_20]), c_0_36])).
cnf(c_0_50, negated_conjecture, (k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k3_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k3_xboole_0(k3_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tarski(esk1_0))))!=k5_xboole_0(k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k5_xboole_0(k1_tarski(esk6_0),k3_xboole_0(k1_tarski(esk6_0),k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_19]), c_0_20]), c_0_43])).
cnf(c_0_51, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k3_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k1_tarski(X1))))=k5_xboole_0(k1_enumset1(X1,X2,X3),k5_xboole_0(k1_enumset1(X4,X5,X6),k3_xboole_0(k1_enumset1(X4,X5,X6),k1_enumset1(X1,X2,X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_19]), c_0_20]), c_0_43])).
cnf(c_0_52, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X2,X3,X4,X5,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_45]), c_0_46])).
cnf(c_0_53, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X5,X4)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_47])).
cnf(c_0_54, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X2,X1,X3,X4,X5)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_40]), c_0_47])).
cnf(c_0_55, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X4,X1,X2,X3,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_40]), c_0_49]), c_0_49]), c_0_46])).
cnf(c_0_56, negated_conjecture, (k5_xboole_0(k1_tarski(esk6_0),k5_xboole_0(k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k3_xboole_0(k3_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0),k1_tarski(esk6_0))))!=k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k3_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k3_xboole_0(k3_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tarski(esk1_0))))), inference(rw,[status(thm)],[c_0_50, c_0_40])).
cnf(c_0_57, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k3_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k1_tarski(X1))))=k5_xboole_0(k1_tarski(X4),k5_xboole_0(k3_enumset1(X5,X6,X1,X2,X3),k3_xboole_0(k3_enumset1(X5,X6,X1,X2,X3),k1_tarski(X4))))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_51]), c_0_51])).
cnf(c_0_58, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X1,X5,X2,X3,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_40]), c_0_49]), c_0_49]), c_0_45])).
cnf(c_0_59, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X2,X3,X5,X4,X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_60, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X4,X3,X5)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_48]), c_0_46])).
cnf(c_0_61, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X1,X4,X2,X3,X5)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_62, negated_conjecture, (k5_xboole_0(k1_tarski(esk3_0),k5_xboole_0(k3_enumset1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),k3_xboole_0(k3_enumset1(esk1_0,esk2_0,esk4_0,esk5_0,esk6_0),k1_tarski(esk3_0))))!=k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k3_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k3_xboole_0(k3_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tarski(esk1_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57]), c_0_55]), c_0_58]), c_0_55]), c_0_58])).
cnf(c_0_63, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k3_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k1_tarski(X1))))=k5_xboole_0(k1_tarski(X4),k5_xboole_0(k3_enumset1(X6,X1,X3,X2,X5),k3_xboole_0(k3_enumset1(X6,X1,X3,X2,X5),k1_tarski(X4))))), inference(spm,[status(thm)],[c_0_57, c_0_59])).
cnf(c_0_64, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X1,X3,X2,X4,X5)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_65, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k3_enumset1(X3,X4,X1,X2,X5)), inference(spm,[status(thm)],[c_0_55, c_0_55])).
cnf(c_0_66, negated_conjecture, (k5_xboole_0(k1_tarski(esk2_0),k5_xboole_0(k3_enumset1(esk1_0,esk3_0,esk4_0,esk5_0,esk6_0),k3_xboole_0(k3_enumset1(esk1_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tarski(esk2_0))))!=k5_xboole_0(k1_tarski(esk1_0),k5_xboole_0(k3_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k3_xboole_0(k3_enumset1(esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tarski(esk1_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63]), c_0_54]), c_0_58]), c_0_54]), c_0_60]), c_0_64]), c_0_54]), c_0_58]), c_0_54]), c_0_60]), c_0_64])).
cnf(c_0_67, plain, (k5_xboole_0(k1_tarski(X1),k5_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k3_xboole_0(k3_enumset1(X2,X3,X4,X5,X6),k1_tarski(X1))))=k5_xboole_0(k1_tarski(X4),k5_xboole_0(k3_enumset1(X1,X2,X5,X6,X3),k3_xboole_0(k3_enumset1(X1,X2,X5,X6,X3),k1_tarski(X4))))), inference(spm,[status(thm)],[c_0_57, c_0_65])).
cnf(c_0_68, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_67]), c_0_58]), c_0_55]), c_0_53]), c_0_60]), c_0_58]), c_0_55]), c_0_53]), c_0_60])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 42
# Proof object formula steps           : 27
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 74
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 8
# Processed clauses                    : 1920
# ...of these trivial                  : 5
# ...subsumed                          : 1819
# ...remaining for further processing  : 96
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 13290
# ...of the previous two non-trivial   : 12688
# Contextual simplify-reflections      : 0
# Paramodulations                      : 13290
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 93
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 84
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 10776
# ...number of literals in the above   : 10776
# Current number of archived formulas  : 0
# Current number of archived clauses   : 8
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 2777
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6118
# BW rewrite match successes           : 6116
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 63399

# -------------------------------------------------
# User time                : 0.056 s
# System time              : 0.007 s
# Total time               : 0.063 s
# Maximum resident set size: 28660 pages
