# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t108_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(![X5, X6]:(r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X1,X2))<=>r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X3,X4)))=>k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', t108_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', d3_tarski)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', t2_xboole_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', l54_zfmisc_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', d5_xboole_0)).
fof(t103_zfmisc_1, axiom, ![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6))))))), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', t103_zfmisc_1)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', t3_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', d10_xboole_0)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/88164b524f5041f4adc329fab9731605.E_input', t2_tarski)).
fof(c_0_9, negated_conjecture, ~(![X1, X2, X3, X4]:(![X5, X6]:(r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X1,X2))<=>r2_hidden(k4_tarski(X5,X6),k2_zfmisc_1(X3,X4)))=>k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4))), inference(assume_negation,[status(cth)],[t108_zfmisc_1])).
fof(c_0_10, plain, ![X17, X18, X19, X20, X21]:((~r1_tarski(X17,X18)|(~r2_hidden(X19,X17)|r2_hidden(X19,X18)))&((r2_hidden(esk5_2(X20,X21),X20)|r1_tarski(X20,X21))&(~r2_hidden(esk5_2(X20,X21),X21)|r1_tarski(X20,X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_11, plain, ![X23]:r1_tarski(k1_xboole_0,X23), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_12, plain, ![X13, X14, X15, X16]:(((r2_hidden(X13,X15)|~r2_hidden(k4_tarski(X13,X14),k2_zfmisc_1(X15,X16)))&(r2_hidden(X14,X16)|~r2_hidden(k4_tarski(X13,X14),k2_zfmisc_1(X15,X16))))&(~r2_hidden(X13,X15)|~r2_hidden(X14,X16)|r2_hidden(k4_tarski(X13,X14),k2_zfmisc_1(X15,X16)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_13, negated_conjecture, ![X11, X12]:(((~r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(esk1_0,esk2_0))|r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(esk3_0,esk4_0)))&(~r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(esk3_0,esk4_0))|r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(esk1_0,esk2_0))))&k2_zfmisc_1(esk1_0,esk2_0)!=k2_zfmisc_1(esk3_0,esk4_0)), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
fof(c_0_14, plain, ![X24, X25, X26, X27, X28, X29, X30, X31]:((((r2_hidden(X27,X24)|~r2_hidden(X27,X26)|X26!=k4_xboole_0(X24,X25))&(~r2_hidden(X27,X25)|~r2_hidden(X27,X26)|X26!=k4_xboole_0(X24,X25)))&(~r2_hidden(X28,X24)|r2_hidden(X28,X25)|r2_hidden(X28,X26)|X26!=k4_xboole_0(X24,X25)))&((~r2_hidden(esk6_3(X29,X30,X31),X31)|(~r2_hidden(esk6_3(X29,X30,X31),X29)|r2_hidden(esk6_3(X29,X30,X31),X30))|X31=k4_xboole_0(X29,X30))&((r2_hidden(esk6_3(X29,X30,X31),X29)|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k4_xboole_0(X29,X30))&(~r2_hidden(esk6_3(X29,X30,X31),X30)|r2_hidden(esk6_3(X29,X30,X31),X31)|X31=k4_xboole_0(X29,X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_15, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, plain, ![X39, X40, X41, X42]:((r2_hidden(esk8_4(X39,X40,X41,X42),X40)|~r2_hidden(X42,X39)|~r1_tarski(X39,k2_zfmisc_1(X40,X41)))&((r2_hidden(esk9_4(X39,X40,X41,X42),X41)|~r2_hidden(X42,X39)|~r1_tarski(X39,k2_zfmisc_1(X40,X41)))&(X42=k4_tarski(esk8_4(X39,X40,X41,X42),esk9_4(X39,X40,X41,X42))|~r2_hidden(X42,X39)|~r1_tarski(X39,k2_zfmisc_1(X40,X41))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t103_zfmisc_1])])])])])])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22, plain, (r2_hidden(esk9_4(X1,X2,X3,X4),X3)|~r2_hidden(X4,X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r2_hidden(esk9_4(X1,X2,k1_xboole_0,X3),X4)|~r1_tarski(X1,k2_zfmisc_1(X2,k1_xboole_0))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X2,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29, plain, (~r1_tarski(X1,k2_zfmisc_1(X2,k1_xboole_0))|~r2_hidden(X3,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_26])).
cnf(c_0_30, plain, (r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk6_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_31, plain, (r2_hidden(esk6_3(X1,X2,X3),X1)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_32, negated_conjecture, (r1_tarski(esk2_0,X1)|r2_hidden(X2,esk3_0)|~r2_hidden(X2,esk1_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_34, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_29, c_0_16])).
cnf(c_0_35, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk6_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_36, plain, ![X35]:(~r1_tarski(X35,k1_xboole_0)|X35=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_37, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_38, negated_conjecture, (r1_tarski(esk1_0,X1)|r1_tarski(esk2_0,X2)|r2_hidden(esk5_2(esk1_0,X1),esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_28])).
fof(c_0_39, plain, ![X33, X34]:(((r1_tarski(X33,X34)|X33!=X34)&(r1_tarski(X34,X33)|X33!=X34))&(~r1_tarski(X33,X34)|~r1_tarski(X34,X33)|X33=X34)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_40, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(k4_tarski(X2,X1),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_33, c_0_19])).
cnf(c_0_41, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_42, plain, (~epred2_0<=>![X1]:~r2_hidden(X1,esk3_0)), introduced(definition)).
cnf(c_0_43, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, negated_conjecture, (r1_tarski(esk1_0,esk3_0)|r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)), inference(spm,[status(thm)],[c_0_40, c_0_24])).
cnf(c_0_48, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_41, c_0_24])).
cnf(c_0_49, negated_conjecture, (epred2_0|~r2_hidden(X1,esk3_0)), inference(split_equiv,[status(thm)],[c_0_42])).
cnf(c_0_50, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (X1=k1_xboole_0|r2_hidden(esk6_3(X2,X2,X1),X1)), inference(rw,[status(thm)],[c_0_35, c_0_43])).
cnf(c_0_52, negated_conjecture, (esk2_0=k1_xboole_0|r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_54, plain, (r2_hidden(esk8_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_55, negated_conjecture, (r1_tarski(esk1_0,X1)|r2_hidden(X2,esk4_0)|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_47, c_0_28])).
cnf(c_0_56, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_18, c_0_48])).
cnf(c_0_57, negated_conjecture, (epred2_0|r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_49, c_0_28])).
cnf(c_0_58, negated_conjecture, (X1=esk2_0|r1_tarski(esk1_0,esk3_0)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_50, c_0_45])).
cnf(c_0_59, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)!=k2_zfmisc_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_60, negated_conjecture, (esk3_0=k1_xboole_0|epred2_0), inference(spm,[status(thm)],[c_0_49, c_0_51])).
cnf(c_0_61, negated_conjecture, (esk2_0=k1_xboole_0|esk3_0=esk1_0|~r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_50, c_0_52])).
cnf(c_0_62, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_29, c_0_53])).
cnf(c_0_63, plain, (~r1_tarski(X1,k2_zfmisc_1(k1_xboole_0,X2))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_34, c_0_54])).
cnf(c_0_64, negated_conjecture, (r1_tarski(esk2_0,X1)|r1_tarski(esk1_0,X2)|r2_hidden(esk5_2(esk2_0,X1),esk4_0)), inference(spm,[status(thm)],[c_0_55, c_0_28])).
cnf(c_0_65, negated_conjecture, (r1_tarski(esk4_0,X1)|r2_hidden(X2,esk1_0)|~r2_hidden(X2,esk3_0)), inference(spm,[status(thm)],[c_0_56, c_0_28])).
cnf(c_0_66, negated_conjecture, (X1=esk3_0|epred2_0|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_50, c_0_57])).
cnf(c_0_67, negated_conjecture, (esk3_0=esk2_0|epred2_0|r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_58, c_0_57])).
cnf(c_0_68, negated_conjecture, (epred2_0|k2_zfmisc_1(esk1_0,esk2_0)!=k2_zfmisc_1(k1_xboole_0,esk4_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_69, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0|epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_60]), c_0_16])])).
cnf(c_0_70, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_51])).
cnf(c_0_71, plain, (~r2_hidden(X1,k2_zfmisc_1(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_63, c_0_53])).
cnf(c_0_72, negated_conjecture, (r1_tarski(esk2_0,esk4_0)|r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_37, c_0_64])).
cnf(c_0_73, negated_conjecture, (r1_tarski(esk3_0,X1)|r1_tarski(esk4_0,X2)|r2_hidden(esk5_2(esk3_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_65, c_0_28])).
cnf(c_0_74, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk4_0)|~r2_hidden(X2,esk3_0)), inference(spm,[status(thm)],[c_0_33, c_0_48])).
cnf(c_0_75, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk6_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_31])).
cnf(c_0_76, negated_conjecture, (esk3_0=esk2_0|esk3_0=esk1_0|epred2_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_77, negated_conjecture, (esk3_0=esk1_0|esk2_0=k1_xboole_0|epred2_0), inference(spm,[status(thm)],[c_0_61, c_0_57])).
cnf(c_0_78, negated_conjecture, (esk1_0=k1_xboole_0|epred2_0|k2_zfmisc_1(k1_xboole_0,esk4_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])).
cnf(c_0_79, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_51])).
cnf(c_0_80, negated_conjecture, (X1=esk1_0|r1_tarski(esk2_0,esk4_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_50, c_0_72])).
cnf(c_0_81, negated_conjecture, (r1_tarski(esk3_0,esk1_0)|r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_37, c_0_73])).
cnf(c_0_82, negated_conjecture, (r1_tarski(esk4_0,X1)|r2_hidden(esk5_2(esk4_0,X1),esk2_0)|~r2_hidden(X2,esk3_0)), inference(spm,[status(thm)],[c_0_74, c_0_28])).
cnf(c_0_83, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_84, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_75])).
cnf(c_0_85, negated_conjecture, (esk3_0=esk1_0|epred2_0|esk2_0!=esk1_0), inference(ef,[status(thm)],[c_0_76])).
cnf(c_0_86, negated_conjecture, (esk2_0=k1_xboole_0|epred2_0|k2_zfmisc_1(esk1_0,esk4_0)!=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_59, c_0_77])).
cnf(c_0_87, negated_conjecture, (esk1_0=k1_xboole_0|epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_79])])).
cnf(c_0_88, negated_conjecture, (esk1_0=k1_xboole_0|r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_80, c_0_16])).
cnf(c_0_89, negated_conjecture, (esk4_0=k1_xboole_0|r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_44, c_0_81])).
cnf(c_0_90, negated_conjecture, (k4_xboole_0(esk3_0,X1)=esk3_0|r1_tarski(esk4_0,X2)|r2_hidden(esk5_2(esk4_0,X2),esk2_0)), inference(spm,[status(thm)],[c_0_82, c_0_75])).
fof(c_0_91, plain, (~epred1_0<=>![X2]:(~r2_hidden(X2,esk4_0)|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0)), introduced(definition)).
cnf(c_0_92, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_34])).
cnf(c_0_93, negated_conjecture, (epred2_0|k2_zfmisc_1(esk1_0,esk4_0)!=k2_zfmisc_1(esk1_0,esk2_0)|esk2_0!=esk1_0), inference(spm,[status(thm)],[c_0_59, c_0_85])).
cnf(c_0_94, negated_conjecture, (esk2_0=k1_xboole_0|epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_79]), c_0_79])])).
cnf(c_0_95, negated_conjecture, (esk1_0=k1_xboole_0|esk4_0=esk2_0|~r1_tarski(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_50, c_0_88])).
cnf(c_0_96, negated_conjecture, (esk4_0=k1_xboole_0|esk3_0=esk1_0|esk2_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_89])).
cnf(c_0_97, negated_conjecture, (k4_xboole_0(esk3_0,X1)=esk3_0|r1_tarski(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_37, c_0_90])).
cnf(c_0_98, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_15, c_0_52])).
cnf(c_0_99, negated_conjecture, (~epred2_0|~epred1_0), inference(apply_def,[status(thm)],[inference(apply_def,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_48]), c_0_91]), c_0_42])).
cnf(c_0_100, negated_conjecture, (epred2_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_87]), c_0_79]), c_0_79])]), c_0_94])).
cnf(c_0_101, negated_conjecture, (esk3_0=esk1_0|esk2_0=k1_xboole_0|esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_16])])).
cnf(c_0_102, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_97])).
cnf(c_0_103, negated_conjecture, (esk2_0=k1_xboole_0|esk3_0!=k1_xboole_0|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_92, c_0_98])).
cnf(c_0_104, negated_conjecture, (epred1_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0|~r2_hidden(X1,esk4_0)), inference(split_equiv,[status(thm)],[c_0_91])).
cnf(c_0_105, negated_conjecture, (~epred1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_106, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0|k2_zfmisc_1(esk1_0,esk4_0)!=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_59, c_0_101])).
cnf(c_0_107, negated_conjecture, (esk3_0=k1_xboole_0|esk4_0=esk2_0|esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_95, c_0_102])).
cnf(c_0_108, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0|esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_103, c_0_51])).
cnf(c_0_109, negated_conjecture, (r1_tarski(esk4_0,X1)|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_28]), c_0_105])).
cnf(c_0_110, negated_conjecture, (esk2_0=k1_xboole_0|esk1_0=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_108])).
cnf(c_0_111, negated_conjecture, (esk1_0=k1_xboole_0|r1_tarski(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_70])])).
fof(c_0_112, plain, ![X36, X37]:((~r2_hidden(esk7_2(X36,X37),X36)|~r2_hidden(esk7_2(X36,X37),X37)|X36=X37)&(r2_hidden(esk7_2(X36,X37),X36)|r2_hidden(esk7_2(X36,X37),X37)|X36=X37)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
cnf(c_0_113, negated_conjecture, (esk4_0=k1_xboole_0|esk3_0=esk1_0|~r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_50, c_0_89])).
cnf(c_0_114, plain, (r1_tarski(X1,X2)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_92, c_0_28])).
cnf(c_0_115, negated_conjecture, (esk4_0=esk2_0|esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_95, c_0_111])).
cnf(c_0_116, plain, (r2_hidden(esk7_2(X1,X2),X1)|r2_hidden(esk7_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_117, negated_conjecture, (esk3_0=esk1_0|esk4_0=k1_xboole_0|esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_118, negated_conjecture, (esk1_0=k1_xboole_0|k2_zfmisc_1(esk3_0,esk2_0)!=k2_zfmisc_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_59, c_0_115])).
cnf(c_0_119, plain, (X1=k1_xboole_0|r2_hidden(esk7_2(X1,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_34, c_0_116])).
cnf(c_0_120, negated_conjecture, (esk3_0=esk1_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0|esk1_0!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_117]), c_0_70])).
cnf(c_0_121, negated_conjecture, (esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_110]), c_0_70]), c_0_70])])).
cnf(c_0_122, negated_conjecture, (esk4_0=k1_xboole_0|k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_119]), c_0_105])).
cnf(c_0_123, negated_conjecture, (esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_121]), c_0_121]), c_0_79]), c_0_121])])).
cnf(c_0_124, negated_conjecture, (esk4_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_121]), c_0_79])])).
cnf(c_0_125, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_121]), c_0_79]), c_0_123]), c_0_124]), c_0_70])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 126
# Proof object clause steps            : 105
# Proof object formula steps           : 21
# Proof object conjectures             : 72
# Proof object clause conjectures      : 69
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 76
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 965
# ...of these trivial                  : 14
# ...subsumed                          : 498
# ...remaining for further processing  : 452
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 73
# Backward-rewritten                   : 261
# Generated clauses                    : 10573
# ...of the previous two non-trivial   : 9770
# Contextual simplify-reflections      : 8
# Paramodulations                      : 10494
# Factorizations                       : 60
# Equation resolutions                 : 16
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 91
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 78
# Current number of unprocessed clauses: 8645
# ...number of literals in the above   : 41590
# Current number of archived formulas  : 0
# Current number of archived clauses   : 358
# Clause-clause subsumption calls (NU) : 18435
# Rec. Clause-clause subsumption calls : 8864
# Non-unit clause-clause subsumptions  : 544
# Unit Clause-clause subsumption calls : 571
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 52
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 138040

# -------------------------------------------------
# User time                : 0.100 s
# System time              : 0.003 s
# Total time               : 0.103 s
# Maximum resident set size: 28604 pages
