# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/88dfcd71389e4ed3b8eb5ed9203bdcf7.E_input', d1_xboole_0)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/88dfcd71389e4ed3b8eb5ed9203bdcf7.E_input', t6_xboole_0)).
fof(t6_relset_1, conjecture, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>~((r2_hidden(X1,X4)&![X5, X6]:~((X1=k4_tarski(X5,X6)&(r2_hidden(X5,X2)&r2_hidden(X6,X3))))))), file('data/MPTP/proofs/88dfcd71389e4ed3b8eb5ed9203bdcf7.E_input', t6_relset_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/88dfcd71389e4ed3b8eb5ed9203bdcf7.E_input', d2_zfmisc_1)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/88dfcd71389e4ed3b8eb5ed9203bdcf7.E_input', d8_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/88dfcd71389e4ed3b8eb5ed9203bdcf7.E_input', t3_subset)).
fof(t65_subset_1, axiom, ![X1, X2, X3, X4]:~((r2_hidden(X4,X3)&(r1_tarski(X3,k2_zfmisc_1(X1,X2))&![X5]:(m1_subset_1(X5,X1)=>![X6]:(m1_subset_1(X6,X2)=>~(X4=k4_tarski(X5,X6))))))), file('data/MPTP/proofs/88dfcd71389e4ed3b8eb5ed9203bdcf7.E_input', t65_subset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/88dfcd71389e4ed3b8eb5ed9203bdcf7.E_input', t2_subset)).
fof(c_0_8, plain, ![X43, X44, X45]:((~v1_xboole_0(X43)|~r2_hidden(X44,X43))&(r2_hidden(esk13_1(X45),X45)|v1_xboole_0(X45))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_9, plain, ![X17, X18]:((r2_hidden(esk5_2(X17,X18),X18)|~r2_xboole_0(X17,X18))&(~r2_hidden(esk5_2(X17,X18),X17)|~r2_xboole_0(X17,X18))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
fof(c_0_10, negated_conjecture, ~(![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>~((r2_hidden(X1,X4)&![X5, X6]:~((X1=k4_tarski(X5,X6)&(r2_hidden(X5,X2)&r2_hidden(X6,X3)))))))), inference(assume_negation,[status(cth)],[t6_relset_1])).
fof(c_0_11, plain, ![X20, X21, X22, X23, X26, X27, X28, X29, X30, X31, X33, X34]:((((r2_hidden(esk6_4(X20,X21,X22,X23),X20)|~r2_hidden(X23,X22)|X22!=k2_zfmisc_1(X20,X21))&((r2_hidden(esk7_4(X20,X21,X22,X23),X21)|~r2_hidden(X23,X22)|X22!=k2_zfmisc_1(X20,X21))&(X23=k4_tarski(esk6_4(X20,X21,X22,X23),esk7_4(X20,X21,X22,X23))|~r2_hidden(X23,X22)|X22!=k2_zfmisc_1(X20,X21))))&(~r2_hidden(X27,X20)|(~r2_hidden(X28,X21)|X26!=k4_tarski(X27,X28))|r2_hidden(X26,X22)|X22!=k2_zfmisc_1(X20,X21)))&((~r2_hidden(esk8_3(X29,X30,X31),X31)|(~r2_hidden(X33,X29)|(~r2_hidden(X34,X30)|esk8_3(X29,X30,X31)!=k4_tarski(X33,X34)))|X31=k2_zfmisc_1(X29,X30))&((r2_hidden(esk9_3(X29,X30,X31),X29)|r2_hidden(esk8_3(X29,X30,X31),X31)|X31=k2_zfmisc_1(X29,X30))&((r2_hidden(esk10_3(X29,X30,X31),X30)|r2_hidden(esk8_3(X29,X30,X31),X31)|X31=k2_zfmisc_1(X29,X30))&(esk8_3(X29,X30,X31)=k4_tarski(esk9_3(X29,X30,X31),esk10_3(X29,X30,X31))|r2_hidden(esk8_3(X29,X30,X31),X31)|X31=k2_zfmisc_1(X29,X30)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_12, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (r2_hidden(esk5_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, plain, ![X13, X14]:(((r1_tarski(X13,X14)|~r2_xboole_0(X13,X14))&(X13!=X14|~r2_xboole_0(X13,X14)))&(~r1_tarski(X13,X14)|X13=X14|r2_xboole_0(X13,X14))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
fof(c_0_15, plain, ![X47, X48]:((~m1_subset_1(X47,k1_zfmisc_1(X48))|r1_tarski(X47,X48))&(~r1_tarski(X47,X48)|m1_subset_1(X47,k1_zfmisc_1(X48)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_16, negated_conjecture, ![X11, X12]:(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))&(r2_hidden(esk1_0,esk4_0)&(esk1_0!=k4_tarski(X11,X12)|(~r2_hidden(X11,esk2_0)|~r2_hidden(X12,esk3_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])])).
cnf(c_0_17, plain, (r2_hidden(esk7_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (~v1_xboole_0(X1)|~r2_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (X1!=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_12, c_0_17])).
cnf(c_0_23, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24, plain, (X1=X2|~v1_xboole_0(X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, negated_conjecture, (r1_tarski(esk4_0,k2_zfmisc_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k2_zfmisc_1(X3,X1))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (r2_hidden(esk13_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_28, plain, (X1!=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_12, c_0_23])).
fof(c_0_29, plain, ![X37, X38, X39, X40]:((m1_subset_1(esk11_4(X37,X38,X39,X40),X37)|~r1_tarski(X39,k2_zfmisc_1(X37,X38))|~r2_hidden(X40,X39))&((m1_subset_1(esk12_4(X37,X38,X39,X40),X38)|~r1_tarski(X39,k2_zfmisc_1(X37,X38))|~r2_hidden(X40,X39))&(X40=k4_tarski(esk11_4(X37,X38,X39,X40),esk12_4(X37,X38,X39,X40))|~r1_tarski(X39,k2_zfmisc_1(X37,X38))|~r2_hidden(X40,X39)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_subset_1])])])])])).
fof(c_0_30, plain, ![X15, X16]:(~m1_subset_1(X15,X16)|(v1_xboole_0(X16)|r2_hidden(X15,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_31, negated_conjecture, (k2_zfmisc_1(esk2_0,esk3_0)=esk4_0|~v1_xboole_0(k2_zfmisc_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk1_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_34, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k2_zfmisc_1(X1,X3))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (esk1_0!=k4_tarski(X1,X2)|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_36, plain, (X1=k4_tarski(esk11_4(X2,X3,X4,X1),esk12_4(X2,X3,X4,X1))|~r1_tarski(X4,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (m1_subset_1(esk12_4(X1,X2,X3,X4),X2)|~r1_tarski(X3,k2_zfmisc_1(X1,X2))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, negated_conjecture, (k2_zfmisc_1(esk2_0,esk3_0)=esk4_0|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_40, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_12, c_0_33])).
cnf(c_0_41, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_34, c_0_27])).
cnf(c_0_42, negated_conjecture, (X1!=esk1_0|~r1_tarski(X2,k2_zfmisc_1(X3,X4))|~r2_hidden(esk12_4(X3,X4,X2,X1),esk3_0)|~r2_hidden(esk11_4(X3,X4,X2,X1),esk2_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, plain, (v1_xboole_0(X1)|r2_hidden(esk12_4(X2,X1,X3,X4),X1)|~r1_tarski(X3,k2_zfmisc_1(X2,X1))|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_39]), c_0_40])).
cnf(c_0_45, plain, (m1_subset_1(esk11_4(X1,X2,X3,X4),X1)|~r1_tarski(X3,k2_zfmisc_1(X1,X2))|~r2_hidden(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_46, negated_conjecture, (k2_zfmisc_1(esk2_0,esk3_0)=esk4_0|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_31, c_0_41])).
cnf(c_0_47, negated_conjecture, (X1!=esk1_0|~r1_tarski(X2,k2_zfmisc_1(X3,esk3_0))|~r2_hidden(esk11_4(X3,esk3_0,X2,X1),esk2_0)|~r2_hidden(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_48, plain, (v1_xboole_0(X1)|r2_hidden(esk11_4(X1,X2,X3,X4),X1)|~r1_tarski(X3,k2_zfmisc_1(X1,X2))|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_37, c_0_45])).
cnf(c_0_49, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_46]), c_0_40])).
cnf(c_0_50, negated_conjecture, (X1!=esk1_0|~r1_tarski(X2,k2_zfmisc_1(esk2_0,esk3_0))|~r2_hidden(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_51, negated_conjecture, (X1!=esk1_0|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_50, c_0_25])).
cnf(c_0_52, negated_conjecture, ($false), inference(spm,[status(thm)],[c_0_51, c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 53
# Proof object clause steps            : 36
# Proof object formula steps           : 17
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 857
# ...of these trivial                  : 0
# ...subsumed                          : 651
# ...remaining for further processing  : 206
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 0
# Generated clauses                    : 4508
# ...of the previous two non-trivial   : 4498
# Contextual simplify-reflections      : 0
# Paramodulations                      : 4488
# Factorizations                       : 0
# Equation resolutions                 : 20
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 174
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 167
# Current number of unprocessed clauses: 3666
# ...number of literals in the above   : 15525
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 10802
# Rec. Clause-clause subsumption calls : 8206
# Non-unit clause-clause subsumptions  : 472
# Unit Clause-clause subsumption calls : 92
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 45589

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.010 s
# Total time               : 0.040 s
# Maximum resident set size: 28644 pages
