# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t63_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&m1_yellow_0(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X5=X4=>((r1_lattice3(X2,X3,X5)=>r1_lattice3(X1,X3,X4))&(r2_lattice3(X2,X3,X5)=>r2_lattice3(X1,X3,X4))))))))), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', t63_yellow_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', t3_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', d3_tarski)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', d9_lattice3)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', d1_xboole_0)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', dt_m1_yellow_0)).
fof(t60_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X5=X3&(X6=X4&r1_orders_2(X2,X5,X6)))=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', t60_yellow_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/88fafac3cb4c4b9c958a607e221a6355.E_input', d8_lattice3)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((~(v2_struct_0(X2))&m1_yellow_0(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X5=X4=>((r1_lattice3(X2,X3,X5)=>r1_lattice3(X1,X3,X4))&(r2_lattice3(X2,X3,X5)=>r2_lattice3(X1,X3,X4)))))))))), inference(assume_negation,[status(cth)],[t63_yellow_0])).
fof(c_0_10, plain, ![X42, X43]:((~m1_subset_1(X42,k1_zfmisc_1(X43))|r1_tarski(X42,X43))&(~r1_tarski(X42,X43)|m1_subset_1(X42,k1_zfmisc_1(X43)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~v2_struct_0(esk2_0)&m1_yellow_0(esk2_0,esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(m1_subset_1(esk5_0,u1_struct_0(esk2_0))&(esk5_0=esk4_0&(((r2_lattice3(esk2_0,esk3_0,esk5_0)|r1_lattice3(esk2_0,esk3_0,esk5_0))&(~r2_lattice3(esk1_0,esk3_0,esk4_0)|r1_lattice3(esk2_0,esk3_0,esk5_0)))&((r2_lattice3(esk2_0,esk3_0,esk5_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0))&(~r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])).
fof(c_0_12, plain, ![X22, X23, X24, X25, X26]:((~r1_tarski(X22,X23)|(~r2_hidden(X24,X22)|r2_hidden(X24,X23)))&((r2_hidden(esk8_2(X25,X26),X25)|r1_tarski(X25,X26))&(~r2_hidden(esk8_2(X25,X26),X26)|r1_tarski(X25,X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_13, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X12, X13, X14, X15]:((~r2_lattice3(X12,X13,X14)|(~m1_subset_1(X15,u1_struct_0(X12))|(~r2_hidden(X15,X13)|r1_orders_2(X12,X15,X14)))|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&((m1_subset_1(esk6_3(X12,X13,X14),u1_struct_0(X12))|r2_lattice3(X12,X13,X14)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&((r2_hidden(esk6_3(X12,X13,X14),X13)|r2_lattice3(X12,X13,X14)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))&(~r1_orders_2(X12,esk6_3(X12,X13,X14),X14)|r2_lattice3(X12,X13,X14)|~m1_subset_1(X14,u1_struct_0(X12))|~l1_orders_2(X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_16, plain, ![X40, X41]:(((~m1_subset_1(X41,X40)|r2_hidden(X41,X40)|v1_xboole_0(X40))&(~r2_hidden(X41,X40)|m1_subset_1(X41,X40)|v1_xboole_0(X40)))&((~m1_subset_1(X41,X40)|v1_xboole_0(X41)|~v1_xboole_0(X40))&(~v1_xboole_0(X41)|m1_subset_1(X41,X40)|~v1_xboole_0(X40)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_17, plain, ![X28, X29, X30]:((~v1_xboole_0(X28)|~r2_hidden(X29,X28))&(r2_hidden(esk9_1(X30),X30)|v1_xboole_0(X30))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_18, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (r1_tarski(esk3_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_23, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, negated_conjecture, (r2_hidden(esk6_3(esk1_0,X1,esk4_0),X1)|r2_lattice3(esk1_0,X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
fof(c_0_27, plain, ![X32, X33]:(~l1_orders_2(X32)|(~m1_yellow_0(X33,X32)|l1_orders_2(X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
cnf(c_0_28, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))|r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_31, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_32, plain, ![X34, X35, X36, X37, X38, X39]:(~l1_orders_2(X34)|(~m1_yellow_0(X35,X34)|(~m1_subset_1(X36,u1_struct_0(X34))|(~m1_subset_1(X37,u1_struct_0(X34))|(~m1_subset_1(X38,u1_struct_0(X35))|(~m1_subset_1(X39,u1_struct_0(X35))|(X38!=X36|(X39!=X37|~r1_orders_2(X35,X38,X39))|r1_orders_2(X34,X36,X37)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_yellow_0])])])).
cnf(c_0_33, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_34, negated_conjecture, (r2_lattice3(esk1_0,esk3_0,esk4_0)|m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_22])])).
fof(c_0_36, plain, ![X17, X18, X19, X20]:((~r1_lattice3(X17,X18,X19)|(~m1_subset_1(X20,u1_struct_0(X17))|(~r2_hidden(X20,X18)|r1_orders_2(X17,X19,X20)))|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&((m1_subset_1(esk7_3(X17,X18,X19),u1_struct_0(X17))|r1_lattice3(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&((r2_hidden(esk7_3(X17,X18,X19),X18)|r1_lattice3(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))&(~r1_orders_2(X17,X19,esk7_3(X17,X18,X19))|r1_lattice3(X17,X18,X19)|~m1_subset_1(X19,u1_struct_0(X17))|~l1_orders_2(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_37, plain, (r1_orders_2(X1,X3,X4)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|X5!=X3|X6!=X4|~r1_orders_2(X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_39, negated_conjecture, (esk5_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_40, negated_conjecture, (r1_orders_2(esk2_0,esk6_3(esk1_0,esk3_0,esk4_0),X1)|r2_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk6_3(esk1_0,esk3_0,esk4_0),X2)|~r2_lattice3(esk2_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_41, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk5_0)|r1_lattice3(esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42, negated_conjecture, (r1_lattice3(esk2_0,esk3_0,esk5_0)|~r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_43, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (r1_orders_2(X1,X2,X3)|~r1_orders_2(X4,X2,X3)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_yellow_0(X4,X1)|~l1_orders_2(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_37])])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(esk2_0,esk6_3(esk1_0,esk3_0,esk4_0),X1)|r2_lattice3(esk1_0,esk3_0,esk4_0)|~r2_lattice3(esk2_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_40, c_0_26])).
cnf(c_0_47, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_39]), c_0_39])).
cnf(c_0_48, negated_conjecture, (r1_lattice3(esk2_0,esk3_0,esk4_0)|~r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_42, c_0_39])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk7_3(esk1_0,X1,esk4_0),X1)|r1_lattice3(esk1_0,X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_21]), c_0_22])])).
cnf(c_0_50, negated_conjecture, (r1_orders_2(X1,X2,esk4_0)|~r1_orders_2(esk2_0,X2,esk4_0)|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(esk4_0,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_yellow_0(esk2_0,X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_51, negated_conjecture, (r1_orders_2(esk2_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_45])]), c_0_48])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk7_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_25, c_0_49])).
cnf(c_0_53, negated_conjecture, (r1_orders_2(X1,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))|~m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(X1))|~m1_subset_1(esk4_0,u1_struct_0(X1))|~m1_yellow_0(esk2_0,X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_54, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_55, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_56, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)|m1_subset_1(esk7_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_52])).
cnf(c_0_57, negated_conjecture, (r1_orders_2(X1,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)|~m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(X1))|~m1_subset_1(esk4_0,u1_struct_0(X1))|~m1_yellow_0(esk2_0,X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_34]), c_0_48])).
cnf(c_0_58, negated_conjecture, (r2_lattice3(esk1_0,X1,esk4_0)|m1_subset_1(esk6_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_21]), c_0_22])])).
cnf(c_0_59, negated_conjecture, (r1_orders_2(esk2_0,X1,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r2_hidden(esk7_3(esk1_0,esk3_0,esk4_0),X2)|~r1_lattice3(esk2_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_35])])).
cnf(c_0_60, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_61, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)|r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_21]), c_0_31]), c_0_22])]), c_0_48])).
cnf(c_0_62, negated_conjecture, (r1_orders_2(esk2_0,X1,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk2_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_59, c_0_49])).
cnf(c_0_63, negated_conjecture, (r1_lattice3(esk2_0,esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_21]), c_0_22])]), c_0_48])).
cnf(c_0_64, negated_conjecture, (r1_orders_2(X1,X2,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)|~r1_orders_2(esk2_0,X2,esk7_3(esk1_0,esk3_0,esk4_0))|~m1_subset_1(esk7_3(esk1_0,esk3_0,esk4_0),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_yellow_0(esk2_0,X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_44, c_0_56])).
cnf(c_0_65, negated_conjecture, (r1_orders_2(esk2_0,esk4_0,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_45])])).
cnf(c_0_66, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_67, negated_conjecture, (r1_orders_2(X1,esk4_0,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)|~m1_subset_1(esk7_3(esk1_0,esk3_0,esk4_0),u1_struct_0(X1))|~m1_subset_1(esk4_0,u1_struct_0(X1))|~m1_yellow_0(esk2_0,X1)|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_45])])).
cnf(c_0_68, negated_conjecture, (r1_lattice3(esk1_0,X1,esk4_0)|m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_21]), c_0_22])])).
cnf(c_0_69, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk5_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_70, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk7_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_71, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk7_3(esk1_0,esk3_0,esk4_0))|r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_21]), c_0_31]), c_0_22])])).
cnf(c_0_72, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_69, c_0_39])).
cnf(c_0_73, negated_conjecture, (r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_21]), c_0_22])])).
cnf(c_0_74, negated_conjecture, (~r2_lattice3(esk1_0,esk3_0,esk4_0)|~r1_lattice3(esk1_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_75, negated_conjecture, (r2_lattice3(esk2_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73])])).
cnf(c_0_76, negated_conjecture, (~r2_lattice3(esk1_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_73])])).
cnf(c_0_77, negated_conjecture, (r1_orders_2(esk2_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_75]), c_0_45])]), c_0_76])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[c_0_34, c_0_76])).
cnf(c_0_79, negated_conjecture, (r1_orders_2(X1,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_subset_1(esk6_3(esk1_0,esk3_0,esk4_0),u1_struct_0(X1))|~m1_subset_1(esk4_0,u1_struct_0(X1))|~m1_yellow_0(esk2_0,X1)|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_77]), c_0_78])])).
cnf(c_0_80, negated_conjecture, (r1_orders_2(esk1_0,esk6_3(esk1_0,esk3_0,esk4_0),esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_58]), c_0_21]), c_0_31]), c_0_22])]), c_0_76])).
cnf(c_0_81, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_80]), c_0_21]), c_0_22])]), c_0_76]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 82
# Proof object clause steps            : 63
# Proof object formula steps           : 19
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 65
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 33
# Processed clauses                    : 3026
# ...of these trivial                  : 3
# ...subsumed                          : 1231
# ...remaining for further processing  : 1792
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 93
# Backward-rewritten                   : 133
# Generated clauses                    : 19557
# ...of the previous two non-trivial   : 19142
# Contextual simplify-reflections      : 27
# Paramodulations                      : 19518
# Factorizations                       : 2
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1496
#    Positive orientable unit clauses  : 24
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 1467
# Current number of unprocessed clauses: 16106
# ...number of literals in the above   : 85835
# Current number of archived formulas  : 0
# Current number of archived clauses   : 295
# Clause-clause subsumption calls (NU) : 411214
# Rec. Clause-clause subsumption calls : 55177
# Non-unit clause-clause subsumptions  : 1151
# Unit Clause-clause subsumption calls : 3858
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 69
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 598180

# -------------------------------------------------
# User time                : 0.378 s
# System time              : 0.007 s
# Total time               : 0.385 s
# Maximum resident set size: 28788 pages
