# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', t3_yellow_1)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', t1_yellow_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', redefinition_r3_orders_2)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', dt_k2_yellow_1)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', fc5_yellow_1)).
fof(fc6_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1)))), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', fc6_yellow_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', t2_xboole_1)).
fof(t13_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k1_xboole_0,X1)=>k3_yellow_0(k2_yellow_1(X1))=k1_xboole_0)), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', t13_yellow_1)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', t30_yellow_0)).
fof(d11_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', d11_yellow_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', t1_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', d1_xboole_0)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', d9_lattice3)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/89263e02dfa24dd1924fe580fc00fe6b.E_input', fc1_xboole_0)).
fof(c_0_14, plain, ![X5, X6, X7]:((~r3_orders_2(k2_yellow_1(X5),X6,X7)|r1_tarski(X6,X7)|~m1_subset_1(X7,u1_struct_0(k2_yellow_1(X5)))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X5)))|v1_xboole_0(X5))&(~r1_tarski(X6,X7)|r3_orders_2(k2_yellow_1(X5),X6,X7)|~m1_subset_1(X7,u1_struct_0(k2_yellow_1(X5)))|~m1_subset_1(X6,u1_struct_0(k2_yellow_1(X5)))|v1_xboole_0(X5))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])).
fof(c_0_15, plain, ![X8]:(u1_struct_0(k2_yellow_1(X8))=X8&u1_orders_2(k2_yellow_1(X8))=k1_yellow_1(X8)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
fof(c_0_16, plain, ![X9, X10, X11]:((~r3_orders_2(X9,X10,X11)|r1_orders_2(X9,X10,X11)|(v2_struct_0(X9)|(~v3_orders_2(X9)|~l1_orders_2(X9))|(~m1_subset_1(X10,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9)))))&(~r1_orders_2(X9,X10,X11)|r3_orders_2(X9,X10,X11)|(v2_struct_0(X9)|(~v3_orders_2(X9)|~l1_orders_2(X9))|(~m1_subset_1(X10,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_17, plain, (r3_orders_2(k2_yellow_1(X3),X1,X2)|v1_xboole_0(X3)|~r1_tarski(X1,X2)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X3)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X13]:(v1_orders_2(k2_yellow_1(X13))&l1_orders_2(k2_yellow_1(X13))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_20, plain, ![X12]:(v1_orders_2(k2_yellow_1(X12))&(v3_orders_2(k2_yellow_1(X12))&(v4_orders_2(k2_yellow_1(X12))&v5_orders_2(k2_yellow_1(X12))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_21, plain, ![X14]:((~v2_struct_0(k2_yellow_1(X14))|v1_xboole_0(X14))&(v1_orders_2(k2_yellow_1(X14))|v1_xboole_0(X14))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
cnf(c_0_22, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r3_orders_2(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_18])).
cnf(c_0_24, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_27, plain, ![X15]:r1_tarski(k1_xboole_0,X15), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_28, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k1_xboole_0,X1)=>k3_yellow_0(k2_yellow_1(X1))=k1_xboole_0))), inference(assume_negation,[status(cth)],[t13_yellow_1])).
fof(c_0_29, plain, ![X17, X18, X19, X20, X21]:(((r2_lattice3(X17,X19,X18)|(X18!=k1_yellow_0(X17,X19)|~r1_yellow_0(X17,X19))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(~m1_subset_1(X20,u1_struct_0(X17))|(~r2_lattice3(X17,X19,X20)|r1_orders_2(X17,X18,X20))|(X18!=k1_yellow_0(X17,X19)|~r1_yellow_0(X17,X19))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17))))&(((X18=k1_yellow_0(X17,X21)|(m1_subset_1(esk2_3(X17,X18,X21),u1_struct_0(X17))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(r1_yellow_0(X17,X21)|(m1_subset_1(esk2_3(X17,X18,X21),u1_struct_0(X17))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17))))&(((X18=k1_yellow_0(X17,X21)|(r2_lattice3(X17,X21,esk2_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(r1_yellow_0(X17,X21)|(r2_lattice3(X17,X21,esk2_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17))))&((X18=k1_yellow_0(X17,X21)|(~r1_orders_2(X17,X18,esk2_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(r1_yellow_0(X17,X21)|(~r1_orders_2(X17,X18,esk2_3(X17,X18,X21))|~r2_lattice3(X17,X21,X18))|~m1_subset_1(X18,u1_struct_0(X17))|(~v5_orders_2(X17)|~l1_orders_2(X17))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_30, plain, (r1_orders_2(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25]), c_0_18]), c_0_18])]), c_0_26])).
cnf(c_0_31, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_32, negated_conjecture, (~v1_xboole_0(esk1_0)&(r2_hidden(k1_xboole_0,esk1_0)&k3_yellow_0(k2_yellow_1(esk1_0))!=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_28])])])])).
fof(c_0_33, plain, ![X23]:(~l1_orders_2(X23)|k3_yellow_0(X23)=k1_yellow_0(X23,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
cnf(c_0_34, plain, (X1=k1_yellow_0(X2,X3)|~r1_orders_2(X2,X1,esk2_3(X2,X1,X3))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (r1_orders_2(k2_yellow_1(X1),k1_xboole_0,X2)|v1_xboole_0(X1)|~m1_subset_1(k1_xboole_0,X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_37, plain, (X1=k1_yellow_0(X2,X3)|m1_subset_1(esk2_3(X2,X1,X3),u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_38, plain, ![X24, X25]:(~r2_hidden(X24,X25)|m1_subset_1(X24,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_39, plain, ![X26, X27, X28]:((~v1_xboole_0(X26)|~r2_hidden(X27,X26))&(r2_hidden(esk3_1(X28),X28)|v1_xboole_0(X28))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_40, plain, ![X30, X31, X32, X33]:((~r2_lattice3(X30,X31,X32)|(~m1_subset_1(X33,u1_struct_0(X30))|(~r2_hidden(X33,X31)|r1_orders_2(X30,X33,X32)))|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&((m1_subset_1(esk4_3(X30,X31,X32),u1_struct_0(X30))|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&((r2_hidden(esk4_3(X30,X31,X32),X31)|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&(~r1_orders_2(X30,esk4_3(X30,X31,X32),X32)|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_41, negated_conjecture, (k3_yellow_0(k2_yellow_1(esk1_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (k1_yellow_0(k2_yellow_1(X1),X2)=k1_xboole_0|v1_xboole_0(X1)|~r2_lattice3(k2_yellow_1(X1),X2,k1_xboole_0)|~m1_subset_1(esk2_3(k2_yellow_1(X1),k1_xboole_0,X2),X1)|~m1_subset_1(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_24]), c_0_18])])).
cnf(c_0_44, plain, (X1=k1_yellow_0(k2_yellow_1(X2),X3)|m1_subset_1(esk2_3(k2_yellow_1(X2),X1,X3),X2)|~r2_lattice3(k2_yellow_1(X2),X3,X1)|~m1_subset_1(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_18]), c_0_36]), c_0_24])])).
cnf(c_0_45, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, negated_conjecture, (r2_hidden(k1_xboole_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_47, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, negated_conjecture, (k1_yellow_0(k2_yellow_1(esk1_0),k1_xboole_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_24])])).
cnf(c_0_50, plain, (k1_yellow_0(k2_yellow_1(X1),X2)=k1_xboole_0|v1_xboole_0(X1)|~r2_lattice3(k2_yellow_1(X1),X2,k1_xboole_0)|~m1_subset_1(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(k1_xboole_0,esk1_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_53, plain, (r2_lattice3(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_54, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_55, negated_conjecture, (~r2_lattice3(k2_yellow_1(esk1_0),k1_xboole_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])]), c_0_52])).
cnf(c_0_56, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_57, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_24]), c_0_18]), c_0_51])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 30
# Proof object formula steps           : 28
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 34
# Processed clauses                    : 91
# ...of these trivial                  : 2
# ...subsumed                          : 3
# ...remaining for further processing  : 86
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 41
# ...of the previous two non-trivial   : 34
# Contextual simplify-reflections      : 2
# Paramodulations                      : 39
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 9
# ...number of literals in the above   : 68
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 789
# Rec. Clause-clause subsumption calls : 83
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3776

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28776 pages
