# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', d10_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t28_xboole_1)).
fof(t147_zfmisc_1, conjecture, ![X1, X2, X3]:(~(X1=X2)=>k4_xboole_0(k2_xboole_0(X3,k1_tarski(X1)),k1_tarski(X2))=k2_xboole_0(k4_xboole_0(X3,k1_tarski(X2)),k1_tarski(X1))), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t147_zfmisc_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t70_enumset1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t36_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t100_xboole_1)).
fof(l97_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', l97_xboole_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t69_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t72_enumset1)).
fof(t94_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t94_xboole_1)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t83_xboole_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t16_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', d7_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', commutativity_k3_xboole_0)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t91_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t5_boole)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', commutativity_k5_xboole_0)).
fof(l31_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', l31_zfmisc_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', d5_xboole_0)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', l27_zfmisc_1)).
fof(l35_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', l35_zfmisc_1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', l1_zfmisc_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', symmetry_r1_xboole_0)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t65_zfmisc_1)).
fof(t64_zfmisc_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))<=>(r2_hidden(X1,X2)&~(X1=X3))), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t64_zfmisc_1)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', t1_xboole_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/89c0cfb14adf474a802ed8c7e9fecaa3.E_input', d1_tarski)).
fof(c_0_29, plain, ![X11, X12]:(((r1_tarski(X11,X12)|X11!=X12)&(r1_tarski(X12,X11)|X11!=X12))&(~r1_tarski(X11,X12)|~r1_tarski(X12,X11)|X11=X12)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_30, plain, ![X58, X59]:(~r1_tarski(X58,X59)|k3_xboole_0(X58,X59)=X58), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_31, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_32, negated_conjecture, ~(![X1, X2, X3]:(~(X1=X2)=>k4_xboole_0(k2_xboole_0(X3,k1_tarski(X1)),k1_tarski(X2))=k2_xboole_0(k4_xboole_0(X3,k1_tarski(X2)),k1_tarski(X1)))), inference(assume_negation,[status(cth)],[t147_zfmisc_1])).
fof(c_0_33, plain, ![X23, X24]:k3_tarski(k2_tarski(X23,X24))=k2_xboole_0(X23,X24), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_34, plain, ![X56, X57]:k1_enumset1(X56,X56,X57)=k2_tarski(X56,X57), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_35, plain, ![X41, X42]:r1_tarski(k4_xboole_0(X41,X42),X41), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_36, plain, ![X79, X80]:k4_xboole_0(X79,X80)=k5_xboole_0(X79,k3_xboole_0(X79,X80)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_37, plain, ![X39, X40]:r1_xboole_0(k3_xboole_0(X39,X40),k5_xboole_0(X39,X40)), inference(variable_rename,[status(thm)],[l97_xboole_1])).
cnf(c_0_38, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_31])).
fof(c_0_40, negated_conjecture, (esk1_0!=esk2_0&k4_xboole_0(k2_xboole_0(esk3_0,k1_tarski(esk1_0)),k1_tarski(esk2_0))!=k2_xboole_0(k4_xboole_0(esk3_0,k1_tarski(esk2_0)),k1_tarski(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_32])])])])).
fof(c_0_41, plain, ![X65]:k2_tarski(X65,X65)=k1_tarski(X65), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_42, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_44, plain, ![X34, X35, X36]:k2_enumset1(X34,X34,X35,X36)=k1_enumset1(X34,X35,X36), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_45, plain, ![X43, X44, X45, X46]:k3_enumset1(X43,X43,X44,X45,X46)=k2_enumset1(X43,X44,X45,X46), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_46, plain, ![X37, X38]:k2_xboole_0(X37,X38)=k5_xboole_0(k5_xboole_0(X37,X38),k3_xboole_0(X37,X38)), inference(variable_rename,[status(thm)],[t94_xboole_1])).
fof(c_0_47, plain, ![X66, X67]:((~r1_xboole_0(X66,X67)|k4_xboole_0(X66,X67)=X66)&(k4_xboole_0(X66,X67)!=X66|r1_xboole_0(X66,X67))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_48, plain, ![X14, X15, X16]:k3_xboole_0(k3_xboole_0(X14,X15),X16)=k3_xboole_0(X14,k3_xboole_0(X15,X16)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_49, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_50, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_51, plain, ![X32, X33]:((~r1_xboole_0(X32,X33)|k3_xboole_0(X32,X33)=k1_xboole_0)&(k3_xboole_0(X32,X33)!=k1_xboole_0|r1_xboole_0(X32,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_52, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_53, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_54, negated_conjecture, (k4_xboole_0(k2_xboole_0(esk3_0,k1_tarski(esk1_0)),k1_tarski(esk2_0))!=k2_xboole_0(k4_xboole_0(esk3_0,k1_tarski(esk2_0)),k1_tarski(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_55, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_56, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_57, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_58, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_59, plain, ![X47, X48]:k3_xboole_0(X47,X48)=k3_xboole_0(X48,X47), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_60, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_61, plain, ![X60, X61, X62]:k5_xboole_0(k5_xboole_0(X60,X61),X62)=k5_xboole_0(X60,k5_xboole_0(X61,X62)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_62, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_63, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_64, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_65, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_66, plain, (r1_xboole_0(X1,k5_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
fof(c_0_67, plain, ![X13]:k5_xboole_0(X13,k1_xboole_0)=X13, inference(variable_rename,[status(thm)],[t5_boole])).
fof(c_0_68, plain, ![X21, X22]:k5_xboole_0(X21,X22)=k5_xboole_0(X22,X21), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_69, negated_conjecture, (k5_xboole_0(k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),k3_xboole_0(k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))!=k3_tarski(k3_enumset1(k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55]), c_0_55]), c_0_55]), c_0_55]), c_0_43]), c_0_43]), c_0_43]), c_0_43]), c_0_56]), c_0_56]), c_0_50]), c_0_50]), c_0_50]), c_0_57]), c_0_57]), c_0_57]), c_0_57]), c_0_57]), c_0_57]), c_0_57]), c_0_57]), c_0_57]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_58]), c_0_58])).
cnf(c_0_70, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_71, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X2))=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_56]), c_0_57]), c_0_58])).
cnf(c_0_72, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_73, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_62, c_0_50])).
cnf(c_0_74, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_75, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_63, c_0_53])).
fof(c_0_76, plain, ![X54, X55]:(~r2_hidden(X54,X55)|k3_xboole_0(X55,k1_tarski(X54))=k1_tarski(X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l31_zfmisc_1])])).
cnf(c_0_77, plain, (r1_tarski(k5_xboole_0(X1,X1),X1)), inference(spm,[status(thm)],[c_0_64, c_0_53])).
cnf(c_0_78, plain, (k3_xboole_0(X1,k5_xboole_0(X1,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_79, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_80, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_81, plain, ![X70, X71, X72, X73, X74, X75, X76, X77]:((((r2_hidden(X73,X70)|~r2_hidden(X73,X72)|X72!=k4_xboole_0(X70,X71))&(~r2_hidden(X73,X71)|~r2_hidden(X73,X72)|X72!=k4_xboole_0(X70,X71)))&(~r2_hidden(X74,X70)|r2_hidden(X74,X71)|r2_hidden(X74,X72)|X72!=k4_xboole_0(X70,X71)))&((~r2_hidden(esk5_3(X75,X76,X77),X77)|(~r2_hidden(esk5_3(X75,X76,X77),X75)|r2_hidden(esk5_3(X75,X76,X77),X76))|X77=k4_xboole_0(X75,X76))&((r2_hidden(esk5_3(X75,X76,X77),X75)|r2_hidden(esk5_3(X75,X76,X77),X77)|X77=k4_xboole_0(X75,X76))&(~r2_hidden(esk5_3(X75,X76,X77),X76)|r2_hidden(esk5_3(X75,X76,X77),X77)|X77=k4_xboole_0(X75,X76))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_82, negated_conjecture, (k3_tarski(k3_enumset1(k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))!=k5_xboole_0(k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k3_tarski(k3_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))))), inference(rw,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_83, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,X2))=k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_84, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_85, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_75, c_0_70])).
fof(c_0_86, plain, ![X68, X69]:(r2_hidden(X68,X69)|r1_xboole_0(k1_tarski(X68),X69)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
cnf(c_0_87, plain, (k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_88, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_77]), c_0_70]), c_0_78])).
cnf(c_0_89, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_79, c_0_80])).
fof(c_0_90, plain, ![X63, X64]:((k4_xboole_0(k1_tarski(X63),X64)!=k1_xboole_0|r2_hidden(X63,X64))&(~r2_hidden(X63,X64)|k4_xboole_0(k1_tarski(X63),X64)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l35_zfmisc_1])])).
cnf(c_0_91, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_92, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)),k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_83]), c_0_70]), c_0_72]), c_0_83]), c_0_83]), c_0_72]), c_0_72])).
cnf(c_0_93, plain, (k5_xboole_0(X1,k3_xboole_0(X2,X1))=X1|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_94, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_95, plain, (k3_xboole_0(X2,k3_enumset1(X1,X1,X1,X1,X1))=k3_enumset1(X1,X1,X1,X1,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_55]), c_0_55]), c_0_43]), c_0_43]), c_0_57]), c_0_57]), c_0_58]), c_0_58])).
cnf(c_0_96, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_88]), c_0_89])).
cnf(c_0_97, plain, (k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_98, plain, (r2_hidden(X1,X2)|X3!=k5_xboole_0(X2,k3_xboole_0(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_91, c_0_50])).
fof(c_0_99, plain, ![X49, X50]:((~r1_tarski(k1_tarski(X49),X50)|r2_hidden(X49,X50))&(~r2_hidden(X49,X50)|r1_tarski(k1_tarski(X49),X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_100, plain, ![X19, X20]:(~r1_xboole_0(X19,X20)|r1_xboole_0(X20,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_101, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)),k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))))|k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_102, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X3))=k5_xboole_0(X3,k5_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_80, c_0_72])).
cnf(c_0_103, plain, (r2_hidden(X1,X2)|r1_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_55]), c_0_43]), c_0_57]), c_0_58])).
cnf(c_0_104, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))))))!=k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))|~r2_hidden(esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_95]), c_0_88]), c_0_79]), c_0_96]), c_0_70])).
cnf(c_0_105, plain, (k5_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),k3_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2))=k1_xboole_0|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_55]), c_0_43]), c_0_50]), c_0_57]), c_0_57]), c_0_58]), c_0_58])).
cnf(c_0_106, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(er,[status(thm)],[c_0_98])).
cnf(c_0_107, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_108, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_109, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),k3_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)),k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))))|k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102]), c_0_80])).
cnf(c_0_110, plain, (k3_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2)=k1_xboole_0|r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_65, c_0_103])).
cnf(c_0_111, negated_conjecture, (~r2_hidden(esk1_0,k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_79])]), c_0_106])).
cnf(c_0_112, plain, (r1_tarski(k3_enumset1(X1,X1,X1,X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_55]), c_0_43]), c_0_57]), c_0_58])).
cnf(c_0_113, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(spm,[status(thm)],[c_0_96, c_0_80])).
cnf(c_0_114, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_108, c_0_74])).
cnf(c_0_115, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)),k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))|k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_79]), c_0_111])).
cnf(c_0_116, plain, (k3_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2)=k3_enumset1(X1,X1,X1,X1,X1)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_38, c_0_112])).
cnf(c_0_117, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X3))=k5_xboole_0(X2,k5_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_80]), c_0_72])).
cnf(c_0_118, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_119, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,X3)))=k5_xboole_0(X2,X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_113]), c_0_72])).
cnf(c_0_120, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_114])).
fof(c_0_121, plain, ![X17, X18]:((k4_xboole_0(X17,k1_tarski(X18))!=X17|~r2_hidden(X18,X17))&(r2_hidden(X18,X17)|k4_xboole_0(X17,k1_tarski(X18))=X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
fof(c_0_122, plain, ![X51, X52, X53]:(((r2_hidden(X51,X52)|~r2_hidden(X51,k4_xboole_0(X52,k1_tarski(X53))))&(X51!=X53|~r2_hidden(X51,k4_xboole_0(X52,k1_tarski(X53)))))&(~r2_hidden(X51,X52)|X51=X53|r2_hidden(X51,k4_xboole_0(X52,k1_tarski(X53))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t64_zfmisc_1])])])])).
cnf(c_0_123, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))|k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))!=k1_xboole_0|~r2_hidden(esk2_0,k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_80]), c_0_117])).
cnf(c_0_124, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_118, c_0_50])).
cnf(c_0_125, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)),k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),k3_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_102]), c_0_80])).
cnf(c_0_126, plain, (k5_xboole_0(X1,k3_xboole_0(X2,X3))=X1|k3_xboole_0(X3,X2)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_113])).
cnf(c_0_127, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_128, plain, (X1=X3|r2_hidden(X1,k4_xboole_0(X2,k1_tarski(X3)))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_129, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))|k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))!=k1_xboole_0|~r2_hidden(esk2_0,k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_123, c_0_93])).
fof(c_0_130, plain, ![X8, X9, X10]:(((~r2_hidden(X8,X9)|~r2_hidden(X8,X10)|~r2_hidden(X8,k5_xboole_0(X9,X10)))&(r2_hidden(X8,X9)|r2_hidden(X8,X10)|~r2_hidden(X8,k5_xboole_0(X9,X10))))&((~r2_hidden(X8,X9)|r2_hidden(X8,X10)|r2_hidden(X8,k5_xboole_0(X9,X10)))&(~r2_hidden(X8,X10)|r2_hidden(X8,X9)|r2_hidden(X8,k5_xboole_0(X9,X10))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_131, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_124])).
cnf(c_0_132, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),k3_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))|k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_70])).
cnf(c_0_133, plain, (k5_xboole_0(X2,k3_xboole_0(X2,k3_enumset1(X1,X1,X1,X1,X1)))=X2|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_55]), c_0_43]), c_0_50]), c_0_57]), c_0_58])).
cnf(c_0_134, plain, (X1=X3|r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,k3_enumset1(X3,X3,X3,X3,X3))))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_55]), c_0_43]), c_0_50]), c_0_57]), c_0_58])).
cnf(c_0_135, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_136, negated_conjecture, (k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))!=k1_xboole_0|~r2_hidden(esk2_0,k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))|~r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_95]), c_0_80])])).
cnf(c_0_137, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_138, plain, (k3_xboole_0(X1,X2)!=k1_xboole_0|~r2_hidden(X3,X1)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_131, c_0_84])).
cnf(c_0_139, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))!=k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))|k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))))!=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_110]), c_0_79]), c_0_111])).
cnf(c_0_140, plain, (k3_xboole_0(X1,k3_enumset1(X2,X2,X2,X2,X2))=k1_xboole_0|r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_133]), c_0_88])).
cnf(c_0_141, negated_conjecture, (~r2_hidden(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_134]), c_0_135])).
cnf(c_0_142, negated_conjecture, (k3_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))!=k1_xboole_0|~r2_hidden(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_138])).
fof(c_0_143, plain, ![X25, X26, X27, X28, X29, X30]:(((~r2_hidden(X27,X26)|X27=X25|X26!=k1_tarski(X25))&(X28!=X25|r2_hidden(X28,X26)|X26!=k1_tarski(X25)))&((~r2_hidden(esk4_2(X29,X30),X30)|esk4_2(X29,X30)!=X29|X30=k1_tarski(X29))&(r2_hidden(esk4_2(X29,X30),X30)|esk4_2(X29,X30)=X29|X30=k1_tarski(X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_144, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k3_xboole_0(esk3_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))!=k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))!=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_79]), c_0_79]), c_0_141])).
cnf(c_0_145, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_140]), c_0_141])).
cnf(c_0_146, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_147, negated_conjecture, (k3_xboole_0(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))!=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_140]), c_0_79])]), c_0_145])).
cnf(c_0_148, plain, (X1=X3|X2!=k3_enumset1(X3,X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_55]), c_0_43]), c_0_57]), c_0_58])).
cnf(c_0_149, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_150, negated_conjecture, (r2_hidden(esk2_0,k5_xboole_0(esk3_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(spm,[status(thm)],[c_0_147, c_0_110])).
cnf(c_0_151, plain, (X1=X2|~r2_hidden(X1,k3_enumset1(X2,X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_148])).
cnf(c_0_152, negated_conjecture, (r2_hidden(esk2_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_145])).
cnf(c_0_153, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_135]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 154
# Proof object clause steps            : 95
# Proof object formula steps           : 59
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 29
# Proof object generating inferences   : 39
# Proof object simplifying inferences  : 128
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 44
# Processed clauses                    : 2416
# ...of these trivial                  : 111
# ...subsumed                          : 1848
# ...remaining for further processing  : 457
# Other redundant clauses eliminated   : 280
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 42
# Backward-rewritten                   : 25
# Generated clauses                    : 16805
# ...of the previous two non-trivial   : 13132
# Contextual simplify-reflections      : 20
# Paramodulations                      : 16524
# Factorizations                       : 2
# Equation resolutions                 : 280
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 340
#    Positive orientable unit clauses  : 70
#    Positive unorientable unit clauses: 13
#    Negative unit clauses             : 16
#    Non-unit-clauses                  : 241
# Current number of unprocessed clauses: 10713
# ...number of literals in the above   : 24982
# Current number of archived formulas  : 0
# Current number of archived clauses   : 115
# Clause-clause subsumption calls (NU) : 13692
# Rec. Clause-clause subsumption calls : 8968
# Non-unit clause-clause subsumptions  : 874
# Unit Clause-clause subsumption calls : 438
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 367
# BW rewrite match successes           : 134
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 201179

# -------------------------------------------------
# User time                : 0.139 s
# System time              : 0.007 s
# Total time               : 0.146 s
# Maximum resident set size: 28624 pages
