# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t49_mcart_1, conjecture, ![X1, X2, X3]:(~((~(r1_tarski(X1,k3_zfmisc_1(X1,X2,X3)))&(~(r1_tarski(X1,k3_zfmisc_1(X2,X3,X1)))&~(r1_tarski(X1,k3_zfmisc_1(X3,X1,X2))))))=>X1=k1_xboole_0), file('data/MPTP/proofs/8ab90d4be2934ce49a756daedbc574ab.E_input', t49_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/8ab90d4be2934ce49a756daedbc574ab.E_input', t7_mcart_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/8ab90d4be2934ce49a756daedbc574ab.E_input', d2_zfmisc_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/8ab90d4be2934ce49a756daedbc574ab.E_input', d3_zfmisc_1)).
fof(t29_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k3_mcart_1(X3,X4,X5))))))), file('data/MPTP/proofs/8ab90d4be2934ce49a756daedbc574ab.E_input', t29_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/8ab90d4be2934ce49a756daedbc574ab.E_input', d3_mcart_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/8ab90d4be2934ce49a756daedbc574ab.E_input', d3_tarski)).
fof(t135_zfmisc_1, axiom, ![X1, X2]:((r1_tarski(X1,k2_zfmisc_1(X1,X2))|r1_tarski(X1,k2_zfmisc_1(X2,X1)))=>X1=k1_xboole_0), file('data/MPTP/proofs/8ab90d4be2934ce49a756daedbc574ab.E_input', t135_zfmisc_1)).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3]:(~((~(r1_tarski(X1,k3_zfmisc_1(X1,X2,X3)))&(~(r1_tarski(X1,k3_zfmisc_1(X2,X3,X1)))&~(r1_tarski(X1,k3_zfmisc_1(X3,X1,X2))))))=>X1=k1_xboole_0)), inference(assume_negation,[status(cth)],[t49_mcart_1])).
fof(c_0_9, plain, ![X10, X11]:(k1_mcart_1(k4_tarski(X10,X11))=X10&k2_mcart_1(k4_tarski(X10,X11))=X11), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_10, plain, ![X12, X13, X14, X15, X18, X19, X20, X21, X22, X23, X25, X26]:((((r2_hidden(esk4_4(X12,X13,X14,X15),X12)|~r2_hidden(X15,X14)|X14!=k2_zfmisc_1(X12,X13))&((r2_hidden(esk5_4(X12,X13,X14,X15),X13)|~r2_hidden(X15,X14)|X14!=k2_zfmisc_1(X12,X13))&(X15=k4_tarski(esk4_4(X12,X13,X14,X15),esk5_4(X12,X13,X14,X15))|~r2_hidden(X15,X14)|X14!=k2_zfmisc_1(X12,X13))))&(~r2_hidden(X19,X12)|(~r2_hidden(X20,X13)|X18!=k4_tarski(X19,X20))|r2_hidden(X18,X14)|X14!=k2_zfmisc_1(X12,X13)))&((~r2_hidden(esk6_3(X21,X22,X23),X23)|(~r2_hidden(X25,X21)|(~r2_hidden(X26,X22)|esk6_3(X21,X22,X23)!=k4_tarski(X25,X26)))|X23=k2_zfmisc_1(X21,X22))&((r2_hidden(esk7_3(X21,X22,X23),X21)|r2_hidden(esk6_3(X21,X22,X23),X23)|X23=k2_zfmisc_1(X21,X22))&((r2_hidden(esk8_3(X21,X22,X23),X22)|r2_hidden(esk6_3(X21,X22,X23),X23)|X23=k2_zfmisc_1(X21,X22))&(esk6_3(X21,X22,X23)=k4_tarski(esk7_3(X21,X22,X23),esk8_3(X21,X22,X23))|r2_hidden(esk6_3(X21,X22,X23),X23)|X23=k2_zfmisc_1(X21,X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_11, negated_conjecture, ((r1_tarski(esk1_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))|(r1_tarski(esk1_0,k3_zfmisc_1(esk2_0,esk3_0,esk1_0))|r1_tarski(esk1_0,k3_zfmisc_1(esk3_0,esk1_0,esk2_0))))&esk1_0!=k1_xboole_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
fof(c_0_12, plain, ![X29, X30, X31]:k3_zfmisc_1(X29,X30,X31)=k2_zfmisc_1(k2_zfmisc_1(X29,X30),X31), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_13, plain, ![X32, X34, X35, X36]:((r2_hidden(esk9_1(X32),X32)|X32=k1_xboole_0)&((~r2_hidden(X34,X32)|esk9_1(X32)!=k3_mcart_1(X34,X35,X36)|X32=k1_xboole_0)&(~r2_hidden(X35,X32)|esk9_1(X32)!=k3_mcart_1(X34,X35,X36)|X32=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_mcart_1])])])])])])])).
fof(c_0_14, plain, ![X37, X38, X39]:k3_mcart_1(X37,X38,X39)=k4_tarski(k4_tarski(X37,X38),X39), inference(variable_rename,[status(thm)],[d3_mcart_1])).
cnf(c_0_15, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (X1=k4_tarski(esk4_4(X2,X3,X4,X1),esk5_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, plain, ![X40, X41, X42, X43, X44]:((~r1_tarski(X40,X41)|(~r2_hidden(X42,X40)|r2_hidden(X42,X41)))&((r2_hidden(esk10_2(X43,X44),X43)|r1_tarski(X43,X44))&(~r2_hidden(esk10_2(X43,X44),X44)|r1_tarski(X43,X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_18, negated_conjecture, (r1_tarski(esk1_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))|r1_tarski(esk1_0,k3_zfmisc_1(esk2_0,esk3_0,esk1_0))|r1_tarski(esk1_0,k3_zfmisc_1(esk3_0,esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk9_1(X2)!=k3_mcart_1(X3,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_23, plain, (k1_mcart_1(X1)=esk4_4(X2,X3,X4,X1)|X4!=k2_zfmisc_1(X2,X3)|~r2_hidden(X1,X4)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_24, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk3_0,esk1_0),esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_19]), c_0_19])).
cnf(c_0_26, plain, (X2=k1_xboole_0|esk9_1(X2)!=k4_tarski(k4_tarski(X3,X1),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (r2_hidden(k1_mcart_1(X1),X2)|X3!=k2_zfmisc_1(X2,X4)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(k2_zfmisc_1(esk3_0,esk1_0),esk2_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, plain, (r2_hidden(esk9_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_31, plain, (X1=k1_xboole_0|esk9_1(X1)!=k4_tarski(X2,X3)|X4!=k2_zfmisc_1(X5,X6)|~r2_hidden(esk5_4(X5,X6,X4,X2),X1)|~r2_hidden(X2,X4)), inference(spm,[status(thm)],[c_0_26, c_0_16])).
cnf(c_0_32, plain, (r2_hidden(esk5_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_33, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk9_1(esk1_0),k2_zfmisc_1(k2_zfmisc_1(esk3_0,esk1_0),esk2_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_35, plain, (X1=k1_xboole_0|esk9_1(X1)!=k4_tarski(X2,X3)|X4!=k2_zfmisc_1(X5,X1)|~r2_hidden(X2,X4)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, negated_conjecture, (r2_hidden(k1_mcart_1(esk9_1(esk1_0)),k2_zfmisc_1(esk3_0,esk1_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_37, plain, (X1=k1_xboole_0|esk9_1(X1)!=k4_tarski(X2,X3)|~r2_hidden(X2,k2_zfmisc_1(X4,X1))), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk4_4(X1,X2,X3,esk9_1(esk1_0)),k2_zfmisc_1(esk3_0,esk1_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))|X3!=k2_zfmisc_1(X1,X2)|~r2_hidden(esk9_1(esk1_0),X3)), inference(spm,[status(thm)],[c_0_36, c_0_23])).
cnf(c_0_39, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|k4_tarski(esk4_4(X1,X2,X3,esk9_1(esk1_0)),X4)!=esk9_1(esk1_0)|X3!=k2_zfmisc_1(X1,X2)|~r2_hidden(esk9_1(esk1_0),X3)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_30])).
cnf(c_0_40, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))|X1!=k2_zfmisc_1(X2,X3)|~r2_hidden(esk9_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_39, c_0_16])).
cnf(c_0_41, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk9_1(X2)!=k3_mcart_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_42, plain, ![X46, X47]:((~r1_tarski(X46,k2_zfmisc_1(X46,X47))|X46=k1_xboole_0)&(~r1_tarski(X46,k2_zfmisc_1(X47,X46))|X46=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t135_zfmisc_1])])])).
cnf(c_0_43, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|k2_zfmisc_1(k2_zfmisc_1(esk3_0,esk1_0),esk2_0)!=k2_zfmisc_1(X1,X2)), inference(spm,[status(thm)],[c_0_40, c_0_34])).
cnf(c_0_44, plain, (X2=k1_xboole_0|esk9_1(X2)!=k4_tarski(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_41, c_0_21])).
cnf(c_0_45, plain, (X1=k1_xboole_0|~r1_tarski(X1,k2_zfmisc_1(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk1_0))), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_47, plain, (X1=k1_xboole_0|esk9_1(X1)!=k4_tarski(X2,X3)|X4!=k2_zfmisc_1(X5,X6)|~r2_hidden(esk4_4(X5,X6,X4,X2),X1)|~r2_hidden(X2,X4)), inference(spm,[status(thm)],[c_0_44, c_0_16])).
cnf(c_0_48, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_30])).
cnf(c_0_49, plain, (X1=k1_xboole_0|esk9_1(X1)!=k4_tarski(X2,X3)|X4!=k2_zfmisc_1(X1,X5)|~r2_hidden(X2,X4)), inference(spm,[status(thm)],[c_0_47, c_0_22])).
cnf(c_0_50, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_48])).
cnf(c_0_51, plain, (X1=k1_xboole_0|esk9_1(X1)!=k4_tarski(X2,X3)|~r2_hidden(X2,k2_zfmisc_1(X1,X4))), inference(er,[status(thm)],[c_0_49])).
cnf(c_0_52, negated_conjecture, (r2_hidden(k1_mcart_1(X1),k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_50])).
cnf(c_0_53, negated_conjecture, (k4_tarski(k1_mcart_1(X1),X2)!=esk9_1(esk1_0)|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_30])).
cnf(c_0_54, negated_conjecture, (k4_tarski(esk4_4(X1,X2,X3,X4),X5)!=esk9_1(esk1_0)|X3!=k2_zfmisc_1(X1,X2)|~r2_hidden(X4,esk1_0)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_53, c_0_23])).
cnf(c_0_55, negated_conjecture, (X1!=esk9_1(esk1_0)|X2!=k2_zfmisc_1(X3,X4)|~r2_hidden(X1,esk1_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_54, c_0_16])).
cnf(c_0_56, negated_conjecture, (X1!=k2_zfmisc_1(X2,X3)|~r2_hidden(esk9_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_29]), c_0_30])).
cnf(c_0_57, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)!=k2_zfmisc_1(X1,X2)|~r2_hidden(esk9_1(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_56, c_0_50])).
cnf(c_0_58, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)!=k2_zfmisc_1(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_29]), c_0_30])).
cnf(c_0_59, negated_conjecture, ($false), inference(er,[status(thm)],[c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 43
# Proof object formula steps           : 17
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 20
# Processed clauses                    : 156
# ...of these trivial                  : 0
# ...subsumed                          : 12
# ...remaining for further processing  : 144
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 52
# Backward-rewritten                   : 1
# Generated clauses                    : 333
# ...of the previous two non-trivial   : 327
# Contextual simplify-reflections      : 0
# Paramodulations                      : 321
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 71
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 64
# Current number of unprocessed clauses: 209
# ...number of literals in the above   : 964
# Current number of archived formulas  : 0
# Current number of archived clauses   : 75
# Clause-clause subsumption calls (NU) : 1363
# Rec. Clause-clause subsumption calls : 323
# Non-unit clause-clause subsumptions  : 64
# Unit Clause-clause subsumption calls : 30
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8957

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 28744 pages
