# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t75_xboole_1, conjecture, ![X1, X2]:~((~(r1_xboole_0(X1,X2))&r1_xboole_0(k3_xboole_0(X1,X2),X2))), file('data/MPTP/proofs/8ac1c41f2373449a99f8623caa5e85ee.E_input', t75_xboole_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/8ac1c41f2373449a99f8623caa5e85ee.E_input', d4_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/8ac1c41f2373449a99f8623caa5e85ee.E_input', t3_xboole_0)).
fof(c_0_3, negated_conjecture, ~(![X1, X2]:~((~(r1_xboole_0(X1,X2))&r1_xboole_0(k3_xboole_0(X1,X2),X2)))), inference(assume_negation,[status(cth)],[t75_xboole_1])).
fof(c_0_4, plain, ![X5, X6, X7, X8, X9, X10, X11, X12]:((((r2_hidden(X8,X5)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6))&(r2_hidden(X8,X6)|~r2_hidden(X8,X7)|X7!=k3_xboole_0(X5,X6)))&(~r2_hidden(X9,X5)|~r2_hidden(X9,X6)|r2_hidden(X9,X7)|X7!=k3_xboole_0(X5,X6)))&((~r2_hidden(esk1_3(X10,X11,X12),X12)|(~r2_hidden(esk1_3(X10,X11,X12),X10)|~r2_hidden(esk1_3(X10,X11,X12),X11))|X12=k3_xboole_0(X10,X11))&((r2_hidden(esk1_3(X10,X11,X12),X10)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))&(r2_hidden(esk1_3(X10,X11,X12),X11)|r2_hidden(esk1_3(X10,X11,X12),X12)|X12=k3_xboole_0(X10,X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_5, plain, ![X14, X15, X17, X18, X19]:(((r2_hidden(esk2_2(X14,X15),X14)|r1_xboole_0(X14,X15))&(r2_hidden(esk2_2(X14,X15),X15)|r1_xboole_0(X14,X15)))&(~r2_hidden(X19,X17)|~r2_hidden(X19,X18)|~r1_xboole_0(X17,X18))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_6, negated_conjecture, (~r1_xboole_0(esk3_0,esk4_0)&r1_xboole_0(k3_xboole_0(esk3_0,esk4_0),esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])).
cnf(c_0_7, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, negated_conjecture, (r1_xboole_0(k3_xboole_0(esk3_0,esk4_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12, negated_conjecture, (~r2_hidden(X1,k3_xboole_0(esk3_0,esk4_0))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X3)|r2_hidden(esk2_2(k3_xboole_0(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, negated_conjecture, (r1_xboole_0(k3_xboole_0(X1,esk4_0),X2)|~r2_hidden(esk2_2(k3_xboole_0(X1,esk4_0),X2),k3_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_15, negated_conjecture, (r1_xboole_0(k3_xboole_0(esk3_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_14, c_0_11])).
cnf(c_0_16, plain, (r2_hidden(esk1_3(X1,X2,X3),X2)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17, negated_conjecture, (~r2_hidden(X1,k3_xboole_0(esk3_0,esk4_0))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_8, c_0_15])).
cnf(c_0_18, plain, (k3_xboole_0(X1,X2)=X2|r2_hidden(esk1_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_16])).
cnf(c_0_19, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_20, negated_conjecture, (k3_xboole_0(X1,k3_xboole_0(esk3_0,esk4_0))=k3_xboole_0(esk3_0,esk4_0)|~r2_hidden(esk1_3(X1,k3_xboole_0(esk3_0,esk4_0),k3_xboole_0(esk3_0,esk4_0)),X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_23, negated_conjecture, (k3_xboole_0(X1,k3_xboole_0(esk3_0,esk4_0))=k3_xboole_0(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_20, c_0_18])).
cnf(c_0_24, plain, (r2_hidden(X1,k3_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (r2_hidden(esk1_3(X1,X2,X3),X1)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_26, negated_conjecture, (~r2_hidden(X1,k3_xboole_0(esk3_0,esk4_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_17])).
cnf(c_0_27, plain, (X1=k3_xboole_0(X2,X3)|r2_hidden(esk1_3(X2,X3,X1),k3_xboole_0(X4,X2))|r2_hidden(esk1_3(X2,X3,X1),X1)|~r2_hidden(esk1_3(X2,X3,X1),X4)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, negated_conjecture, (k3_xboole_0(esk3_0,esk4_0)=k3_xboole_0(X1,X2)|r2_hidden(esk1_3(X1,X2,k3_xboole_0(esk3_0,esk4_0)),X2)), inference(spm,[status(thm)],[c_0_26, c_0_16])).
cnf(c_0_29, negated_conjecture, (k3_xboole_0(esk3_0,esk4_0)=k3_xboole_0(X1,X2)|r2_hidden(esk1_3(X1,X2,k3_xboole_0(esk3_0,esk4_0)),k3_xboole_0(X2,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_26])).
cnf(c_0_30, negated_conjecture, (k3_xboole_0(esk3_0,esk4_0)=k3_xboole_0(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_26, c_0_29])).
cnf(c_0_31, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),k3_xboole_0(X3,X1))|~r2_hidden(esk2_2(X1,X2),X3)), inference(spm,[status(thm)],[c_0_24, c_0_11])).
cnf(c_0_32, plain, (r2_hidden(esk2_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_33, negated_conjecture, (~r2_hidden(X1,k3_xboole_0(esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_26, c_0_30])).
cnf(c_0_34, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_35, negated_conjecture, (~r1_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_36, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 30
# Proof object formula steps           : 7
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 11
# Processed clauses                    : 752
# ...of these trivial                  : 54
# ...subsumed                          : 438
# ...remaining for further processing  : 260
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 178
# Generated clauses                    : 2637
# ...of the previous two non-trivial   : 2390
# Contextual simplify-reflections      : 3
# Paramodulations                      : 2609
# Factorizations                       : 20
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 66
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 54
# Current number of unprocessed clauses: 1182
# ...number of literals in the above   : 3339
# Current number of archived formulas  : 0
# Current number of archived clauses   : 194
# Clause-clause subsumption calls (NU) : 4905
# Rec. Clause-clause subsumption calls : 4608
# Non-unit clause-clause subsumptions  : 375
# Unit Clause-clause subsumption calls : 177
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 210
# BW rewrite match successes           : 23
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 45968

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.004 s
# Total time               : 0.034 s
# Maximum resident set size: 28716 pages
