# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', t69_enumset1)).
fof(t36_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', t36_tex_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', dt_l1_pre_topc)).
fof(t63_subset_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', t63_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', d1_xboole_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', fc2_struct_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', redefinition_k6_domain_1)).
fof(t10_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>?[X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&m1_pre_topc(X3,X1)))&X2=u1_struct_0(X3)))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', t10_tsep_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', t2_subset)).
fof(d4_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&m1_pre_topc(X3,X1)))=>(X3=k1_tex_2(X1,X2)<=>u1_struct_0(X3)=k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', d4_tex_2)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', dt_k6_domain_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', d2_subset_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', dt_m1_pre_topc)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', fc1_struct_0)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', cc1_pre_topc)).
fof(t26_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>(v2_tex_2(X3,X1)<=>v1_tdlat_3(X2)))))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', t26_tex_2)).
fof(t24_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v2_pre_topc(k1_tex_2(X1,X2))=>(v1_tdlat_3(k1_tex_2(X1,X2))&v2_tdlat_3(k1_tex_2(X1,X2)))))), file('data/MPTP/proofs/8b36a0ddc3154291891ae2eea736e791.E_input', t24_tex_2)).
fof(c_0_18, plain, ![X22, X23, X24, X25, X26, X27]:(((~r2_hidden(X24,X23)|X24=X22|X23!=k1_tarski(X22))&(X25!=X22|r2_hidden(X25,X23)|X23!=k1_tarski(X22)))&((~r2_hidden(esk4_2(X26,X27),X27)|esk4_2(X26,X27)!=X26|X27=k1_tarski(X26))&(r2_hidden(esk4_2(X26,X27),X27)|esk4_2(X26,X27)=X26|X27=k1_tarski(X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_19, plain, ![X29]:k2_tarski(X29,X29)=k1_tarski(X29), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_20, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)))), inference(assume_negation,[status(cth)],[t36_tex_2])).
cnf(c_0_23, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_24, plain, ![X30]:(~l1_pre_topc(X30)|l1_struct_0(X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_25, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&~v2_tex_2(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
fof(c_0_26, plain, ![X19, X20]:(~r2_hidden(X19,X20)|m1_subset_1(k1_tarski(X19),k1_zfmisc_1(X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_subset_1])])).
fof(c_0_27, plain, ![X8, X9, X10]:((~v1_xboole_0(X8)|~r2_hidden(X9,X8))&(r2_hidden(esk3_1(X10),X10)|v1_xboole_0(X10))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X1)), inference(er,[status(thm)],[c_0_23])).
fof(c_0_29, plain, ![X49]:(v2_struct_0(X49)|~l1_struct_0(X49)|~v1_xboole_0(u1_struct_0(X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_30, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X39, X40]:(v1_xboole_0(X39)|~m1_subset_1(X40,X39)|k6_domain_1(X39,X40)=k1_tarski(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_33, plain, ![X46, X47]:(((~v2_struct_0(esk6_2(X46,X47))|(v1_xboole_0(X47)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46))))|(v2_struct_0(X46)|~l1_pre_topc(X46)))&((v1_pre_topc(esk6_2(X46,X47))|(v1_xboole_0(X47)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46))))|(v2_struct_0(X46)|~l1_pre_topc(X46)))&(m1_pre_topc(esk6_2(X46,X47),X46)|(v1_xboole_0(X47)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46))))|(v2_struct_0(X46)|~l1_pre_topc(X46)))))&(X47=u1_struct_0(esk6_2(X46,X47))|(v1_xboole_0(X47)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46))))|(v2_struct_0(X46)|~l1_pre_topc(X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_tsep_1])])])])])])).
cnf(c_0_34, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[c_0_28])).
fof(c_0_37, plain, ![X6, X7]:(~m1_subset_1(X6,X7)|(v1_xboole_0(X7)|r2_hidden(X6,X7))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_41, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_42, plain, ![X16, X17, X18]:((X18!=k1_tex_2(X16,X17)|u1_struct_0(X18)=k6_domain_1(u1_struct_0(X16),X17)|(v2_struct_0(X18)|(~v1_pre_topc(X18)|~m1_pre_topc(X18,X16)))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|~l1_pre_topc(X16)))&(u1_struct_0(X18)!=k6_domain_1(u1_struct_0(X16),X17)|X18=k1_tex_2(X16,X17)|(v2_struct_0(X18)|(~v1_pre_topc(X18)|~m1_pre_topc(X18,X16)))|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|~l1_pre_topc(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tex_2])])])])])).
fof(c_0_43, plain, ![X14, X15]:(v1_xboole_0(X14)|~m1_subset_1(X15,X14)|m1_subset_1(k6_domain_1(X14,X15),k1_zfmisc_1(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_44, plain, (m1_pre_topc(esk6_2(X1,X2),X1)|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_45, plain, (m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_34, c_0_21])).
cnf(c_0_46, plain, (~v1_xboole_0(k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_47, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_49, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
fof(c_0_50, plain, ![X12, X13]:(((~m1_subset_1(X13,X12)|r2_hidden(X13,X12)|v1_xboole_0(X12))&(~r2_hidden(X13,X12)|m1_subset_1(X13,X12)|v1_xboole_0(X12)))&((~m1_subset_1(X13,X12)|v1_xboole_0(X13)|~v1_xboole_0(X12))&(~v1_xboole_0(X13)|m1_subset_1(X13,X12)|~v1_xboole_0(X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_51, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_41, c_0_21])).
cnf(c_0_52, plain, (X1=k1_tex_2(X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(X1)!=k6_domain_1(u1_struct_0(X2),X3)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_54, plain, ![X34, X35]:(~l1_pre_topc(X34)|(~m1_pre_topc(X35,X34)|l1_pre_topc(X35))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_55, plain, (m1_pre_topc(esk6_2(X1,k2_tarski(X2,X2)),X1)|v2_struct_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])).
cnf(c_0_56, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_57, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, plain, (v1_xboole_0(X1)|~v1_xboole_0(k6_domain_1(X1,X2))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_46, c_0_51])).
fof(c_0_59, plain, ![X31]:(~v2_struct_0(X31)|~l1_struct_0(X31)|v1_xboole_0(u1_struct_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_60, plain, (X1=k1_tex_2(X2,X3)|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(X1)!=k2_tarski(X3,X3)|~m1_pre_topc(X1,X2)|~v1_pre_topc(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_52, c_0_51])).
cnf(c_0_61, plain, (X1=u1_struct_0(esk6_2(X2,X1))|v1_xboole_0(X1)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_62, plain, (v1_xboole_0(X1)|m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_53, c_0_51])).
cnf(c_0_63, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64, negated_conjecture, (m1_pre_topc(esk6_2(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_31])]), c_0_38])).
cnf(c_0_65, plain, (v1_pre_topc(esk6_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_66, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_53]), c_0_58])).
cnf(c_0_67, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_68, negated_conjecture, (X1=k1_tex_2(esk1_0,esk2_0)|v2_struct_0(X1)|u1_struct_0(X1)!=k2_tarski(esk2_0,esk2_0)|~m1_pre_topc(X1,esk1_0)|~v1_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_48]), c_0_31])]), c_0_49]), c_0_38])).
cnf(c_0_69, plain, (u1_struct_0(esk6_2(X1,k2_tarski(X2,X2)))=k2_tarski(X2,X2)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_46])).
cnf(c_0_70, negated_conjecture, (l1_pre_topc(esk6_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_31])])).
cnf(c_0_71, plain, (v1_pre_topc(esk6_2(X1,k2_tarski(X2,X2)))|v2_struct_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_45]), c_0_46])).
cnf(c_0_72, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_73, plain, (r2_hidden(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_53]), c_0_66])).
cnf(c_0_74, negated_conjecture, (X1=k1_tex_2(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k2_tarski(esk2_0,esk2_0)|~l1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)|~v1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_75, negated_conjecture, (u1_struct_0(esk6_2(esk1_0,k2_tarski(esk2_0,esk2_0)))=k2_tarski(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_48]), c_0_31])]), c_0_49]), c_0_38])).
cnf(c_0_76, negated_conjecture, (l1_struct_0(esk6_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(spm,[status(thm)],[c_0_30, c_0_70])).
cnf(c_0_77, negated_conjecture, (v1_pre_topc(esk6_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_56]), c_0_31])]), c_0_38])).
cnf(c_0_78, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|~v2_struct_0(esk6_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_79, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_72, c_0_35])).
cnf(c_0_80, plain, (r2_hidden(k2_tarski(X1,X1),k1_zfmisc_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(spm,[status(thm)],[c_0_73, c_0_51])).
cnf(c_0_81, plain, (u1_struct_0(X1)=k6_domain_1(u1_struct_0(X2),X3)|v2_struct_0(X1)|v2_struct_0(X2)|X1!=k1_tex_2(X2,X3)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_82, negated_conjecture, (esk6_2(esk1_0,k2_tarski(esk2_0,esk2_0))=k1_tex_2(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76]), c_0_64]), c_0_77])]), c_0_46])).
cnf(c_0_83, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|~r2_hidden(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_struct_0(esk6_2(X2,X1))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_84, negated_conjecture, (r2_hidden(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_48]), c_0_49])).
fof(c_0_85, plain, ![X32, X33]:(~v2_pre_topc(X32)|~l1_pre_topc(X32)|(~m1_pre_topc(X33,X32)|v2_pre_topc(X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
fof(c_0_86, plain, ![X36, X37, X38]:((~v2_tex_2(X38,X36)|v1_tdlat_3(X37)|X38!=u1_struct_0(X37)|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X36)))|(v2_struct_0(X37)|~m1_pre_topc(X37,X36))|(v2_struct_0(X36)|~l1_pre_topc(X36)))&(~v1_tdlat_3(X37)|v2_tex_2(X38,X36)|X38!=u1_struct_0(X37)|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X36)))|(v2_struct_0(X37)|~m1_pre_topc(X37,X36))|(v2_struct_0(X36)|~l1_pre_topc(X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_tex_2])])])])])).
cnf(c_0_87, plain, (k6_domain_1(u1_struct_0(X1),X2)=u1_struct_0(k1_tex_2(X1,X2))|v2_struct_0(k1_tex_2(X1,X2))|v2_struct_0(X1)|~m1_pre_topc(k1_tex_2(X1,X2),X1)|~v1_pre_topc(k1_tex_2(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(er,[status(thm)],[c_0_81])).
cnf(c_0_88, negated_conjecture, (m1_pre_topc(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(rw,[status(thm)],[c_0_64, c_0_82])).
cnf(c_0_89, negated_conjecture, (v1_pre_topc(k1_tex_2(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_77, c_0_82])).
cnf(c_0_90, negated_conjecture, (~v2_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_82]), c_0_84]), c_0_31])]), c_0_46]), c_0_38])).
fof(c_0_91, plain, ![X44, X45]:((v1_tdlat_3(k1_tex_2(X44,X45))|~v2_pre_topc(k1_tex_2(X44,X45))|~m1_subset_1(X45,u1_struct_0(X44))|(v2_struct_0(X44)|~l1_pre_topc(X44)))&(v2_tdlat_3(k1_tex_2(X44,X45))|~v2_pre_topc(k1_tex_2(X44,X45))|~m1_subset_1(X45,u1_struct_0(X44))|(v2_struct_0(X44)|~l1_pre_topc(X44)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_tex_2])])])])])).
cnf(c_0_92, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_93, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_94, plain, (v2_tex_2(X2,X3)|v2_struct_0(X1)|v2_struct_0(X3)|~v1_tdlat_3(X1)|X2!=u1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_95, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=u1_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89]), c_0_48]), c_0_31])]), c_0_90]), c_0_38])).
cnf(c_0_96, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,esk2_0))=k2_tarski(esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_75, c_0_82])).
cnf(c_0_97, negated_conjecture, (~v2_tex_2(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_98, plain, (v1_tdlat_3(k1_tex_2(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(k1_tex_2(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_99, negated_conjecture, (v2_pre_topc(esk6_2(esk1_0,k2_tarski(esk2_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_64]), c_0_31]), c_0_93])])).
cnf(c_0_100, plain, (v1_xboole_0(u1_struct_0(X1))|v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)|v2_struct_0(X3)|v2_struct_0(X1)|k6_domain_1(u1_struct_0(X1),X2)!=u1_struct_0(X3)|~v1_tdlat_3(X3)|~m1_pre_topc(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_94, c_0_53])).
cnf(c_0_101, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_102, negated_conjecture, (~v2_tex_2(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_51]), c_0_48])]), c_0_49])).
cnf(c_0_103, negated_conjecture, (v1_tdlat_3(k1_tex_2(esk1_0,esk2_0))|~v2_pre_topc(k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_48]), c_0_31])]), c_0_38])).
cnf(c_0_104, negated_conjecture, (v2_pre_topc(k1_tex_2(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_99, c_0_82])).
cnf(c_0_105, negated_conjecture, (v2_struct_0(X1)|u1_struct_0(X1)!=k2_tarski(esk2_0,esk2_0)|~v1_tdlat_3(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_48]), c_0_31])]), c_0_49]), c_0_102]), c_0_38])).
cnf(c_0_106, negated_conjecture, (v1_tdlat_3(k1_tex_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_107, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_105]), c_0_96]), c_0_106]), c_0_88])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 108
# Proof object clause steps            : 71
# Proof object formula steps           : 37
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 72
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 38
# Processed clauses                    : 1311
# ...of these trivial                  : 16
# ...subsumed                          : 703
# ...remaining for further processing  : 592
# Other redundant clauses eliminated   : 135
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 28
# Backward-rewritten                   : 26
# Generated clauses                    : 12050
# ...of the previous two non-trivial   : 11406
# Contextual simplify-reflections      : 55
# Paramodulations                      : 11887
# Factorizations                       : 10
# Equation resolutions                 : 153
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 500
#    Positive orientable unit clauses  : 33
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 15
#    Non-unit-clauses                  : 452
# Current number of unprocessed clauses: 10106
# ...number of literals in the above   : 63943
# Current number of archived formulas  : 0
# Current number of archived clauses   : 92
# Clause-clause subsumption calls (NU) : 43536
# Rec. Clause-clause subsumption calls : 12387
# Non-unit clause-clause subsumptions  : 630
# Unit Clause-clause subsumption calls : 1483
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 248656

# -------------------------------------------------
# User time                : 0.205 s
# System time              : 0.003 s
# Total time               : 0.209 s
# Maximum resident set size: 28744 pages
