# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', t70_enumset1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', t43_xboole_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', redefinition_k6_subset_1)).
fof(t51_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>k5_relat_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k5_relat_1(X1,X2),k5_relat_1(X1,X3))))), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', t51_relat_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', t39_xboole_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', t10_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', commutativity_k2_tarski)).
fof(t53_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>r1_tarski(k6_subset_1(k5_relat_1(X1,X2),k5_relat_1(X1,X3)),k5_relat_1(X1,k6_subset_1(X2,X3)))))), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', t53_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', d10_xboole_0)).
fof(fc2_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k4_xboole_0(X1,X2))), file('data/MPTP/proofs/8b3a846a2c534e698173926434bdec3c.E_input', fc2_relat_1)).
fof(c_0_11, plain, ![X7, X8]:k3_tarski(k2_tarski(X7,X8))=k2_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_12, plain, ![X15, X16]:k1_enumset1(X15,X15,X16)=k2_tarski(X15,X16), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_13, plain, ![X22, X23, X24]:(~r1_tarski(X22,k2_xboole_0(X23,X24))|r1_tarski(k4_xboole_0(X22,X23),X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
fof(c_0_14, plain, ![X11, X12]:k6_subset_1(X11,X12)=k4_xboole_0(X11,X12), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_15, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, ![X27, X28, X29]:(~v1_relat_1(X27)|(~v1_relat_1(X28)|(~v1_relat_1(X29)|k5_relat_1(X27,k2_xboole_0(X28,X29))=k2_xboole_0(k5_relat_1(X27,X28),k5_relat_1(X27,X29))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t51_relat_1])])])).
cnf(c_0_18, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21, plain, (k5_relat_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k5_relat_1(X1,X2),k5_relat_1(X1,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X20, X21]:k2_xboole_0(X20,k4_xboole_0(X21,X20))=k2_xboole_0(X20,X21), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_23, plain, ![X17, X18, X19]:(~r1_tarski(X17,X18)|r1_tarski(X17,k2_xboole_0(X19,X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
fof(c_0_24, plain, ![X9, X10]:k2_tarski(X9,X10)=k2_tarski(X10,X9), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_25, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>r1_tarski(k6_subset_1(k5_relat_1(X1,X2),k5_relat_1(X1,X3)),k5_relat_1(X1,k6_subset_1(X2,X3))))))), inference(assume_negation,[status(cth)],[t53_relat_1])).
cnf(c_0_26, plain, (r1_tarski(k6_subset_1(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_27, plain, (k5_relat_1(X1,k3_tarski(k1_enumset1(X2,X2,X3)))=k3_tarski(k1_enumset1(k5_relat_1(X1,X2),k5_relat_1(X1,X2),k5_relat_1(X1,X3)))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_20]), c_0_20])).
cnf(c_0_28, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_31, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(v1_relat_1(esk3_0)&~r1_tarski(k6_subset_1(k5_relat_1(esk1_0,esk2_0),k5_relat_1(esk1_0,esk3_0)),k5_relat_1(esk1_0,k6_subset_1(esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])).
cnf(c_0_32, plain, (r1_tarski(k6_subset_1(X1,k5_relat_1(X2,X3)),k5_relat_1(X2,X4))|~r1_tarski(X1,k5_relat_1(X2,k3_tarski(k1_enumset1(X3,X3,X4))))|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, plain, (k3_tarski(k1_enumset1(X1,X1,k6_subset_1(X2,X1)))=k3_tarski(k1_enumset1(X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_19]), c_0_20]), c_0_20])).
cnf(c_0_34, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X3,X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_29, c_0_20])).
cnf(c_0_35, plain, (k1_enumset1(X1,X1,X2)=k1_enumset1(X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_16]), c_0_16])).
fof(c_0_36, plain, ![X25, X26]:(((r1_tarski(X25,X26)|X25!=X26)&(r1_tarski(X26,X25)|X25!=X26))&(~r1_tarski(X25,X26)|~r1_tarski(X26,X25)|X25=X26)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_37, negated_conjecture, (~r1_tarski(k6_subset_1(k5_relat_1(esk1_0,esk2_0),k5_relat_1(esk1_0,esk3_0)),k5_relat_1(esk1_0,k6_subset_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (r1_tarski(k6_subset_1(X1,k5_relat_1(X2,X3)),k5_relat_1(X2,k6_subset_1(X4,X3)))|~r1_tarski(X1,k5_relat_1(X2,k3_tarski(k1_enumset1(X3,X3,X4))))|~v1_relat_1(k6_subset_1(X4,X3))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_43, plain, ![X13, X14]:(~v1_relat_1(X13)|v1_relat_1(k4_xboole_0(X13,X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_relat_1])])).
cnf(c_0_44, negated_conjecture, (~r1_tarski(k5_relat_1(esk1_0,esk2_0),k5_relat_1(esk1_0,k3_tarski(k1_enumset1(esk2_0,esk2_0,esk3_0))))|~v1_relat_1(k6_subset_1(esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40])]), c_0_35])).
cnf(c_0_45, plain, (r1_tarski(X1,k5_relat_1(X2,k3_tarski(k1_enumset1(X3,X3,X4))))|~r1_tarski(X1,k5_relat_1(X2,X3))|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_41, c_0_27])).
cnf(c_0_46, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_47, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_48, plain, (v1_relat_1(k4_xboole_0(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, negated_conjecture, (~v1_relat_1(k6_subset_1(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_39]), c_0_47]), c_0_40])])).
cnf(c_0_50, plain, (v1_relat_1(k6_subset_1(X1,X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_48, c_0_19])).
cnf(c_0_51, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_47])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 29
# Proof object formula steps           : 23
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 13
# Processed clauses                    : 1598
# ...of these trivial                  : 19
# ...subsumed                          : 1298
# ...remaining for further processing  : 281
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 27
# Backward-rewritten                   : 15
# Generated clauses                    : 12888
# ...of the previous two non-trivial   : 12547
# Contextual simplify-reflections      : 0
# Paramodulations                      : 12886
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 225
#    Positive orientable unit clauses  : 28
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 189
# Current number of unprocessed clauses: 10869
# ...number of literals in the above   : 38976
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 18776
# Rec. Clause-clause subsumption calls : 8414
# Non-unit clause-clause subsumptions  : 1293
# Unit Clause-clause subsumption calls : 303
# Rewrite failures with RHS unbound    : 40
# BW rewrite match attempts            : 472
# BW rewrite match successes           : 22
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 194838

# -------------------------------------------------
# User time                : 0.106 s
# System time              : 0.013 s
# Total time               : 0.119 s
# Maximum resident set size: 28656 pages
