# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t70_enumset1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t17_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t73_enumset1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t48_xboole_1)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t47_xboole_1)).
fof(fc1_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k3_xboole_0(X1,X2))), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', fc1_relat_1)).
fof(t124_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k8_relat_1(X1,X2)=k3_xboole_0(X2,k2_zfmisc_1(k1_relat_1(X2),X1))), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t124_relat_1)).
fof(t127_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k8_relat_1(X1,k8_relat_1(X2,X3))=k8_relat_1(k3_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t127_relat_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t1_xboole_1)).
fof(t25_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&r1_tarski(k2_relat_1(X1),k2_relat_1(X2)))))), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t25_relat_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t71_relat_1)).
fof(dt_k6_relat_1, axiom, ![X1]:v1_relat_1(k6_relat_1(X1)), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', dt_k6_relat_1)).
fof(t126_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k8_relat_1(k2_relat_1(X1),X1)=X1), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t126_relat_1)).
fof(t76_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k5_relat_1(X2,k6_relat_1(X1)),X2)&r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2))), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t76_relat_1)).
fof(t123_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k8_relat_1(X1,X2)=k5_relat_1(X2,k6_relat_1(X1))), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t123_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', d10_xboole_0)).
fof(t54_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k4_relat_1(k5_relat_1(X1,X2))=k5_relat_1(k4_relat_1(X2),k4_relat_1(X1)))), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t54_relat_1)).
fof(t72_relat_1, axiom, ![X1]:k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t72_relat_1)).
fof(t77_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k1_relat_1(X2),X1)=>k5_relat_1(k6_relat_1(X1),X2)=X2)), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t77_relat_1)).
fof(t43_funct_1, conjecture, ![X1, X2]:k5_relat_1(k6_relat_1(X2),k6_relat_1(X1))=k6_relat_1(k3_xboole_0(X1,X2)), file('data/MPTP/proofs/8b9b1cad2489431a82416cfb0441dc43.E_input', t43_funct_1)).
fof(c_0_23, plain, ![X15, X16]:k1_setfam_1(k2_tarski(X15,X16))=k3_xboole_0(X15,X16), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_24, plain, ![X52, X53]:k1_enumset1(X52,X52,X53)=k2_tarski(X52,X53), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_25, plain, ![X54, X55]:r1_tarski(k3_xboole_0(X54,X55),X54), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_26, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_28, plain, ![X36, X37, X38]:k2_enumset1(X36,X36,X37,X38)=k1_enumset1(X36,X37,X38), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_29, plain, ![X43, X44, X45, X46]:k3_enumset1(X43,X43,X44,X45,X46)=k2_enumset1(X43,X44,X45,X46), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_30, plain, ![X29, X30, X31, X32, X33]:k4_enumset1(X29,X29,X30,X31,X32,X33)=k3_enumset1(X29,X30,X31,X32,X33), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_31, plain, ![X27, X28]:k4_xboole_0(X27,k4_xboole_0(X27,X28))=k3_xboole_0(X27,X28), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_32, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_38, plain, ![X19, X20]:k4_xboole_0(X19,k3_xboole_0(X19,X20))=k4_xboole_0(X19,X20), inference(variable_rename,[status(thm)],[t47_xboole_1])).
fof(c_0_39, plain, ![X13, X14]:(~v1_relat_1(X13)|v1_relat_1(k3_xboole_0(X13,X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_relat_1])])).
fof(c_0_40, plain, ![X34, X35]:(~v1_relat_1(X35)|k8_relat_1(X34,X35)=k3_xboole_0(X35,k2_zfmisc_1(k1_relat_1(X35),X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t124_relat_1])])).
cnf(c_0_41, plain, (r1_tarski(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_42, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_43, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44, plain, (v1_relat_1(k3_xboole_0(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_45, plain, (k8_relat_1(X2,X1)=k3_xboole_0(X1,k2_zfmisc_1(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_46, plain, ![X24, X25, X26]:(~v1_relat_1(X26)|k8_relat_1(X24,k8_relat_1(X25,X26))=k8_relat_1(k3_xboole_0(X24,X25),X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t127_relat_1])])).
fof(c_0_47, plain, ![X47, X48, X49]:(~r1_tarski(X47,X48)|~r1_tarski(X48,X49)|r1_tarski(X47,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_48, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49, plain, (k4_xboole_0(X1,k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
fof(c_0_50, plain, ![X21, X22]:((r1_tarski(k1_relat_1(X21),k1_relat_1(X22))|~r1_tarski(X21,X22)|~v1_relat_1(X22)|~v1_relat_1(X21))&(r1_tarski(k2_relat_1(X21),k2_relat_1(X22))|~r1_tarski(X21,X22)|~v1_relat_1(X22)|~v1_relat_1(X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_relat_1])])])])).
fof(c_0_51, plain, ![X8]:(k1_relat_1(k6_relat_1(X8))=X8&k2_relat_1(k6_relat_1(X8))=X8), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_52, plain, ![X50]:v1_relat_1(k6_relat_1(X50)), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
cnf(c_0_53, plain, (v1_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_54, plain, (k8_relat_1(X2,X1)=k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,k2_zfmisc_1(k1_relat_1(X1),X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_55, plain, (k8_relat_1(X2,k8_relat_1(X3,X1))=k8_relat_1(k3_xboole_0(X2,X3),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_56, plain, ![X23]:(~v1_relat_1(X23)|k8_relat_1(k2_relat_1(X23),X23)=X23), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t126_relat_1])])).
cnf(c_0_57, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_42]), c_0_49])).
cnf(c_0_59, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_60, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_62, plain, ![X41, X42]:((r1_tarski(k5_relat_1(X42,k6_relat_1(X41)),X42)|~v1_relat_1(X42))&(r1_tarski(k5_relat_1(k6_relat_1(X41),X42),X42)|~v1_relat_1(X42))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t76_relat_1])])])).
fof(c_0_63, plain, ![X11, X12]:(~v1_relat_1(X12)|k8_relat_1(X11,X12)=k5_relat_1(X12,k6_relat_1(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t123_relat_1])])).
cnf(c_0_64, plain, (v1_relat_1(k8_relat_1(X1,X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_65, plain, (k8_relat_1(X2,k8_relat_1(X3,X1))=k8_relat_1(k1_setfam_1(k4_enumset1(X2,X2,X2,X2,X2,X3)),X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_66, plain, (k8_relat_1(k2_relat_1(X1),X1)=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_67, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_68, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_69, plain, (r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)|~r1_tarski(X1,k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])])).
cnf(c_0_70, plain, (r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_71, plain, (k8_relat_1(X2,X1)=k5_relat_1(X1,k6_relat_1(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_72, plain, (v1_relat_1(k8_relat_1(X1,k8_relat_1(X2,X3)))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_73, plain, (k8_relat_1(X1,k6_relat_1(X1))=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_61])])).
cnf(c_0_74, plain, (r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)|~r1_tarski(X1,k6_relat_1(k4_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, plain, (r1_tarski(k8_relat_1(X1,k6_relat_1(X2)),k6_relat_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_61]), c_0_61])])).
cnf(c_0_76, plain, (v1_relat_1(k8_relat_1(X1,k6_relat_1(X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_61])])).
cnf(c_0_77, plain, (r1_tarski(k1_relat_1(k8_relat_1(k4_xboole_0(X1,X2),k6_relat_1(X3))),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])])).
cnf(c_0_78, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k8_relat_1(k4_xboole_0(X2,X3),k6_relat_1(X4))))), inference(spm,[status(thm)],[c_0_57, c_0_77])).
cnf(c_0_79, plain, (k8_relat_1(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X3)=k8_relat_1(X1,k8_relat_1(X2,X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_65, c_0_42])).
fof(c_0_80, plain, ![X9, X10]:(((r1_tarski(X9,X10)|X9!=X10)&(r1_tarski(X10,X9)|X9!=X10))&(~r1_tarski(X9,X10)|~r1_tarski(X10,X9)|X9=X10)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_81, plain, ![X17, X18]:(~v1_relat_1(X17)|(~v1_relat_1(X18)|k4_relat_1(k5_relat_1(X17,X18))=k5_relat_1(k4_relat_1(X18),k4_relat_1(X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_relat_1])])])).
fof(c_0_82, plain, ![X51]:k4_relat_1(k6_relat_1(X51))=k6_relat_1(X51), inference(variable_rename,[status(thm)],[t72_relat_1])).
cnf(c_0_83, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k8_relat_1(X2,k8_relat_1(X3,k6_relat_1(X4)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_61])])).
cnf(c_0_84, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_85, plain, (k4_relat_1(k5_relat_1(X1,X2))=k5_relat_1(k4_relat_1(X2),k4_relat_1(X1))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_86, plain, (k4_relat_1(k6_relat_1(X1))=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
fof(c_0_87, plain, ![X39, X40]:(~v1_relat_1(X40)|(~r1_tarski(k1_relat_1(X40),X39)|k5_relat_1(k6_relat_1(X39),X40)=X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t77_relat_1])])).
cnf(c_0_88, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k8_relat_1(X2,k6_relat_1(X3))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_66]), c_0_61])])).
cnf(c_0_89, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_84])).
cnf(c_0_90, plain, (k4_relat_1(k5_relat_1(k6_relat_1(X1),X2))=k5_relat_1(k4_relat_1(X2),k6_relat_1(X1))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_61])])).
cnf(c_0_91, plain, (k8_relat_1(X1,k8_relat_1(X2,k6_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))))=k6_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_65]), c_0_61])])).
cnf(c_0_92, plain, (k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_42]), c_0_49])).
cnf(c_0_93, plain, (k5_relat_1(k6_relat_1(X2),X1)=X1|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_94, plain, (r1_tarski(k1_relat_1(k8_relat_1(X1,k6_relat_1(X2))),X1)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_95, plain, (k4_relat_1(k8_relat_1(X1,k6_relat_1(X2)))=k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_71]), c_0_86]), c_0_61]), c_0_61])])).
cnf(c_0_96, plain, (k8_relat_1(X1,k6_relat_1(k4_xboole_0(X1,X2)))=k6_relat_1(k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_73])).
cnf(c_0_97, plain, (k5_relat_1(k6_relat_1(X1),k8_relat_1(X1,k6_relat_1(X2)))=k8_relat_1(X1,k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_76])])).
cnf(c_0_98, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(k4_xboole_0(X1,X2)))=k6_relat_1(k4_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_86])).
cnf(c_0_99, plain, (k5_relat_1(k5_relat_1(k6_relat_1(X1),k6_relat_1(X2)),k6_relat_1(X1))=k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_97]), c_0_95]), c_0_95]), c_0_76])])).
fof(c_0_100, negated_conjecture, ~(![X1, X2]:k5_relat_1(k6_relat_1(X2),k6_relat_1(X1))=k6_relat_1(k3_xboole_0(X1,X2))), inference(assume_negation,[status(cth)],[t43_funct_1])).
cnf(c_0_101, plain, (k8_relat_1(k4_xboole_0(X1,X2),k6_relat_1(X1))=k6_relat_1(k4_xboole_0(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_98]), c_0_61])])).
cnf(c_0_102, plain, (k5_relat_1(k8_relat_1(X1,k6_relat_1(X2)),k6_relat_1(X2))=k8_relat_1(X1,k6_relat_1(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_71]), c_0_61])])).
fof(c_0_103, negated_conjecture, k5_relat_1(k6_relat_1(esk2_0),k6_relat_1(esk1_0))!=k6_relat_1(k3_xboole_0(esk1_0,esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_100])])])).
cnf(c_0_104, plain, (k6_relat_1(k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k8_relat_1(X1,k8_relat_1(X2,k6_relat_1(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_79]), c_0_61])])).
cnf(c_0_105, plain, (k8_relat_1(X1,k8_relat_1(X2,k6_relat_1(X1)))=k8_relat_1(X2,k6_relat_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_102]), c_0_76])])).
cnf(c_0_106, negated_conjecture, (k5_relat_1(k6_relat_1(esk2_0),k6_relat_1(esk1_0))!=k6_relat_1(k3_xboole_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_107, plain, (k6_relat_1(k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k8_relat_1(X2,k6_relat_1(X1))), inference(rw,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_108, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k6_relat_1(X2)|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_60]), c_0_61])])).
cnf(c_0_109, negated_conjecture, (k5_relat_1(k6_relat_1(esk2_0),k6_relat_1(esk1_0))!=k6_relat_1(k1_setfam_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_110, plain, (k5_relat_1(k6_relat_1(X1),k6_relat_1(X2))=k8_relat_1(X1,k6_relat_1(X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_107]), c_0_95])).
cnf(c_0_111, plain, (k8_relat_1(X1,k6_relat_1(X2))=k6_relat_1(X1)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_108]), c_0_61])])).
cnf(c_0_112, negated_conjecture, (k6_relat_1(k1_setfam_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))!=k8_relat_1(esk2_0,k6_relat_1(esk1_0))), inference(rw,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_113, plain, (k8_relat_1(X1,k6_relat_1(X2))=k8_relat_1(X2,k6_relat_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_110]), c_0_61])])).
cnf(c_0_114, plain, (k8_relat_1(X1,k8_relat_1(X2,k6_relat_1(X3)))=k6_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))|~r1_tarski(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)),X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_65]), c_0_61])])).
cnf(c_0_115, negated_conjecture, (k6_relat_1(k1_setfam_1(k4_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)))!=k8_relat_1(esk1_0,k6_relat_1(esk2_0))), inference(rw,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_116, plain, (k6_relat_1(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)))=k8_relat_1(X2,k6_relat_1(X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_41]), c_0_105])).
cnf(c_0_117, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_116]), c_0_113])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 118
# Proof object clause steps            : 71
# Proof object formula steps           : 47
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 87
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 23
# Processed clauses                    : 5231
# ...of these trivial                  : 293
# ...subsumed                          : 3837
# ...remaining for further processing  : 1101
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 19
# Backward-rewritten                   : 142
# Generated clauses                    : 52244
# ...of the previous two non-trivial   : 45393
# Contextual simplify-reflections      : 53
# Paramodulations                      : 52242
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 916
#    Positive orientable unit clauses  : 154
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 759
# Current number of unprocessed clauses: 39816
# ...number of literals in the above   : 95076
# Current number of archived formulas  : 0
# Current number of archived clauses   : 188
# Clause-clause subsumption calls (NU) : 77143
# Rec. Clause-clause subsumption calls : 62947
# Non-unit clause-clause subsumptions  : 3906
# Unit Clause-clause subsumption calls : 207
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 501
# BW rewrite match successes           : 105
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 789789

# -------------------------------------------------
# User time                : 0.437 s
# System time              : 0.007 s
# Total time               : 0.444 s
# Maximum resident set size: 34904 pages
