# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d15_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:((v1_orders_2(X3)&(v4_yellow_0(X3,X1)&m1_yellow_0(X3,X1)))=>(X3=k5_yellow_0(X1,X2)<=>u1_struct_0(X3)=X2)))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', d15_yellow_0)).
fof(dt_k5_yellow_0, axiom, ![X1, X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_orders_2(k5_yellow_0(X1,X2))&(v4_yellow_0(k5_yellow_0(X1,X2),X1)&m1_yellow_0(k5_yellow_0(X1,X2),X1)))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', dt_k5_yellow_0)).
fof(d16_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v5_yellow_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,u1_struct_0(X2))&(r2_hidden(X4,u1_struct_0(X2))&r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))))=>r2_hidden(k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X2)))))))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', d16_yellow_0)).
fof(t63_waybel_0, conjecture, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v1_xboole_0(X2))&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))))=>((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))<=>v5_yellow_0(k5_yellow_0(X1,X2),X1)))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', t63_waybel_0)).
fof(t40_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', t40_yellow_0)).
fof(t41_waybel_0, axiom, ![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:((v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v2_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r2_hidden(X4,X2))=>r2_hidden(k12_lattice3(X1,X3,X4),X2))))))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', t41_waybel_0)).
fof(redefinition_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', redefinition_k7_domain_1)).
fof(t21_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>(v2_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r2_yellow_0(X1,k2_tarski(X2,X3)))))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', t21_yellow_0)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', t5_subset)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/8d75d4b9a0cd4571ad1e5595ab3dc266.E_input', cc2_lattice3)).
fof(c_0_10, plain, ![X7, X8, X9]:((X9!=k5_yellow_0(X7,X8)|u1_struct_0(X9)=X8|(~v1_orders_2(X9)|(~v4_yellow_0(X9,X7)|~m1_yellow_0(X9,X7)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|~l1_orders_2(X7))&(u1_struct_0(X9)!=X8|X9=k5_yellow_0(X7,X8)|(~v1_orders_2(X9)|(~v4_yellow_0(X9,X7)|~m1_yellow_0(X9,X7)))|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|~l1_orders_2(X7))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_yellow_0])])])])).
fof(c_0_11, plain, ![X30, X31]:((v1_orders_2(k5_yellow_0(X30,X31))|(~l1_orders_2(X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))))&((v4_yellow_0(k5_yellow_0(X30,X31),X30)|(~l1_orders_2(X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))))&(m1_yellow_0(k5_yellow_0(X30,X31),X30)|(~l1_orders_2(X30)|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_yellow_0])])])).
fof(c_0_12, plain, ![X18, X19, X20, X21]:((~v5_yellow_0(X19,X18)|(~m1_subset_1(X20,u1_struct_0(X18))|(~m1_subset_1(X21,u1_struct_0(X18))|(~r2_hidden(X20,u1_struct_0(X19))|(~r2_hidden(X21,u1_struct_0(X19))|~r2_yellow_0(X18,k7_domain_1(u1_struct_0(X18),X20,X21)))|r2_hidden(k2_yellow_0(X18,k7_domain_1(u1_struct_0(X18),X20,X21)),u1_struct_0(X19)))))|~m1_yellow_0(X19,X18)|(v2_struct_0(X18)|~l1_orders_2(X18)))&((m1_subset_1(esk5_2(X18,X19),u1_struct_0(X18))|v5_yellow_0(X19,X18)|~m1_yellow_0(X19,X18)|(v2_struct_0(X18)|~l1_orders_2(X18)))&((m1_subset_1(esk6_2(X18,X19),u1_struct_0(X18))|v5_yellow_0(X19,X18)|~m1_yellow_0(X19,X18)|(v2_struct_0(X18)|~l1_orders_2(X18)))&(((r2_hidden(esk5_2(X18,X19),u1_struct_0(X19))|v5_yellow_0(X19,X18)|~m1_yellow_0(X19,X18)|(v2_struct_0(X18)|~l1_orders_2(X18)))&((r2_hidden(esk6_2(X18,X19),u1_struct_0(X19))|v5_yellow_0(X19,X18)|~m1_yellow_0(X19,X18)|(v2_struct_0(X18)|~l1_orders_2(X18)))&(r2_yellow_0(X18,k7_domain_1(u1_struct_0(X18),esk5_2(X18,X19),esk6_2(X18,X19)))|v5_yellow_0(X19,X18)|~m1_yellow_0(X19,X18)|(v2_struct_0(X18)|~l1_orders_2(X18)))))&(~r2_hidden(k2_yellow_0(X18,k7_domain_1(u1_struct_0(X18),esk5_2(X18,X19),esk6_2(X18,X19))),u1_struct_0(X19))|v5_yellow_0(X19,X18)|~m1_yellow_0(X19,X18)|(v2_struct_0(X18)|~l1_orders_2(X18))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_yellow_0])])])])])])).
cnf(c_0_13, plain, (u1_struct_0(X1)=X3|X1!=k5_yellow_0(X2,X3)|~v1_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (v1_orders_2(k5_yellow_0(X1,X2))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (v4_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (m1_yellow_0(k5_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (v5_yellow_0(X2,X1)|v2_struct_0(X1)|~r2_hidden(k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),esk5_2(X1,X2),esk6_2(X1,X2))),u1_struct_0(X2))|~m1_yellow_0(X2,X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (u1_struct_0(k5_yellow_0(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_14]), c_0_15]), c_0_16])).
fof(c_0_19, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v1_xboole_0(X2))&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))))=>((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))<=>v5_yellow_0(k5_yellow_0(X1,X2),X1))))), inference(assume_negation,[status(cth)],[t63_waybel_0])).
cnf(c_0_20, plain, (v2_struct_0(X1)|v5_yellow_0(k5_yellow_0(X2,X3),X1)|~r2_hidden(k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),esk5_2(X1,k5_yellow_0(X2,X3)),esk6_2(X1,k5_yellow_0(X2,X3)))),X3)|~m1_yellow_0(k5_yellow_0(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X1)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_21, negated_conjecture, ((v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0)))))&((~v1_xboole_0(esk2_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))))&((v1_xboole_0(esk2_0)|(~v2_waybel_0(esk2_0,esk1_0)|(~v13_waybel_0(esk2_0,esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))))|~v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0))&((~v1_xboole_0(esk2_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0))&((v2_waybel_0(esk2_0,esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0))&((v13_waybel_0(esk2_0,esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])])).
cnf(c_0_22, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v5_yellow_0(X2,X1)|v2_struct_0(X1)|~m1_yellow_0(X2,X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, plain, (m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|v5_yellow_0(X2,X1)|v2_struct_0(X1)|~m1_yellow_0(X2,X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24, plain, (r2_hidden(esk5_2(X1,X2),u1_struct_0(X2))|v5_yellow_0(X2,X1)|v2_struct_0(X1)|~m1_yellow_0(X2,X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, plain, (r2_hidden(esk6_2(X1,X2),u1_struct_0(X2))|v5_yellow_0(X2,X1)|v2_struct_0(X1)|~m1_yellow_0(X2,X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_26, plain, (v2_struct_0(X1)|v5_yellow_0(k5_yellow_0(X1,X2),X1)|~r2_hidden(k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),esk5_2(X1,k5_yellow_0(X1,X2)),esk6_2(X1,k5_yellow_0(X1,X2)))),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_16])).
cnf(c_0_27, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_29, plain, ![X15, X16, X17]:(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|(~v2_lattice3(X15)|~l1_orders_2(X15))))|(~m1_subset_1(X16,u1_struct_0(X15))|(~m1_subset_1(X17,u1_struct_0(X15))|k2_yellow_0(X15,k7_domain_1(u1_struct_0(X15),X16,X17))=k12_lattice3(X15,X16,X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t40_yellow_0])])])).
cnf(c_0_30, plain, (v2_struct_0(X1)|v5_yellow_0(k5_yellow_0(X1,X2),X1)|m1_subset_1(esk5_2(X1,k5_yellow_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_22, c_0_16])).
cnf(c_0_31, plain, (v2_struct_0(X1)|v5_yellow_0(k5_yellow_0(X1,X2),X1)|m1_subset_1(esk6_2(X1,k5_yellow_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_23, c_0_16])).
cnf(c_0_32, plain, (r2_hidden(esk5_2(X1,k5_yellow_0(X1,X2)),u1_struct_0(k5_yellow_0(X1,X2)))|v2_struct_0(X1)|v5_yellow_0(k5_yellow_0(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_24, c_0_16])).
cnf(c_0_33, plain, (r2_hidden(esk6_2(X1,k5_yellow_0(X1,X2)),u1_struct_0(k5_yellow_0(X1,X2)))|v2_struct_0(X1)|v5_yellow_0(k5_yellow_0(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_16])).
cnf(c_0_34, negated_conjecture, (v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)|~r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk5_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)),esk6_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)))),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_35, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_37, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_38, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_39, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_40, negated_conjecture, (v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)|m1_subset_1(esk5_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_27]), c_0_28])])).
cnf(c_0_41, negated_conjecture, (v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)|m1_subset_1(esk6_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_27]), c_0_28])])).
fof(c_0_42, plain, ![X32, X33, X34, X35]:((~v2_waybel_0(X33,X32)|(~m1_subset_1(X34,u1_struct_0(X32))|(~m1_subset_1(X35,u1_struct_0(X32))|(~r2_hidden(X34,X33)|~r2_hidden(X35,X33)|r2_hidden(k12_lattice3(X32,X34,X35),X33))))|(~v13_waybel_0(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32))))|(~v5_orders_2(X32)|(~v2_lattice3(X32)|~l1_orders_2(X32))))&((m1_subset_1(esk7_2(X32,X33),u1_struct_0(X32))|v2_waybel_0(X33,X32)|(~v13_waybel_0(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32))))|(~v5_orders_2(X32)|(~v2_lattice3(X32)|~l1_orders_2(X32))))&((m1_subset_1(esk8_2(X32,X33),u1_struct_0(X32))|v2_waybel_0(X33,X32)|(~v13_waybel_0(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32))))|(~v5_orders_2(X32)|(~v2_lattice3(X32)|~l1_orders_2(X32))))&(((r2_hidden(esk7_2(X32,X33),X33)|v2_waybel_0(X33,X32)|(~v13_waybel_0(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32))))|(~v5_orders_2(X32)|(~v2_lattice3(X32)|~l1_orders_2(X32))))&(r2_hidden(esk8_2(X32,X33),X33)|v2_waybel_0(X33,X32)|(~v13_waybel_0(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32))))|(~v5_orders_2(X32)|(~v2_lattice3(X32)|~l1_orders_2(X32)))))&(~r2_hidden(k12_lattice3(X32,esk7_2(X32,X33),esk8_2(X32,X33)),X33)|v2_waybel_0(X33,X32)|(~v13_waybel_0(X33,X32)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32))))|(~v5_orders_2(X32)|(~v2_lattice3(X32)|~l1_orders_2(X32)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t41_waybel_0])])])])])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk5_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)),u1_struct_0(k5_yellow_0(esk1_0,esk2_0)))|v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_27]), c_0_28])])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk6_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)),u1_struct_0(k5_yellow_0(esk1_0,esk2_0)))|v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_27]), c_0_28])])).
cnf(c_0_45, negated_conjecture, (v1_xboole_0(esk2_0)|~v2_waybel_0(esk2_0,esk1_0)|~v13_waybel_0(esk2_0,esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_46, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_47, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_48, negated_conjecture, (v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)|~r2_hidden(k12_lattice3(esk1_0,esk5_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)),esk6_2(esk1_0,k5_yellow_0(esk1_0,esk2_0))),esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_28]), c_0_36]), c_0_37]), c_0_38]), c_0_39])]), c_0_40]), c_0_41])).
cnf(c_0_49, plain, (r2_hidden(k12_lattice3(X2,X3,X4),X1)|~v2_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, negated_conjecture, (v2_waybel_0(esk2_0,esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk5_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)),esk2_0)|v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_18]), c_0_27]), c_0_28])])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk6_2(esk1_0,k5_yellow_0(esk1_0,esk2_0)),esk2_0)|v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_18]), c_0_27]), c_0_28])])).
fof(c_0_53, plain, ![X24, X25, X26]:(v1_xboole_0(X24)|(~m1_subset_1(X25,X24)|~m1_subset_1(X26,X24))|k7_domain_1(X24,X25,X26)=k2_tarski(X25,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k7_domain_1])])])).
cnf(c_0_54, plain, (m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_55, negated_conjecture, (~v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)|~v2_waybel_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46]), c_0_27])]), c_0_47])).
cnf(c_0_56, negated_conjecture, (v2_struct_0(esk1_0)|v5_yellow_0(k5_yellow_0(esk1_0,esk2_0),esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_27]), c_0_46]), c_0_28]), c_0_36]), c_0_37])]), c_0_40]), c_0_41]), c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_57, plain, (v1_xboole_0(X1)|k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_58, negated_conjecture, (v2_waybel_0(esk2_0,esk1_0)|m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_27]), c_0_46]), c_0_28]), c_0_36]), c_0_37])])).
cnf(c_0_59, plain, (m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_60, plain, (r2_hidden(k2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X1))|v2_struct_0(X2)|~v5_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,u1_struct_0(X1))|~r2_hidden(X4,u1_struct_0(X1))|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_61, negated_conjecture, (v2_struct_0(esk1_0)|~v2_waybel_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_62, plain, ![X10, X11, X12]:((~v2_lattice3(X10)|(~m1_subset_1(X11,u1_struct_0(X10))|(~m1_subset_1(X12,u1_struct_0(X10))|r2_yellow_0(X10,k2_tarski(X11,X12))))|(~v5_orders_2(X10)|~l1_orders_2(X10)))&((m1_subset_1(esk3_1(X10),u1_struct_0(X10))|v2_lattice3(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&((m1_subset_1(esk4_1(X10),u1_struct_0(X10))|v2_lattice3(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(~r2_yellow_0(X10,k2_tarski(esk3_1(X10),esk4_1(X10)))|v2_lattice3(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_yellow_0])])])])])).
cnf(c_0_63, negated_conjecture, (k2_tarski(X1,esk8_2(esk1_0,esk2_0))=k7_domain_1(u1_struct_0(esk1_0),X1,esk8_2(esk1_0,esk2_0))|v2_waybel_0(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_64, negated_conjecture, (v2_waybel_0(esk2_0,esk1_0)|m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_27]), c_0_46]), c_0_28]), c_0_36]), c_0_37])])).
cnf(c_0_65, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk8_2(esk1_0,esk2_0))),u1_struct_0(X2))|v2_struct_0(esk1_0)|~r2_hidden(esk8_2(esk1_0,esk2_0),u1_struct_0(X2))|~r2_hidden(X1,u1_struct_0(X2))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk8_2(esk1_0,esk2_0)))|~m1_yellow_0(X2,esk1_0)|~v5_yellow_0(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_58]), c_0_28])]), c_0_61])).
cnf(c_0_66, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~v2_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_67, negated_conjecture, (k2_tarski(esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))=k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))|v2_waybel_0(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk8_2(esk1_0,esk2_0))),u1_struct_0(k5_yellow_0(esk1_0,X2)))|v2_struct_0(esk1_0)|~r2_hidden(esk8_2(esk1_0,esk2_0),u1_struct_0(k5_yellow_0(esk1_0,X2)))|~r2_hidden(X1,u1_struct_0(k5_yellow_0(esk1_0,X2)))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk8_2(esk1_0,esk2_0)))|~v5_yellow_0(k5_yellow_0(esk1_0,X2),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_16]), c_0_28])])).
cnf(c_0_69, negated_conjecture, (r2_yellow_0(X1,k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0)))|v2_waybel_0(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(X1))|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_70, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))),u1_struct_0(k5_yellow_0(esk1_0,X1)))|v2_struct_0(esk1_0)|~r2_hidden(esk8_2(esk1_0,esk2_0),u1_struct_0(k5_yellow_0(esk1_0,X1)))|~r2_hidden(esk7_2(esk1_0,esk2_0),u1_struct_0(k5_yellow_0(esk1_0,X1)))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0)))|~v5_yellow_0(k5_yellow_0(esk1_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_64]), c_0_61])).
cnf(c_0_71, negated_conjecture, (r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0)))|v2_waybel_0(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_58]), c_0_28]), c_0_36]), c_0_37])]), c_0_64])).
fof(c_0_72, plain, ![X27, X28, X29]:(~r2_hidden(X27,X28)|(~m1_subset_1(X28,k1_zfmisc_1(X29))|~v1_xboole_0(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_73, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))),u1_struct_0(k5_yellow_0(esk1_0,X1)))|v2_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(esk8_2(esk1_0,esk2_0),u1_struct_0(k5_yellow_0(esk1_0,X1)))|~r2_hidden(esk7_2(esk1_0,esk2_0),u1_struct_0(k5_yellow_0(esk1_0,X1)))|~v5_yellow_0(k5_yellow_0(esk1_0,X1),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_61])).
cnf(c_0_74, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_75, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))),u1_struct_0(k5_yellow_0(esk1_0,esk2_0)))|v2_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(esk8_2(esk1_0,esk2_0),u1_struct_0(k5_yellow_0(esk1_0,esk2_0)))|~r2_hidden(esk7_2(esk1_0,esk2_0),u1_struct_0(k5_yellow_0(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_56]), c_0_27])])).
cnf(c_0_76, negated_conjecture, (~r2_hidden(X1,esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_74, c_0_27])).
cnf(c_0_77, plain, (r2_hidden(esk8_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_78, plain, (r2_hidden(esk7_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_79, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))),esk2_0)|v2_struct_0(esk1_0)|~r2_hidden(esk8_2(esk1_0,esk2_0),esk2_0)|~r2_hidden(esk7_2(esk1_0,esk2_0),esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_18]), c_0_27]), c_0_28])]), c_0_76])).
cnf(c_0_80, negated_conjecture, (r2_hidden(esk8_2(esk1_0,esk2_0),esk2_0)|v2_waybel_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_27]), c_0_46]), c_0_28]), c_0_36]), c_0_37])])).
cnf(c_0_81, negated_conjecture, (r2_hidden(esk7_2(esk1_0,esk2_0),esk2_0)|v2_waybel_0(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_27]), c_0_46]), c_0_28]), c_0_36]), c_0_37])])).
cnf(c_0_82, negated_conjecture, (r2_hidden(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))),esk2_0)|v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81]), c_0_61])).
cnf(c_0_83, negated_conjecture, (r2_hidden(k12_lattice3(esk1_0,esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0)),esk2_0)|v2_struct_0(esk1_0)|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_35]), c_0_28]), c_0_36]), c_0_37]), c_0_38]), c_0_39])])).
fof(c_0_84, plain, ![X38]:(~l1_orders_2(X38)|(~v2_lattice3(X38)|~v2_struct_0(X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
cnf(c_0_85, plain, (v2_waybel_0(X2,X1)|~r2_hidden(k12_lattice3(X1,esk7_2(X1,X2),esk8_2(X1,X2)),X2)|~v13_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_86, negated_conjecture, (r2_hidden(k12_lattice3(esk1_0,esk7_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0)),esk2_0)|v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_58]), c_0_64]), c_0_61])).
cnf(c_0_87, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_88, negated_conjecture, (v2_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_27]), c_0_46]), c_0_28]), c_0_36]), c_0_37])]), c_0_61])).
cnf(c_0_89, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_28]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 69
# Proof object formula steps           : 21
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 37
# Proof object simplifying inferences  : 100
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 184
# ...of these trivial                  : 3
# ...subsumed                          : 23
# ...remaining for further processing  : 158
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 18
# Backward-rewritten                   : 21
# Generated clauses                    : 181
# ...of the previous two non-trivial   : 151
# Contextual simplify-reflections      : 39
# Paramodulations                      : 178
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 83
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 74
# Current number of unprocessed clauses: 37
# ...number of literals in the above   : 319
# Current number of archived formulas  : 0
# Current number of archived clauses   : 75
# Clause-clause subsumption calls (NU) : 3495
# Rec. Clause-clause subsumption calls : 253
# Non-unit clause-clause subsumptions  : 79
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12146

# -------------------------------------------------
# User time                : 0.018 s
# System time              : 0.000 s
# Total time               : 0.018 s
# Maximum resident set size: 28640 pages
