# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t52_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X3))))))), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', t52_orders_2)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', t35_orders_2)).
fof(fraenkel_a_2_1_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&l1_orders_2(X2)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_orders_2(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r2_hidden(X5,X3)=>r2_orders_2(X2,X4,X5))))))), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', fraenkel_a_2_1_orders_2)).
fof(d13_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2))), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', d13_orders_2)).
fof(t50_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2)))))), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', t50_orders_2)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', t69_enumset1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', t5_subset)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/8dd43894f61348b5a6843e5dbaf27f50.E_input', d1_tarski)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X3)))))))), inference(assume_negation,[status(cth)],[t52_orders_2])).
fof(c_0_10, plain, ![X9, X10]:((v6_orders_2(k6_domain_1(u1_struct_0(X9),X10),X9)|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|(~v3_orders_2(X9)|~l1_orders_2(X9))))&(m1_subset_1(k6_domain_1(u1_struct_0(X9),X10),k1_zfmisc_1(u1_struct_0(X9)))|~m1_subset_1(X10,u1_struct_0(X9))|(v2_struct_0(X9)|(~v3_orders_2(X9)|~l1_orders_2(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r2_orders_2(esk1_0,esk2_0,esk3_0)|~r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))))&(r2_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
fof(c_0_12, plain, ![X11, X12, X13, X15, X16]:(((m1_subset_1(esk4_3(X11,X12,X13),u1_struct_0(X12))|~r2_hidden(X11,a_2_1_orders_2(X12,X13))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|(~v5_orders_2(X12)|~l1_orders_2(X12))))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&((X11=esk4_3(X11,X12,X13)|~r2_hidden(X11,a_2_1_orders_2(X12,X13))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|(~v5_orders_2(X12)|~l1_orders_2(X12))))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&(~m1_subset_1(X15,u1_struct_0(X12))|(~r2_hidden(X15,X13)|r2_orders_2(X12,esk4_3(X11,X12,X13),X15))|~r2_hidden(X11,a_2_1_orders_2(X12,X13))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|(~v5_orders_2(X12)|~l1_orders_2(X12))))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))))&((m1_subset_1(esk5_4(X11,X12,X13,X16),u1_struct_0(X12))|X11!=X16|~m1_subset_1(X16,u1_struct_0(X12))|r2_hidden(X11,a_2_1_orders_2(X12,X13))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|(~v5_orders_2(X12)|~l1_orders_2(X12))))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&((r2_hidden(esk5_4(X11,X12,X13,X16),X13)|X11!=X16|~m1_subset_1(X16,u1_struct_0(X12))|r2_hidden(X11,a_2_1_orders_2(X12,X13))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|(~v5_orders_2(X12)|~l1_orders_2(X12))))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))&(~r2_orders_2(X12,X16,esk5_4(X11,X12,X13,X16))|X11!=X16|~m1_subset_1(X16,u1_struct_0(X12))|r2_hidden(X11,a_2_1_orders_2(X12,X13))|(v2_struct_0(X12)|(~v3_orders_2(X12)|(~v4_orders_2(X12)|(~v5_orders_2(X12)|~l1_orders_2(X12))))|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X12)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_orders_2])])])])])])).
fof(c_0_13, plain, ![X18, X19]:(v2_struct_0(X18)|(~v3_orders_2(X18)|(~v4_orders_2(X18)|(~v5_orders_2(X18)|~l1_orders_2(X18))))|(~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X18)))|k2_orders_2(X18,X19)=a_2_1_orders_2(X18,X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_orders_2])])])])).
cnf(c_0_14, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (r2_hidden(esk5_4(X1,X2,X3,X4),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|X1!=X4|~m1_subset_1(X4,u1_struct_0(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_20, plain, ![X20, X21]:(v2_struct_0(X20)|(~v3_orders_2(X20)|(~v4_orders_2(X20)|(~v5_orders_2(X20)|~l1_orders_2(X20))))|(~m1_subset_1(X21,u1_struct_0(X20))|~r2_hidden(X21,k2_orders_2(X20,k6_domain_1(u1_struct_0(X20),X21))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_orders_2])])])])).
cnf(c_0_21, plain, (v2_struct_0(X1)|k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_23, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25, plain, (r2_hidden(esk5_4(X1,X2,X3,X1),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2)))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, negated_conjecture, (k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))=a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_16]), c_0_23]), c_0_24]), c_0_17])]), c_0_18])).
fof(c_0_28, plain, ![X22, X23]:(v1_xboole_0(X22)|~m1_subset_1(X23,X22)|k6_domain_1(X22,X23)=k1_tarski(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_29, plain, ![X24]:k2_tarski(X24,X24)=k1_tarski(X24), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_30, plain, ![X25, X26, X27]:(~r2_hidden(X25,X26)|(~m1_subset_1(X26,k1_zfmisc_1(X27))|~v1_xboole_0(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk5_4(esk3_0,esk1_0,X1,esk3_0),X1)|r2_hidden(esk3_0,a_2_1_orders_2(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_15]), c_0_16]), c_0_23]), c_0_24]), c_0_17])]), c_0_18])).
cnf(c_0_32, negated_conjecture, (~r2_hidden(esk3_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_15]), c_0_16]), c_0_23]), c_0_24]), c_0_17])]), c_0_18])).
fof(c_0_33, plain, ![X28, X29, X30, X31, X32, X33]:(((~r2_hidden(X30,X29)|X30=X28|X29!=k1_tarski(X28))&(X31!=X28|r2_hidden(X31,X29)|X29!=k1_tarski(X28)))&((~r2_hidden(esk6_2(X32,X33),X33)|esk6_2(X32,X33)!=X32|X33=k1_tarski(X32))&(r2_hidden(esk6_2(X32,X33),X33)|esk6_2(X32,X33)=X32|X33=k1_tarski(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_34, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk5_4(esk3_0,esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0),esk3_0),k6_domain_1(u1_struct_0(esk1_0),esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_22]), c_0_32])).
cnf(c_0_38, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, negated_conjecture, (~v1_xboole_0(X1)|~m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk3_0),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_38, c_0_35])).
cnf(c_0_42, negated_conjecture, (k2_tarski(esk3_0,esk3_0)=k6_domain_1(u1_struct_0(esk1_0),esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_39, c_0_15])).
cnf(c_0_43, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_22])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_45, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_47, negated_conjecture, (k2_tarski(esk3_0,esk3_0)=k6_domain_1(u1_struct_0(esk1_0),esk3_0)), inference(sr,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk5_4(esk2_0,esk1_0,X1,esk2_0),X1)|r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_44]), c_0_16]), c_0_23]), c_0_24]), c_0_17])]), c_0_18])).
cnf(c_0_49, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_45, c_0_35])).
cnf(c_0_50, plain, (r2_hidden(X3,a_2_1_orders_2(X1,X4))|v2_struct_0(X1)|~r2_orders_2(X1,X2,esk5_4(X3,X1,X4,X2))|X3!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_51, negated_conjecture, (X1=esk3_0|~r2_hidden(X1,k6_domain_1(u1_struct_0(esk1_0),esk3_0))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk5_4(esk2_0,esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0),esk2_0),k6_domain_1(u1_struct_0(esk1_0),esk3_0))|r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(spm,[status(thm)],[c_0_48, c_0_22])).
cnf(c_0_53, negated_conjecture, (r2_orders_2(esk1_0,esk2_0,esk3_0)|r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_54, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_49])])).
cnf(c_0_55, plain, (r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|~r2_orders_2(X2,X1,esk5_4(X1,X2,X3,X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_56, negated_conjecture, (esk5_4(esk2_0,esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0),esk2_0)=esk3_0|r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))|r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_53, c_0_27])).
cnf(c_0_58, negated_conjecture, (~r2_orders_2(esk1_0,esk2_0,esk3_0)|~r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_59, plain, (r2_orders_2(X2,esk4_3(X4,X2,X3),X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_hidden(X4,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk3_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))), inference(spm,[status(thm)],[c_0_54, c_0_47])).
cnf(c_0_61, plain, (X1=esk4_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_62, negated_conjecture, (r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_22]), c_0_44]), c_0_16]), c_0_23]), c_0_24]), c_0_17])]), c_0_18]), c_0_57])).
cnf(c_0_63, negated_conjecture, (~r2_hidden(esk2_0,a_2_1_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))|~r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_58, c_0_27])).
cnf(c_0_64, negated_conjecture, (r2_orders_2(X1,esk4_3(X2,X1,k6_domain_1(u1_struct_0(esk1_0),esk3_0)),esk3_0)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_orders_2(X1,k6_domain_1(u1_struct_0(esk1_0),esk3_0)))|~m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk3_0),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk3_0,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_65, negated_conjecture, (esk4_3(esk2_0,esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_22]), c_0_16]), c_0_23]), c_0_24]), c_0_17])]), c_0_18])).
cnf(c_0_66, negated_conjecture, (~r2_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_62])])).
cnf(c_0_67, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_62]), c_0_65]), c_0_22]), c_0_15]), c_0_16]), c_0_23]), c_0_24]), c_0_17])]), c_0_66]), c_0_18]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 49
# Proof object formula steps           : 19
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 69
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 25
# Processed clauses                    : 143
# ...of these trivial                  : 1
# ...subsumed                          : 16
# ...remaining for further processing  : 126
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 17
# Generated clauses                    : 169
# ...of the previous two non-trivial   : 134
# Contextual simplify-reflections      : 2
# Paramodulations                      : 147
# Factorizations                       : 6
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 77
#    Positive orientable unit clauses  : 21
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 49
# Current number of unprocessed clauses: 34
# ...number of literals in the above   : 165
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 1177
# Rec. Clause-clause subsumption calls : 214
# Non-unit clause-clause subsumptions  : 14
# Unit Clause-clause subsumption calls : 43
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7440

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28640 pages
