# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t16_yellow_0, conjecture, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/8de94f09d4094ad2a72207ee7e01b044.E_input', t16_yellow_0)).
fof(d8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X5)=>r1_orders_2(X1,X5,X4))))=>X4=X3))))))), file('data/MPTP/proofs/8de94f09d4094ad2a72207ee7e01b044.E_input', d8_yellow_0)).
fof(t25_orders_2, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X3)&r1_orders_2(X1,X3,X2))=>X2=X3)))), file('data/MPTP/proofs/8de94f09d4094ad2a72207ee7e01b044.E_input', t25_orders_2)).
fof(c_0_3, negated_conjecture, ~(![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), inference(assume_negation,[status(cth)],[t16_yellow_0])).
fof(c_0_4, negated_conjecture, ![X8, X11]:((v5_orders_2(esk1_0)&l1_orders_2(esk1_0))&(((m1_subset_1(esk3_1(X8),u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk2_0,X8)|~m1_subset_1(X8,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0))&((r1_lattice3(esk1_0,esk2_0,esk3_1(X8))|~r1_lattice3(esk1_0,esk2_0,X8)|~m1_subset_1(X8,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0))&(~r1_orders_2(esk1_0,esk3_1(X8),X8)|~r1_lattice3(esk1_0,esk2_0,X8)|~m1_subset_1(X8,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0))))&((m1_subset_1(esk4_0,u1_struct_0(esk1_0))|r2_yellow_0(esk1_0,esk2_0))&((r1_lattice3(esk1_0,esk2_0,esk4_0)|r2_yellow_0(esk1_0,esk2_0))&(~m1_subset_1(X11,u1_struct_0(esk1_0))|(~r1_lattice3(esk1_0,esk2_0,X11)|r1_orders_2(esk1_0,X11,esk4_0))|r2_yellow_0(esk1_0,esk2_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
fof(c_0_5, plain, ![X12, X13, X15, X16, X18, X19, X22]:(((m1_subset_1(esk5_2(X12,X13),u1_struct_0(X12))|~r2_yellow_0(X12,X13)|~l1_orders_2(X12))&((r1_lattice3(X12,X13,esk5_2(X12,X13))|~r2_yellow_0(X12,X13)|~l1_orders_2(X12))&((~m1_subset_1(X15,u1_struct_0(X12))|(~r1_lattice3(X12,X13,X15)|r1_orders_2(X12,X15,esk5_2(X12,X13)))|~r2_yellow_0(X12,X13)|~l1_orders_2(X12))&((m1_subset_1(esk6_3(X12,X13,X16),u1_struct_0(X12))|~r1_lattice3(X12,X13,X16)|X16=esk5_2(X12,X13)|~m1_subset_1(X16,u1_struct_0(X12))|~r2_yellow_0(X12,X13)|~l1_orders_2(X12))&((r1_lattice3(X12,X13,esk6_3(X12,X13,X16))|~r1_lattice3(X12,X13,X16)|X16=esk5_2(X12,X13)|~m1_subset_1(X16,u1_struct_0(X12))|~r2_yellow_0(X12,X13)|~l1_orders_2(X12))&(~r1_orders_2(X12,esk6_3(X12,X13,X16),X16)|~r1_lattice3(X12,X13,X16)|X16=esk5_2(X12,X13)|~m1_subset_1(X16,u1_struct_0(X12))|~r2_yellow_0(X12,X13)|~l1_orders_2(X12)))))))&(((m1_subset_1(esk8_3(X12,X18,X19),u1_struct_0(X12))|m1_subset_1(esk7_3(X12,X18,X19),u1_struct_0(X12))|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))&(((r1_lattice3(X12,X18,esk8_3(X12,X18,X19))|m1_subset_1(esk7_3(X12,X18,X19),u1_struct_0(X12))|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))&(~m1_subset_1(X22,u1_struct_0(X12))|(~r1_lattice3(X12,X18,X22)|r1_orders_2(X12,X22,esk8_3(X12,X18,X19)))|m1_subset_1(esk7_3(X12,X18,X19),u1_struct_0(X12))|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12)))&(esk8_3(X12,X18,X19)!=X19|m1_subset_1(esk7_3(X12,X18,X19),u1_struct_0(X12))|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))))&(((m1_subset_1(esk8_3(X12,X18,X19),u1_struct_0(X12))|r1_lattice3(X12,X18,esk7_3(X12,X18,X19))|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))&(((r1_lattice3(X12,X18,esk8_3(X12,X18,X19))|r1_lattice3(X12,X18,esk7_3(X12,X18,X19))|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))&(~m1_subset_1(X22,u1_struct_0(X12))|(~r1_lattice3(X12,X18,X22)|r1_orders_2(X12,X22,esk8_3(X12,X18,X19)))|r1_lattice3(X12,X18,esk7_3(X12,X18,X19))|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12)))&(esk8_3(X12,X18,X19)!=X19|r1_lattice3(X12,X18,esk7_3(X12,X18,X19))|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))))&((m1_subset_1(esk8_3(X12,X18,X19),u1_struct_0(X12))|~r1_orders_2(X12,esk7_3(X12,X18,X19),X19)|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))&(((r1_lattice3(X12,X18,esk8_3(X12,X18,X19))|~r1_orders_2(X12,esk7_3(X12,X18,X19),X19)|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))&(~m1_subset_1(X22,u1_struct_0(X12))|(~r1_lattice3(X12,X18,X22)|r1_orders_2(X12,X22,esk8_3(X12,X18,X19)))|~r1_orders_2(X12,esk7_3(X12,X18,X19),X19)|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12)))&(esk8_3(X12,X18,X19)!=X19|~r1_orders_2(X12,esk7_3(X12,X18,X19),X19)|~r1_lattice3(X12,X18,X19)|~m1_subset_1(X19,u1_struct_0(X12))|r2_yellow_0(X12,X18)|~l1_orders_2(X12))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_0])])])])])])).
cnf(c_0_6, negated_conjecture, (~r1_orders_2(esk1_0,esk3_1(X1),X1)|~r1_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, plain, (r1_orders_2(X2,X1,esk5_2(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9, negated_conjecture, (m1_subset_1(esk3_1(X1),u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10, negated_conjecture, (~r1_lattice3(esk1_0,X1,esk3_1(esk5_2(esk1_0,X1)))|~r1_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,X1))|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0)|~r2_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8])]), c_0_9])).
cnf(c_0_11, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_1(X1))|~r1_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12, negated_conjecture, (~r1_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_13, plain, (r1_lattice3(X1,X2,esk5_2(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14, negated_conjecture, (~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_8])])).
cnf(c_0_15, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))|r2_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17, negated_conjecture, (~r2_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_8])])).
cnf(c_0_18, plain, (r1_orders_2(X2,X1,esk8_3(X2,X3,X4))|r1_lattice3(X2,X3,esk7_3(X2,X3,X4))|r2_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r1_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_20, plain, ![X23, X24, X25]:(~v5_orders_2(X23)|~l1_orders_2(X23)|(~m1_subset_1(X24,u1_struct_0(X23))|(~m1_subset_1(X25,u1_struct_0(X23))|(~r1_orders_2(X23,X24,X25)|~r1_orders_2(X23,X25,X24)|X24=X25)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_orders_2])])])).
cnf(c_0_21, negated_conjecture, (r1_orders_2(esk1_0,X1,esk8_3(esk1_0,X2,esk4_0))|r1_lattice3(esk1_0,X2,esk7_3(esk1_0,X2,esk4_0))|r2_yellow_0(esk1_0,X2)|~r1_lattice3(esk1_0,X2,esk4_0)|~r1_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_8])])).
cnf(c_0_22, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_0)|r2_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_23, plain, (X2=X3|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, negated_conjecture, (r1_orders_2(esk1_0,X1,esk4_0)|r2_yellow_0(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_25, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_26, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk8_3(esk1_0,X1,esk4_0))|r1_lattice3(esk1_0,X1,esk7_3(esk1_0,X1,esk4_0))|r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_21, c_0_19])).
cnf(c_0_27, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk4_0)), inference(sr,[status(thm)],[c_0_22, c_0_17])).
cnf(c_0_28, negated_conjecture, (esk4_0=X1|r2_yellow_0(esk1_0,esk2_0)|~r1_orders_2(esk1_0,esk4_0,X1)|~r1_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_8]), c_0_25])]), c_0_16])).
cnf(c_0_29, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,esk8_3(esk1_0,esk2_0,esk4_0))|r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_17])).
cnf(c_0_30, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_31, plain, (r1_lattice3(X1,X2,esk8_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_32, plain, (r1_orders_2(X2,X1,esk8_3(X2,X3,X4))|r2_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r1_orders_2(X2,esk7_3(X2,X3,X4),X4)|~r1_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_33, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk4_0)=esk4_0|r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~r1_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_17])).
cnf(c_0_34, plain, (r1_lattice3(X1,X2,esk8_3(X1,X2,X3))|r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk8_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|r2_yellow_0(esk1_0,esk2_0)|r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r1_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_24]), c_0_8])]), c_0_16])).
cnf(c_0_36, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_37, negated_conjecture, (r1_lattice3(esk1_0,X1,esk8_3(esk1_0,X1,esk4_0))|r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r1_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_24]), c_0_8])]), c_0_19])]), c_0_17])).
cnf(c_0_38, plain, (r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|r2_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_39, plain, (r1_orders_2(X2,X1,esk8_3(X2,X3,X4))|m1_subset_1(esk7_3(X2,X3,X4),u1_struct_0(X2))|r2_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r1_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_40, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk1_0,X1,esk8_3(esk1_0,X2,esk4_0))|r2_yellow_0(esk1_0,X2)|~r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X2,esk4_0))|~r1_lattice3(esk1_0,X2,esk4_0)|~r1_lattice3(esk1_0,X2,X1)|~m1_subset_1(esk7_3(esk1_0,X2,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_24]), c_0_8])]), c_0_19])]), c_0_17])).
cnf(c_0_42, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk4_0)=esk4_0|r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_27]), c_0_19]), c_0_8])]), c_0_17])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_8])]), c_0_27]), c_0_19])]), c_0_17])).
cnf(c_0_44, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,esk4_0))|esk8_3(esk1_0,esk2_0,esk4_0)!=esk4_0|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_27]), c_0_19]), c_0_8])]), c_0_17])).
cnf(c_0_45, plain, (esk8_3(X1,X2,X3)=X4|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk8_3(X1,X2,X3),X4)|~r1_lattice3(X1,X2,X3)|~r1_lattice3(X1,X2,X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_39]), c_0_40])).
cnf(c_0_46, plain, (r2_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r1_orders_2(X1,esk7_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_47, negated_conjecture, (esk8_3(esk1_0,X1,esk4_0)=esk4_0|r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,X1,esk4_0))|~r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r1_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk8_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_41]), c_0_19])]), c_0_17])).
cnf(c_0_48, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_42]), c_0_27])]), c_0_17]), c_0_43]), c_0_44])).
cnf(c_0_49, negated_conjecture, (esk8_3(esk1_0,X1,X2)=esk4_0|m1_subset_1(esk7_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,X1,X2))|~r1_lattice3(esk1_0,X1,esk4_0)|~r1_lattice3(esk1_0,X1,X2)|~m1_subset_1(esk8_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_24]), c_0_19]), c_0_8]), c_0_25])]), c_0_17])).
cnf(c_0_50, plain, (r1_lattice3(X1,X2,esk8_3(X1,X2,X3))|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_51, negated_conjecture, (r2_yellow_0(esk1_0,esk2_0)|r2_yellow_0(esk1_0,X1)|esk8_3(esk1_0,X1,esk4_0)!=esk4_0|~r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r1_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_24]), c_0_8])]), c_0_16])).
cnf(c_0_52, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk4_0)=esk4_0|~r1_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_27])]), c_0_17]), c_0_43])).
cnf(c_0_53, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk4_0|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_27]), c_0_8])]), c_0_17])).
cnf(c_0_54, negated_conjecture, (r2_yellow_0(esk1_0,esk2_0)|esk8_3(esk1_0,esk2_0,esk4_0)!=esk4_0|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_38]), c_0_8])]), c_0_16]), c_0_22])).
cnf(c_0_55, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_56, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk4_0)=esk4_0|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_42]), c_0_43])).
cnf(c_0_57, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk4_0|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_40]), c_0_8])]), c_0_17])).
cnf(c_0_58, negated_conjecture, (r2_yellow_0(esk1_0,esk2_0)|esk8_3(esk1_0,esk2_0,esk4_0)!=esk4_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_8])]), c_0_16]), c_0_22])).
cnf(c_0_59, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk4_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_27]), c_0_19])])).
cnf(c_0_60, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])]), c_0_17]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 61
# Proof object clause steps            : 54
# Proof object formula steps           : 7
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 89
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 102
# ...of these trivial                  : 1
# ...subsumed                          : 4
# ...remaining for further processing  : 97
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 8
# Generated clauses                    : 109
# ...of the previous two non-trivial   : 86
# Contextual simplify-reflections      : 22
# Paramodulations                      : 107
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 52
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 46
# Current number of unprocessed clauses: 32
# ...number of literals in the above   : 323
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 1262
# Rec. Clause-clause subsumption calls : 67
# Non-unit clause-clause subsumptions  : 32
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7450

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28644 pages
