# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t97_mcart_1, conjecture, ![X1, X2, X3]:k1_relat_1(k1_relat_1(k1_tarski(k3_mcart_1(X1,X2,X3))))=k1_tarski(X1), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t97_mcart_1)).
fof(t36_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))&k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t36_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t75_enumset1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', d3_mcart_1)).
fof(t195_relat_1, axiom, ![X1, X2]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&~((k1_relat_1(k2_zfmisc_1(X1,X2))=X1&k2_relat_1(k2_zfmisc_1(X1,X2))=X2))))), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t195_relat_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t113_zfmisc_1)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', d6_enumset1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/8deba997011943f7ae16e9dd90cb5c30.E_input', t152_zfmisc_1)).
fof(c_0_14, negated_conjecture, ~(![X1, X2, X3]:k1_relat_1(k1_relat_1(k1_tarski(k3_mcart_1(X1,X2,X3))))=k1_tarski(X1)), inference(assume_negation,[status(cth)],[t97_mcart_1])).
fof(c_0_15, plain, ![X30, X31, X32]:(k2_zfmisc_1(k1_tarski(X30),k2_tarski(X31,X32))=k2_tarski(k4_tarski(X30,X31),k4_tarski(X30,X32))&k2_zfmisc_1(k2_tarski(X30,X31),k1_tarski(X32))=k2_tarski(k4_tarski(X30,X32),k4_tarski(X31,X32))), inference(variable_rename,[status(thm)],[t36_zfmisc_1])).
fof(c_0_16, plain, ![X60]:k2_tarski(X60,X60)=k1_tarski(X60), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_17, plain, ![X58, X59]:k1_enumset1(X58,X58,X59)=k2_tarski(X58,X59), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_18, plain, ![X21, X22, X23]:k2_enumset1(X21,X21,X22,X23)=k1_enumset1(X21,X22,X23), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_19, plain, ![X54, X55, X56, X57]:k3_enumset1(X54,X54,X55,X56,X57)=k2_enumset1(X54,X55,X56,X57), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_20, plain, ![X61, X62, X63, X64, X65]:k4_enumset1(X61,X61,X62,X63,X64,X65)=k3_enumset1(X61,X62,X63,X64,X65), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_21, plain, ![X24, X25, X26, X27, X28, X29]:k5_enumset1(X24,X24,X25,X26,X27,X28,X29)=k4_enumset1(X24,X25,X26,X27,X28,X29), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_22, plain, ![X68, X69, X70, X71, X72, X73, X74]:k6_enumset1(X68,X68,X69,X70,X71,X72,X73,X74)=k5_enumset1(X68,X69,X70,X71,X72,X73,X74), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_23, negated_conjecture, k1_relat_1(k1_relat_1(k1_tarski(k3_mcart_1(esk1_0,esk2_0,esk3_0))))!=k1_tarski(esk1_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_14])])])).
fof(c_0_24, plain, ![X18, X19, X20]:k3_mcart_1(X18,X19,X20)=k4_tarski(k4_tarski(X18,X19),X20), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_25, plain, ![X66, X67]:((k1_relat_1(k2_zfmisc_1(X66,X67))=X66|X67=k1_xboole_0|X66=k1_xboole_0)&(k2_relat_1(k2_zfmisc_1(X66,X67))=X67|X67=k1_xboole_0|X66=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t195_relat_1])])])])).
cnf(c_0_26, plain, (k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_33, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, negated_conjecture, (k1_relat_1(k1_relat_1(k1_tarski(k3_mcart_1(esk1_0,esk2_0,esk3_0))))!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (k1_relat_1(k2_zfmisc_1(X1,X2))=X1|X2=k1_xboole_0|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_37, plain, (k2_zfmisc_1(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2),k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X3))=k6_enumset1(k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_28]), c_0_28]), c_0_29]), c_0_29]), c_0_29]), c_0_30]), c_0_30]), c_0_30]), c_0_31]), c_0_31]), c_0_31]), c_0_32]), c_0_32]), c_0_32]), c_0_33]), c_0_33]), c_0_33])).
fof(c_0_38, plain, ![X14, X15]:((k2_zfmisc_1(X14,X15)!=k1_xboole_0|(X14=k1_xboole_0|X15=k1_xboole_0))&((X14!=k1_xboole_0|k2_zfmisc_1(X14,X15)=k1_xboole_0)&(X15!=k1_xboole_0|k2_zfmisc_1(X14,X15)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_39, plain, (k2_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_40, plain, ![X33, X34, X35, X36, X37, X38, X39, X40, X41, X42, X43, X44, X45, X46, X47, X48, X49, X50, X51, X52]:(((~r2_hidden(X42,X41)|(X42=X33|(X42=X34|(X42=X35|(X42=X36|(X42=X37|(X42=X38|(X42=X39|X42=X40)))))))|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))&((X43!=X33|r2_hidden(X43,X41)|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))&((X43!=X34|r2_hidden(X43,X41)|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))&((X43!=X35|r2_hidden(X43,X41)|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))&((X43!=X36|r2_hidden(X43,X41)|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))&((X43!=X37|r2_hidden(X43,X41)|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))&((X43!=X38|r2_hidden(X43,X41)|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))&((X43!=X39|r2_hidden(X43,X41)|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))&(X43!=X40|r2_hidden(X43,X41)|X41!=k6_enumset1(X33,X34,X35,X36,X37,X38,X39,X40))))))))))&(((esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)!=X44|~r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51))&((esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)!=X45|~r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51))&((esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)!=X46|~r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51))&((esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)!=X47|~r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51))&((esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)!=X48|~r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51))&((esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)!=X49|~r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51))&((esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)!=X50|~r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51))&(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)!=X51|~r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51)))))))))&(r2_hidden(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52),X52)|(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)=X44|(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)=X45|(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)=X46|(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)=X47|(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)=X48|(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)=X49|(esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)=X50|esk4_9(X44,X45,X46,X47,X48,X49,X50,X51,X52)=X51)))))))|X52=k6_enumset1(X44,X45,X46,X47,X48,X49,X50,X51)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
cnf(c_0_41, negated_conjecture, (k1_relat_1(k1_relat_1(k6_enumset1(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0))))!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_27]), c_0_27]), c_0_28]), c_0_28]), c_0_29]), c_0_29]), c_0_35]), c_0_35]), c_0_35]), c_0_35]), c_0_30]), c_0_30]), c_0_31]), c_0_31]), c_0_32]), c_0_32]), c_0_33]), c_0_33])).
cnf(c_0_42, plain, (k1_relat_1(k6_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X3,X2)))=k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X3)|k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X3)=k1_xboole_0|k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44, plain, (k2_zfmisc_1(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3))=k6_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_27]), c_0_28]), c_0_28]), c_0_28]), c_0_29]), c_0_29]), c_0_29]), c_0_30]), c_0_30]), c_0_30]), c_0_31]), c_0_31]), c_0_31]), c_0_32]), c_0_32]), c_0_32]), c_0_33]), c_0_33]), c_0_33])).
fof(c_0_45, plain, ![X16, X17]:~r2_hidden(X16,k2_zfmisc_1(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_46, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_47, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X2,X4,X5,X6,X7,X8,X9,X10)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, negated_conjecture, (k6_enumset1(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0))=k1_xboole_0|k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0|k1_relat_1(k6_enumset1(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0)))!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49, plain, (k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)=k1_xboole_0|k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)=k1_xboole_0|k6_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X3))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (r2_hidden(X1,k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_47])])).
cnf(c_0_53, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k1_xboole_0|k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0|k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_42]), c_0_49])).
cnf(c_0_54, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_55, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0|k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_56, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_55]), c_0_54])).
cnf(c_0_57, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_56]), c_0_54]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 29
# Proof object formula steps           : 29
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 63
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 8
# Initial clauses in saturation        : 27
# Processed clauses                    : 78
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 75
# Other redundant clauses eliminated   : 19
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 3
# Generated clauses                    : 78
# ...of the previous two non-trivial   : 57
# Contextual simplify-reflections      : 1
# Paramodulations                      : 67
# Factorizations                       : 0
# Equation resolutions                 : 19
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 33
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 31
# ...number of literals in the above   : 86
# Current number of archived formulas  : 0
# Current number of archived clauses   : 39
# Clause-clause subsumption calls (NU) : 132
# Rec. Clause-clause subsumption calls : 122
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 59
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3255

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28672 pages
