# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(m1_pre_topc(X2,X3)=>((r1_tsep_1(X2,X4)&r1_tsep_1(X4,X2))|(~(r1_tsep_1(X3,X4))&~(r1_tsep_1(X4,X3))))))))), file('data/MPTP/proofs/8ec45df2a4df4de58a768f367443b87b.E_input', t23_tmap_1)).
fof(dt_k2_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k2_tsep_1(X1,X2,X3)))&(v1_pre_topc(k2_tsep_1(X1,X2,X3))&m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)))), file('data/MPTP/proofs/8ec45df2a4df4de58a768f367443b87b.E_input', dt_k2_tsep_1)).
fof(t41_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(~(r1_tsep_1(X2,X3))=>((~(r1_tsep_1(k2_tsep_1(X1,X2,X3),X4))=>(~(r1_tsep_1(X2,X4))&~(r1_tsep_1(X3,X4))))&(~(r1_tsep_1(X4,k2_tsep_1(X1,X2,X3)))=>(~(r1_tsep_1(X4,X2))&~(r1_tsep_1(X4,X3)))))))))), file('data/MPTP/proofs/8ec45df2a4df4de58a768f367443b87b.E_input', t41_tmap_1)).
fof(t30_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>(m1_pre_topc(k2_tsep_1(X1,X2,X3),X2)&m1_pre_topc(k2_tsep_1(X1,X2,X3),X3)))))), file('data/MPTP/proofs/8ec45df2a4df4de58a768f367443b87b.E_input', t30_tsep_1)).
fof(c_0_4, plain, ![X9, X10, X11, X12]:(((~r1_tsep_1(X11,X12)|r1_tsep_1(X10,X12)|~m1_pre_topc(X10,X11)|(v2_struct_0(X12)|~m1_pre_topc(X12,X9))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9))|(v2_struct_0(X10)|~m1_pre_topc(X10,X9))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&(~r1_tsep_1(X12,X11)|r1_tsep_1(X10,X12)|~m1_pre_topc(X10,X11)|(v2_struct_0(X12)|~m1_pre_topc(X12,X9))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9))|(v2_struct_0(X10)|~m1_pre_topc(X10,X9))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))))&((~r1_tsep_1(X11,X12)|r1_tsep_1(X12,X10)|~m1_pre_topc(X10,X11)|(v2_struct_0(X12)|~m1_pre_topc(X12,X9))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9))|(v2_struct_0(X10)|~m1_pre_topc(X10,X9))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&(~r1_tsep_1(X12,X11)|r1_tsep_1(X12,X10)|~m1_pre_topc(X10,X11)|(v2_struct_0(X12)|~m1_pre_topc(X12,X9))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9))|(v2_struct_0(X10)|~m1_pre_topc(X10,X9))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t23_tmap_1])])])])])).
fof(c_0_5, plain, ![X13, X14, X15]:((~v2_struct_0(k2_tsep_1(X13,X14,X15))|(v2_struct_0(X13)|~l1_pre_topc(X13)|(v2_struct_0(X14)|~m1_pre_topc(X14,X13)|(v2_struct_0(X15)|~m1_pre_topc(X15,X13)))))&((v1_pre_topc(k2_tsep_1(X13,X14,X15))|(v2_struct_0(X13)|~l1_pre_topc(X13)|(v2_struct_0(X14)|~m1_pre_topc(X14,X13)|(v2_struct_0(X15)|~m1_pre_topc(X15,X13)))))&(m1_pre_topc(k2_tsep_1(X13,X14,X15),X13)|(v2_struct_0(X13)|~l1_pre_topc(X13)|(v2_struct_0(X14)|~m1_pre_topc(X14,X13)|(v2_struct_0(X15)|~m1_pre_topc(X15,X13))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_tsep_1])])])])).
fof(c_0_6, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(~(r1_tsep_1(X2,X3))=>((~(r1_tsep_1(k2_tsep_1(X1,X2,X3),X4))=>(~(r1_tsep_1(X2,X4))&~(r1_tsep_1(X3,X4))))&(~(r1_tsep_1(X4,k2_tsep_1(X1,X2,X3)))=>(~(r1_tsep_1(X4,X2))&~(r1_tsep_1(X4,X3))))))))))), inference(assume_negation,[status(cth)],[t41_tmap_1])).
cnf(c_0_7, plain, (r1_tsep_1(X3,X2)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_tsep_1(X1,X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~m1_pre_topc(X3,X4)|~v2_pre_topc(X4)|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(k2_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_10, plain, ![X16, X17, X18]:((m1_pre_topc(k2_tsep_1(X16,X17,X18),X17)|r1_tsep_1(X17,X18)|(v2_struct_0(X18)|~m1_pre_topc(X18,X16))|(v2_struct_0(X17)|~m1_pre_topc(X17,X16))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&(m1_pre_topc(k2_tsep_1(X16,X17,X18),X18)|r1_tsep_1(X17,X18)|(v2_struct_0(X18)|~m1_pre_topc(X18,X16))|(v2_struct_0(X17)|~m1_pre_topc(X17,X16))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_tsep_1])])])])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&(~r1_tsep_1(esk2_0,esk3_0)&(((~r1_tsep_1(esk4_0,k2_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0))&(r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)))&((~r1_tsep_1(esk4_0,k2_tsep_1(esk1_0,esk2_0,esk3_0))|(r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(esk3_0,esk4_0)))&(r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,esk3_0)|(r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(esk3_0,esk4_0)))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])).
cnf(c_0_12, plain, (r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X1)|~r1_tsep_1(X5,X4)|~m1_pre_topc(k2_tsep_1(X1,X2,X3),X5)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])).
cnf(c_0_13, plain, (m1_pre_topc(k2_tsep_1(X1,X2,X3),X2)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (r1_tsep_1(X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_tsep_1(X1,X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X1,X4)|~m1_pre_topc(X3,X4)|~v2_pre_topc(X4)|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15, plain, (m1_pre_topc(k2_tsep_1(X1,X2,X3),X3)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (~r1_tsep_1(esk4_0,k2_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_tsep_1(X2,X4)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_23, negated_conjecture, (~r1_tsep_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|v2_struct_0(X1)|v2_struct_0(X2)|~r1_tsep_1(X5,X1)|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X5)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X5,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_8]), c_0_9])).
cnf(c_0_29, plain, (r1_tsep_1(X1,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r1_tsep_1(X1,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X4)|~m1_pre_topc(X2,X4)|~m1_pre_topc(X3,X4)|~v2_pre_topc(X4)|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_30, plain, (r1_tsep_1(k2_tsep_1(X1,X2,X3),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|~r1_tsep_1(X3,X4)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_12, c_0_15])).
cnf(c_0_31, negated_conjecture, (~r1_tsep_1(esk4_0,k2_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(esk2_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_32, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_tsep_1(X3,X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X4,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_28, c_0_13])).
cnf(c_0_33, negated_conjecture, (r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_34, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X2)|~r1_tsep_1(X1,X5)|~m1_pre_topc(k2_tsep_1(X2,X3,X4),X5)|~m1_pre_topc(X5,X2)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_8]), c_0_9])).
cnf(c_0_35, negated_conjecture, (~r1_tsep_1(esk4_0,k2_tsep_1(esk1_0,esk2_0,esk3_0))|~r1_tsep_1(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_30]), c_0_18]), c_0_20]), c_0_19]), c_0_21]), c_0_22])]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_36, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_tsep_1(X4,X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_28, c_0_15])).
cnf(c_0_37, negated_conjecture, (r1_tsep_1(esk4_0,esk2_0)|r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_38, negated_conjecture, (~r1_tsep_1(esk2_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_39, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,esk2_0)|~r1_tsep_1(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_30]), c_0_18]), c_0_20]), c_0_19]), c_0_21]), c_0_22])]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_40, negated_conjecture, (r1_tsep_1(esk2_0,esk4_0)|r1_tsep_1(esk3_0,esk4_0)|~r1_tsep_1(esk4_0,k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_41, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_tsep_1(X1,X3)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_34, c_0_13])).
cnf(c_0_42, negated_conjecture, (~r1_tsep_1(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_18]), c_0_20]), c_0_19]), c_0_21]), c_0_22])]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_43, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)|r1_tsep_1(esk4_0,esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_44, negated_conjecture, (~r1_tsep_1(esk4_0,esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_19]), c_0_18]), c_0_20]), c_0_21]), c_0_22])]), c_0_38]), c_0_42]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_45, plain, (r1_tsep_1(X1,k2_tsep_1(X2,X3,X4))|r1_tsep_1(X3,X4)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X4)|v2_struct_0(X3)|~r1_tsep_1(X1,X4)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_34, c_0_15])).
cnf(c_0_46, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)), inference(sr,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_45]), c_0_46]), c_0_20]), c_0_18]), c_0_19]), c_0_21]), c_0_22])]), c_0_38]), c_0_42]), c_0_23]), c_0_24]), c_0_25]), c_0_26]), c_0_27]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 39
# Proof object formula steps           : 9
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 88
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 963
# ...of these trivial                  : 0
# ...subsumed                          : 690
# ...remaining for further processing  : 273
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 132
# Backward-rewritten                   : 1
# Generated clauses                    : 1795
# ...of the previous two non-trivial   : 1750
# Contextual simplify-reflections      : 16
# Paramodulations                      : 1793
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 115
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 96
# Current number of unprocessed clauses: 454
# ...number of literals in the above   : 5280
# Current number of archived formulas  : 0
# Current number of archived clauses   : 158
# Clause-clause subsumption calls (NU) : 8726
# Rec. Clause-clause subsumption calls : 3442
# Non-unit clause-clause subsumptions  : 799
# Unit Clause-clause subsumption calls : 107
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 50014

# -------------------------------------------------
# User time                : 0.042 s
# System time              : 0.003 s
# Total time               : 0.045 s
# Maximum resident set size: 28644 pages
