# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t149_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t149_zfmisc_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t106_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t100_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', d7_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t17_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t3_boole)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t65_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', symmetry_r1_xboole_0)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t71_enumset1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t98_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', commutativity_k2_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t3_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', commutativity_k3_xboole_0)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', d2_tarski)).
fof(l3_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2))), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', l3_zfmisc_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t16_xboole_1)).
fof(t78_xboole_1, axiom, ![X1, X2, X3]:(r1_xboole_0(X1,X2)=>k3_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(X1,X3)), file('data/MPTP/proofs/8f04ccba436c44548a4c77289857c809.E_input', t78_xboole_1)).
fof(c_0_19, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2))), inference(assume_negation,[status(cth)],[t149_zfmisc_1])).
fof(c_0_20, plain, ![X54, X55, X56]:((r1_tarski(X54,X55)|~r1_tarski(X54,k4_xboole_0(X55,X56)))&(r1_xboole_0(X54,X56)|~r1_tarski(X54,k4_xboole_0(X55,X56)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
fof(c_0_21, plain, ![X46, X47]:k4_xboole_0(X46,X47)=k5_xboole_0(X46,k3_xboole_0(X46,X47)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_22, plain, ![X9, X10]:((~r1_xboole_0(X9,X10)|k3_xboole_0(X9,X10)=k1_xboole_0)&(k3_xboole_0(X9,X10)!=k1_xboole_0|r1_xboole_0(X9,X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_23, negated_conjecture, ((r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))&(r2_hidden(esk4_0,esk3_0)&r1_xboole_0(esk3_0,esk2_0)))&~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
cnf(c_0_24, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, negated_conjecture, (r1_xboole_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k5_xboole_0(X3,k3_xboole_0(X3,X2)))), inference(rw,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, negated_conjecture, (k3_xboole_0(esk3_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_30, plain, ![X19, X20]:r1_tarski(k3_xboole_0(X19,X20),X19), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_31, plain, ![X38]:k4_xboole_0(X38,k1_xboole_0)=X38, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_32, plain, ![X39]:r1_xboole_0(X39,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_33, plain, ![X27, X28]:(~r1_xboole_0(X27,X28)|r1_xboole_0(X28,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_34, negated_conjecture, (r1_xboole_0(X1,esk2_0)|~r1_tarski(X1,k5_xboole_0(esk3_0,k1_xboole_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, negated_conjecture, (r1_xboole_0(k3_xboole_0(k5_xboole_0(esk3_0,k1_xboole_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_36, c_0_25])).
cnf(c_0_41, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_37])).
fof(c_0_42, plain, ![X23]:k2_tarski(X23,X23)=k1_tarski(X23), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_43, plain, ![X52, X53]:k1_enumset1(X52,X52,X53)=k2_tarski(X52,X53), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_44, plain, ![X11, X12, X13]:k2_enumset1(X11,X11,X12,X13)=k1_enumset1(X11,X12,X13), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_45, plain, ![X50, X51]:k2_xboole_0(X50,X51)=k5_xboole_0(X50,k4_xboole_0(X51,X50)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_46, plain, ![X48, X49]:k2_xboole_0(X48,X49)=k2_xboole_0(X49,X48), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_47, plain, ![X40, X41, X43, X44, X45]:(((r2_hidden(esk6_2(X40,X41),X40)|r1_xboole_0(X40,X41))&(r2_hidden(esk6_2(X40,X41),X41)|r1_xboole_0(X40,X41)))&(~r2_hidden(X45,X43)|~r2_hidden(X45,X44)|~r1_xboole_0(X43,X44))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_48, negated_conjecture, (r1_xboole_0(esk2_0,k3_xboole_0(k5_xboole_0(esk3_0,k1_xboole_0),X1))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_49, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_50, plain, ![X14, X15]:k3_xboole_0(X14,X15)=k3_xboole_0(X15,X14), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_51, plain, ![X29, X30, X31, X32, X33, X34, X35, X36]:(((~r2_hidden(X32,X31)|(X32=X29|X32=X30)|X31!=k2_tarski(X29,X30))&((X33!=X29|r2_hidden(X33,X31)|X31!=k2_tarski(X29,X30))&(X33!=X30|r2_hidden(X33,X31)|X31!=k2_tarski(X29,X30))))&(((esk5_3(X34,X35,X36)!=X34|~r2_hidden(esk5_3(X34,X35,X36),X36)|X36=k2_tarski(X34,X35))&(esk5_3(X34,X35,X36)!=X35|~r2_hidden(esk5_3(X34,X35,X36),X36)|X36=k2_tarski(X34,X35)))&(r2_hidden(esk5_3(X34,X35,X36),X36)|(esk5_3(X34,X35,X36)=X34|esk5_3(X34,X35,X36)=X35)|X36=k2_tarski(X34,X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_52, plain, ![X21, X22]:((~r1_tarski(X21,k1_tarski(X22))|(X21=k1_xboole_0|X21=k1_tarski(X22)))&((X21!=k1_xboole_0|r1_tarski(X21,k1_tarski(X22)))&(X21!=k1_tarski(X22)|r1_tarski(X21,k1_tarski(X22))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_zfmisc_1])])])).
cnf(c_0_53, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_54, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_55, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_56, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_57, negated_conjecture, (~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_58, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_59, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_60, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_61, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_62, plain, ![X16, X17, X18]:k3_xboole_0(k3_xboole_0(X16,X17),X18)=k3_xboole_0(X16,k3_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_63, negated_conjecture, (r1_xboole_0(esk2_0,k3_xboole_0(esk3_0,X1))), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_64, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_65, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_66, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|~r1_tarski(X1,k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_67, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_68, negated_conjecture, (~r1_xboole_0(k5_xboole_0(esk1_0,k5_xboole_0(esk3_0,k3_xboole_0(esk3_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_58]), c_0_25])).
cnf(c_0_69, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))=k5_xboole_0(X2,k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_58]), c_0_58]), c_0_25]), c_0_25])).
fof(c_0_70, plain, ![X24, X25, X26]:(~r1_xboole_0(X24,X25)|k3_xboole_0(X24,k2_xboole_0(X25,X26))=k3_xboole_0(X24,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t78_xboole_1])])).
cnf(c_0_71, plain, (k3_xboole_0(X1,X2)!=k1_xboole_0|~r2_hidden(X3,X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_72, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_73, negated_conjecture, (r1_xboole_0(esk2_0,k3_xboole_0(X1,esk3_0))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_74, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X2,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_55]), c_0_56])).
cnf(c_0_75, plain, (X1=k1_xboole_0|X1=k2_enumset1(X2,X2,X2,X2)|~r1_tarski(X1,k2_enumset1(X2,X2,X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_54]), c_0_54]), c_0_55]), c_0_55]), c_0_56]), c_0_56])).
cnf(c_0_76, negated_conjecture, (r1_tarski(k3_xboole_0(esk2_0,esk1_0),k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[c_0_67, c_0_64])).
cnf(c_0_77, negated_conjecture, (~r1_xboole_0(k5_xboole_0(esk3_0,k5_xboole_0(esk1_0,k3_xboole_0(esk3_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69]), c_0_64])).
cnf(c_0_78, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X3))=k3_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_79, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))!=k1_xboole_0|~r2_hidden(X4,k3_xboole_0(X1,X2))|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_80, negated_conjecture, (k3_xboole_0(esk2_0,k3_xboole_0(X1,esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_73])).
cnf(c_0_81, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_74])])).
cnf(c_0_82, negated_conjecture, (k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0)=k3_xboole_0(esk2_0,esk1_0)|k3_xboole_0(esk2_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_83, negated_conjecture, (k3_xboole_0(k5_xboole_0(esk3_0,k5_xboole_0(esk1_0,k3_xboole_0(esk3_0,esk1_0))),esk2_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_77, c_0_61])).
cnf(c_0_84, plain, (k3_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X3,X2))))=k3_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_58]), c_0_25])).
cnf(c_0_85, negated_conjecture, (~r2_hidden(X1,k3_xboole_0(esk2_0,X2))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_86, negated_conjecture, (k3_xboole_0(esk2_0,esk1_0)=k1_xboole_0|r2_hidden(esk4_0,k3_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_87, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_88, negated_conjecture, (k3_xboole_0(esk2_0,k5_xboole_0(esk3_0,k5_xboole_0(esk1_0,k3_xboole_0(esk3_0,esk1_0))))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_83, c_0_64])).
cnf(c_0_89, plain, (k3_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X2,X3))))=k3_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_84, c_0_64])).
cnf(c_0_90, negated_conjecture, (k3_xboole_0(esk2_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87])])).
cnf(c_0_91, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_27])).
cnf(c_0_92, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_91])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 93
# Proof object clause steps            : 54
# Proof object formula steps           : 39
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 28
# Processed clauses                    : 886
# ...of these trivial                  : 163
# ...subsumed                          : 467
# ...remaining for further processing  : 256
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 24
# Generated clauses                    : 4866
# ...of the previous two non-trivial   : 3204
# Contextual simplify-reflections      : 0
# Paramodulations                      : 4859
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 198
#    Positive orientable unit clauses  : 91
#    Positive unorientable unit clauses: 6
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 99
# Current number of unprocessed clauses: 2312
# ...number of literals in the above   : 5882
# Current number of archived formulas  : 0
# Current number of archived clauses   : 58
# Clause-clause subsumption calls (NU) : 3038
# Rec. Clause-clause subsumption calls : 2540
# Non-unit clause-clause subsumptions  : 460
# Unit Clause-clause subsumption calls : 185
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 135
# BW rewrite match successes           : 104
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 54336

# -------------------------------------------------
# User time                : 0.039 s
# System time              : 0.000 s
# Total time               : 0.039 s
# Maximum resident set size: 28660 pages
