# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k1_wellord1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(~(X4=X2)&r2_hidden(k4_tarski(X4,X2),X1))))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', d1_wellord1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', d3_tarski)).
fof(t33_wellord1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(v2_wellord1(X3)=>r3_xboole_0(k1_wellord1(X3,X1),k1_wellord1(X3,X2)))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', t33_wellord1)).
fof(t30_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', t30_relat_1)).
fof(d9_xboole_0, axiom, ![X1, X2]:(r3_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)|r1_tarski(X2,X1))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', d9_xboole_0)).
fof(d4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)<=>(v1_relat_2(X1)&(v8_relat_2(X1)&(v4_relat_2(X1)&(v6_relat_2(X1)&v1_wellord1(X1))))))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', d4_wellord1)).
fof(l4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v6_relat_2(X1)<=>![X2, X3]:~((r2_hidden(X2,k3_relat_1(X1))&(r2_hidden(X3,k3_relat_1(X1))&(~(X2=X3)&(~(r2_hidden(k4_tarski(X2,X3),X1))&~(r2_hidden(k4_tarski(X3,X2),X1))))))))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', l4_wellord1)).
fof(l2_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v8_relat_2(X1)<=>![X2, X3, X4]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X3,X4),X1))=>r2_hidden(k4_tarski(X2,X4),X1)))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', l2_wellord1)).
fof(l3_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v4_relat_2(X1)<=>![X2, X3]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X3,X2),X1))=>X2=X3))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', l3_wellord1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/8f562d0297e44638ab225de6ccea21fc.E_input', d10_xboole_0)).
fof(c_0_10, plain, ![X5, X6, X7, X8, X9, X10, X11]:((((X8!=X6|~r2_hidden(X8,X7)|X7!=k1_wellord1(X5,X6)|~v1_relat_1(X5))&(r2_hidden(k4_tarski(X8,X6),X5)|~r2_hidden(X8,X7)|X7!=k1_wellord1(X5,X6)|~v1_relat_1(X5)))&(X9=X6|~r2_hidden(k4_tarski(X9,X6),X5)|r2_hidden(X9,X7)|X7!=k1_wellord1(X5,X6)|~v1_relat_1(X5)))&((~r2_hidden(esk1_3(X5,X10,X11),X11)|(esk1_3(X5,X10,X11)=X10|~r2_hidden(k4_tarski(esk1_3(X5,X10,X11),X10),X5))|X11=k1_wellord1(X5,X10)|~v1_relat_1(X5))&((esk1_3(X5,X10,X11)!=X10|r2_hidden(esk1_3(X5,X10,X11),X11)|X11=k1_wellord1(X5,X10)|~v1_relat_1(X5))&(r2_hidden(k4_tarski(esk1_3(X5,X10,X11),X10),X5)|r2_hidden(esk1_3(X5,X10,X11),X11)|X11=k1_wellord1(X5,X10)|~v1_relat_1(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_wellord1])])])])])])])).
cnf(c_0_11, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_12, plain, ![X13, X14, X15, X16, X17]:((~r1_tarski(X13,X14)|(~r2_hidden(X15,X13)|r2_hidden(X15,X14)))&((r2_hidden(esk2_2(X16,X17),X16)|r1_tarski(X16,X17))&(~r2_hidden(esk2_2(X16,X17),X17)|r1_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_13, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(v2_wellord1(X3)=>r3_xboole_0(k1_wellord1(X3,X1),k1_wellord1(X3,X2))))), inference(assume_negation,[status(cth)],[t33_wellord1])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,k1_wellord1(X3,X2))|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, negated_conjecture, (v1_relat_1(esk5_0)&(v2_wellord1(esk5_0)&~r3_xboole_0(k1_wellord1(esk5_0,esk3_0),k1_wellord1(esk5_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])).
fof(c_0_17, plain, ![X22, X23, X24]:((r2_hidden(X22,k3_relat_1(X24))|~r2_hidden(k4_tarski(X22,X23),X24)|~v1_relat_1(X24))&(r2_hidden(X23,k3_relat_1(X24))|~r2_hidden(k4_tarski(X22,X23),X24)|~v1_relat_1(X24))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_relat_1])])])).
cnf(c_0_18, plain, (r1_tarski(k1_wellord1(X1,X2),X3)|r2_hidden(k4_tarski(esk2_2(k1_wellord1(X1,X2),X3),X2),X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, negated_conjecture, (v1_relat_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_20, plain, ![X25, X26]:((~r3_xboole_0(X25,X26)|(r1_tarski(X25,X26)|r1_tarski(X26,X25)))&((~r1_tarski(X25,X26)|r3_xboole_0(X25,X26))&(~r1_tarski(X26,X25)|r3_xboole_0(X25,X26)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_xboole_0])])])).
cnf(c_0_21, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (r1_tarski(k1_wellord1(esk5_0,X1),X2)|r2_hidden(k4_tarski(esk2_2(k1_wellord1(esk5_0,X1),X2),X1),esk5_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, plain, (r3_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, negated_conjecture, (r1_tarski(k1_wellord1(esk5_0,X1),X2)|r2_hidden(X1,k3_relat_1(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_19])])).
fof(c_0_25, plain, ![X27]:(((v1_relat_2(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))&((v8_relat_2(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))&((v4_relat_2(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))&((v6_relat_2(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))&(v1_wellord1(X27)|~v2_wellord1(X27)|~v1_relat_1(X27))))))&(~v1_relat_2(X27)|(~v8_relat_2(X27)|(~v4_relat_2(X27)|(~v6_relat_2(X27)|~v1_wellord1(X27))))|v2_wellord1(X27)|~v1_relat_1(X27))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_wellord1])])])).
fof(c_0_26, plain, ![X28, X29, X30]:((~v6_relat_2(X28)|(~r2_hidden(X29,k3_relat_1(X28))|(~r2_hidden(X30,k3_relat_1(X28))|(X29=X30|(r2_hidden(k4_tarski(X29,X30),X28)|r2_hidden(k4_tarski(X30,X29),X28)))))|~v1_relat_1(X28))&((r2_hidden(esk6_1(X28),k3_relat_1(X28))|v6_relat_2(X28)|~v1_relat_1(X28))&((r2_hidden(esk7_1(X28),k3_relat_1(X28))|v6_relat_2(X28)|~v1_relat_1(X28))&((esk6_1(X28)!=esk7_1(X28)|v6_relat_2(X28)|~v1_relat_1(X28))&((~r2_hidden(k4_tarski(esk6_1(X28),esk7_1(X28)),X28)|v6_relat_2(X28)|~v1_relat_1(X28))&(~r2_hidden(k4_tarski(esk7_1(X28),esk6_1(X28)),X28)|v6_relat_2(X28)|~v1_relat_1(X28))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l4_wellord1])])])])])])])).
cnf(c_0_27, negated_conjecture, (~r3_xboole_0(k1_wellord1(esk5_0,esk3_0),k1_wellord1(esk5_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, negated_conjecture, (r3_xboole_0(k1_wellord1(esk5_0,X1),X2)|r2_hidden(X1,k3_relat_1(esk5_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, plain, (v6_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (v2_wellord1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_31, plain, (r3_xboole_0(X2,X1)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_32, plain, ![X33, X34, X35, X36]:((~v8_relat_2(X33)|(~r2_hidden(k4_tarski(X34,X35),X33)|~r2_hidden(k4_tarski(X35,X36),X33)|r2_hidden(k4_tarski(X34,X36),X33))|~v1_relat_1(X33))&(((r2_hidden(k4_tarski(esk8_1(X33),esk9_1(X33)),X33)|v8_relat_2(X33)|~v1_relat_1(X33))&(r2_hidden(k4_tarski(esk9_1(X33),esk10_1(X33)),X33)|v8_relat_2(X33)|~v1_relat_1(X33)))&(~r2_hidden(k4_tarski(esk8_1(X33),esk10_1(X33)),X33)|v8_relat_2(X33)|~v1_relat_1(X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l2_wellord1])])])])])).
cnf(c_0_33, plain, (v8_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (X2=X3|r2_hidden(k4_tarski(X2,X3),X1)|r2_hidden(k4_tarski(X3,X2),X1)|~v6_relat_2(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk3_0,k3_relat_1(esk5_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36, negated_conjecture, (v6_relat_2(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_19])])).
cnf(c_0_37, negated_conjecture, (r3_xboole_0(X1,k1_wellord1(esk5_0,X2))|r2_hidden(X2,k3_relat_1(esk5_0))), inference(spm,[status(thm)],[c_0_31, c_0_24])).
cnf(c_0_38, plain, (r2_hidden(k4_tarski(X2,X4),X1)|~v8_relat_2(X1)|~r2_hidden(k4_tarski(X2,X3),X1)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, negated_conjecture, (v8_relat_2(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_30]), c_0_19])])).
cnf(c_0_40, negated_conjecture, (esk3_0=X1|r2_hidden(k4_tarski(X1,esk3_0),esk5_0)|r2_hidden(k4_tarski(esk3_0,X1),esk5_0)|~r2_hidden(X1,k3_relat_1(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_19])])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk4_0,k3_relat_1(esk5_0))), inference(spm,[status(thm)],[c_0_27, c_0_37])).
cnf(c_0_42, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),esk5_0)|~r2_hidden(k4_tarski(X3,X2),esk5_0)|~r2_hidden(k4_tarski(X1,X3),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_19])])).
cnf(c_0_43, negated_conjecture, (esk4_0=esk3_0|r2_hidden(k4_tarski(esk3_0,esk4_0),esk5_0)|r2_hidden(k4_tarski(esk4_0,esk3_0),esk5_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_44, plain, (X1=X2|r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_wellord1(X3,X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_45, negated_conjecture, (esk4_0=esk3_0|r2_hidden(k4_tarski(esk3_0,esk4_0),esk5_0)|r2_hidden(k4_tarski(X1,esk3_0),esk5_0)|~r2_hidden(k4_tarski(X1,esk4_0),esk5_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_46, plain, (X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_47, negated_conjecture, (esk4_0=esk3_0|r1_tarski(k1_wellord1(esk5_0,esk4_0),X1)|r2_hidden(k4_tarski(esk2_2(k1_wellord1(esk5_0,esk4_0),X1),esk3_0),esk5_0)|r2_hidden(k4_tarski(esk3_0,esk4_0),esk5_0)), inference(spm,[status(thm)],[c_0_45, c_0_22])).
cnf(c_0_48, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_49, negated_conjecture, (esk2_2(k1_wellord1(esk5_0,esk4_0),X1)=esk3_0|esk4_0=esk3_0|r1_tarski(k1_wellord1(esk5_0,esk4_0),X1)|r2_hidden(esk2_2(k1_wellord1(esk5_0,esk4_0),X1),k1_wellord1(esk5_0,esk3_0))|r2_hidden(k4_tarski(esk3_0,esk4_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_19])])).
cnf(c_0_50, negated_conjecture, (esk2_2(k1_wellord1(esk5_0,esk4_0),k1_wellord1(esk5_0,esk3_0))=esk3_0|esk4_0=esk3_0|r1_tarski(k1_wellord1(esk5_0,esk4_0),k1_wellord1(esk5_0,esk3_0))|r2_hidden(k4_tarski(esk3_0,esk4_0),esk5_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_51, negated_conjecture, (esk4_0=esk3_0|r1_tarski(k1_wellord1(esk5_0,esk4_0),k1_wellord1(esk5_0,esk3_0))|r2_hidden(k4_tarski(esk3_0,esk4_0),esk5_0)), inference(spm,[status(thm)],[c_0_22, c_0_50])).
cnf(c_0_52, negated_conjecture, (esk4_0=esk3_0|r2_hidden(k4_tarski(esk3_0,esk4_0),esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_51]), c_0_27])).
cnf(c_0_53, negated_conjecture, (esk4_0=esk3_0|r2_hidden(k4_tarski(X1,esk4_0),esk5_0)|~r2_hidden(k4_tarski(X1,esk3_0),esk5_0)), inference(spm,[status(thm)],[c_0_42, c_0_52])).
cnf(c_0_54, negated_conjecture, (esk4_0=esk3_0|r1_tarski(k1_wellord1(esk5_0,esk3_0),X1)|r2_hidden(k4_tarski(esk2_2(k1_wellord1(esk5_0,esk3_0),X1),esk4_0),esk5_0)), inference(spm,[status(thm)],[c_0_53, c_0_22])).
cnf(c_0_55, negated_conjecture, (esk2_2(k1_wellord1(esk5_0,esk3_0),X1)=esk4_0|esk4_0=esk3_0|r1_tarski(k1_wellord1(esk5_0,esk3_0),X1)|r2_hidden(esk2_2(k1_wellord1(esk5_0,esk3_0),X1),k1_wellord1(esk5_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_54]), c_0_19])])).
cnf(c_0_56, negated_conjecture, (esk2_2(k1_wellord1(esk5_0,esk3_0),k1_wellord1(esk5_0,esk4_0))=esk4_0|esk4_0=esk3_0|r1_tarski(k1_wellord1(esk5_0,esk3_0),k1_wellord1(esk5_0,esk4_0))), inference(spm,[status(thm)],[c_0_48, c_0_55])).
fof(c_0_57, plain, ![X40, X41, X42]:((~v4_relat_2(X40)|(~r2_hidden(k4_tarski(X41,X42),X40)|~r2_hidden(k4_tarski(X42,X41),X40)|X41=X42)|~v1_relat_1(X40))&(((r2_hidden(k4_tarski(esk11_1(X40),esk12_1(X40)),X40)|v4_relat_2(X40)|~v1_relat_1(X40))&(r2_hidden(k4_tarski(esk12_1(X40),esk11_1(X40)),X40)|v4_relat_2(X40)|~v1_relat_1(X40)))&(esk11_1(X40)!=esk12_1(X40)|v4_relat_2(X40)|~v1_relat_1(X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_wellord1])])])])])).
cnf(c_0_58, plain, (v4_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_59, negated_conjecture, (esk4_0=esk3_0|r1_tarski(k1_wellord1(esk5_0,esk3_0),k1_wellord1(esk5_0,esk4_0))|r2_hidden(esk4_0,k1_wellord1(esk5_0,esk3_0))), inference(spm,[status(thm)],[c_0_15, c_0_56])).
fof(c_0_60, plain, ![X45, X46]:(((r1_tarski(X45,X46)|X45!=X46)&(r1_tarski(X46,X45)|X45!=X46))&(~r1_tarski(X45,X46)|~r1_tarski(X46,X45)|X45=X46)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_61, plain, (X2=X3|~v4_relat_2(X1)|~r2_hidden(k4_tarski(X2,X3),X1)|~r2_hidden(k4_tarski(X3,X2),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_62, negated_conjecture, (v4_relat_2(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_30]), c_0_19])])).
cnf(c_0_63, negated_conjecture, (esk4_0=esk3_0|r2_hidden(esk4_0,k1_wellord1(esk5_0,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_59]), c_0_27])).
cnf(c_0_64, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_65, negated_conjecture, (X1=X2|~r2_hidden(k4_tarski(X2,X1),esk5_0)|~r2_hidden(k4_tarski(X1,X2),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_19])])).
cnf(c_0_66, negated_conjecture, (esk4_0=esk3_0|r2_hidden(k4_tarski(esk4_0,esk3_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_63]), c_0_19])])).
cnf(c_0_67, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_68, negated_conjecture, (esk4_0=esk3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_52])).
cnf(c_0_69, plain, (r3_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_23, c_0_67])).
cnf(c_0_70, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_68]), c_0_69])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 50
# Proof object formula steps           : 21
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 40
# Processed clauses                    : 175
# ...of these trivial                  : 9
# ...subsumed                          : 35
# ...remaining for further processing  : 131
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 20
# Backward-rewritten                   : 24
# Generated clauses                    : 282
# ...of the previous two non-trivial   : 249
# Contextual simplify-reflections      : 1
# Paramodulations                      : 269
# Factorizations                       : 0
# Equation resolutions                 : 13
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 84
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 70
# Current number of unprocessed clauses: 99
# ...number of literals in the above   : 431
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 2253
# Rec. Clause-clause subsumption calls : 868
# Non-unit clause-clause subsumptions  : 56
# Unit Clause-clause subsumption calls : 122
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8418

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28756 pages
