# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/8f5d1c31769c4bd7803f66b54b28ef8d.E_input', d10_xboole_0)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/8f5d1c31769c4bd7803f66b54b28ef8d.E_input', t106_xboole_1)).
fof(t44_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_xboole_0(X2,k3_subset_1(X1,X3))<=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/8f5d1c31769c4bd7803f66b54b28ef8d.E_input', t44_subset_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/8f5d1c31769c4bd7803f66b54b28ef8d.E_input', d5_subset_1)).
fof(t86_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X1,X3))=>r1_tarski(X1,k4_xboole_0(X2,X3))), file('data/MPTP/proofs/8f5d1c31769c4bd7803f66b54b28ef8d.E_input', t86_xboole_1)).
fof(t85_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_xboole_0(X1,k4_xboole_0(X3,X2))), file('data/MPTP/proofs/8f5d1c31769c4bd7803f66b54b28ef8d.E_input', t85_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/8f5d1c31769c4bd7803f66b54b28ef8d.E_input', symmetry_r1_xboole_0)).
fof(t31_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,X3)<=>r1_tarski(k3_subset_1(X1,X3),k3_subset_1(X1,X2))))), file('data/MPTP/proofs/8f5d1c31769c4bd7803f66b54b28ef8d.E_input', t31_subset_1)).
fof(c_0_8, plain, ![X12, X13]:(((r1_tarski(X12,X13)|X12!=X13)&(r1_tarski(X13,X12)|X12!=X13))&(~r1_tarski(X12,X13)|~r1_tarski(X13,X12)|X12=X13)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_9, plain, ![X19, X20, X21]:((r1_tarski(X19,X20)|~r1_tarski(X19,k4_xboole_0(X20,X21)))&(r1_xboole_0(X19,X21)|~r1_tarski(X19,k4_xboole_0(X20,X21)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
cnf(c_0_10, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_xboole_0(X2,k3_subset_1(X1,X3))<=>r1_tarski(X2,X3))))), inference(assume_negation,[status(cth)],[t44_subset_1])).
cnf(c_0_12, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X7, X8]:(~m1_subset_1(X8,k1_zfmisc_1(X7))|k3_subset_1(X7,X8)=k4_xboole_0(X7,X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
fof(c_0_15, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((~r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|~r1_tarski(esk2_0,esk3_0))&(r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|r1_tarski(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_16, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_18, plain, ![X14, X15, X16]:(~r1_tarski(X14,X15)|~r1_xboole_0(X14,X16)|r1_tarski(X14,k4_xboole_0(X15,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_xboole_1])])).
fof(c_0_19, plain, ![X9, X10, X11]:(~r1_tarski(X9,X10)|r1_xboole_0(X9,k4_xboole_0(X11,X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t85_xboole_1])])).
cnf(c_0_20, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (k4_xboole_0(X1,X2)=X1|~r1_tarski(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23, plain, (r1_tarski(X1,k4_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_24, plain, ![X17, X18]:(~r1_xboole_0(X17,X18)|r1_xboole_0(X18,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_25, plain, (r1_xboole_0(X1,k4_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (k4_xboole_0(esk1_0,esk3_0)=k3_subset_1(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_13])])).
fof(c_0_29, plain, ![X22, X23, X24]:((~r1_tarski(X23,X24)|r1_tarski(k3_subset_1(X22,X24),k3_subset_1(X22,X23))|~m1_subset_1(X24,k1_zfmisc_1(X22))|~m1_subset_1(X23,k1_zfmisc_1(X22)))&(~r1_tarski(k3_subset_1(X22,X24),k3_subset_1(X22,X23))|r1_tarski(X23,X24)|~m1_subset_1(X24,k1_zfmisc_1(X22))|~m1_subset_1(X23,k1_zfmisc_1(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_subset_1])])])])).
cnf(c_0_30, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_32, negated_conjecture, (~r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|~r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_33, negated_conjecture, (r1_xboole_0(X1,k3_subset_1(esk1_0,esk3_0))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (r1_tarski(k3_subset_1(X3,X2),k3_subset_1(X3,X1))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_xboole_0(k3_subset_1(esk1_0,esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_30, c_0_25])).
cnf(c_0_39, plain, (r1_xboole_0(k3_subset_1(X1,X2),X3)|~r1_tarski(X4,X2)|~r1_xboole_0(k3_subset_1(X1,X4),X3)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, negated_conjecture, (r1_xboole_0(k3_subset_1(esk1_0,esk3_0),esk2_0)), inference(sr,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X2,X3)|~r1_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_38, c_0_28])).
cnf(c_0_42, negated_conjecture, (r1_tarski(X1,k3_subset_1(esk1_0,esk3_0))|~r1_tarski(X1,esk1_0)|~r1_xboole_0(X1,esk3_0)), inference(spm,[status(thm)],[c_0_23, c_0_26])).
cnf(c_0_43, negated_conjecture, (r1_xboole_0(k3_subset_1(esk1_0,X1),esk2_0)|~r1_tarski(esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_21])])).
cnf(c_0_44, negated_conjecture, (r1_xboole_0(X1,X2)|~r1_tarski(X2,esk1_0)|~r1_xboole_0(X1,k3_subset_1(esk1_0,esk3_0))|~r1_xboole_0(X2,esk3_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_45, negated_conjecture, (r1_xboole_0(esk2_0,k3_subset_1(esk1_0,X1))|~r1_tarski(esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_30, c_0_43])).
cnf(c_0_46, negated_conjecture, (r1_xboole_0(X1,esk3_0)|~r1_tarski(X1,k3_subset_1(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_27, c_0_26])).
cnf(c_0_47, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,k3_subset_1(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_12, c_0_26])).
cnf(c_0_48, negated_conjecture, (r1_xboole_0(esk2_0,X1)|~r1_tarski(X1,esk1_0)|~r1_xboole_0(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_13]), c_0_21])])).
cnf(c_0_49, negated_conjecture, (r1_xboole_0(k4_xboole_0(k3_subset_1(esk1_0,esk3_0),X1),esk3_0)), inference(spm,[status(thm)],[c_0_46, c_0_17])).
cnf(c_0_50, negated_conjecture, (r1_tarski(k4_xboole_0(k3_subset_1(esk1_0,esk3_0),X1),esk1_0)), inference(spm,[status(thm)],[c_0_47, c_0_17])).
cnf(c_0_51, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(k4_xboole_0(X3,X4),X2)|~r1_xboole_0(X1,X4)), inference(spm,[status(thm)],[c_0_34, c_0_23])).
cnf(c_0_52, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_41, c_0_17])).
cnf(c_0_53, negated_conjecture, (r1_xboole_0(esk2_0,k4_xboole_0(k3_subset_1(esk1_0,esk3_0),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_54, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_tarski(X1,X4)|~r1_xboole_0(k4_xboole_0(X4,X5),X2)|~r1_xboole_0(X1,X5)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_55, negated_conjecture, (r1_xboole_0(k4_xboole_0(k3_subset_1(esk1_0,esk3_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_53])).
cnf(c_0_56, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_27, c_0_13])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_58, negated_conjecture, (r1_xboole_0(X1,k4_xboole_0(esk2_0,X2))|~r1_tarski(X1,k3_subset_1(esk1_0,esk3_0))|~r1_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59, negated_conjecture, (r1_xboole_0(k3_subset_1(esk1_0,esk3_0),esk3_0)), inference(spm,[status(thm)],[c_0_56, c_0_26])).
cnf(c_0_60, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=k3_subset_1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_20, c_0_57])).
cnf(c_0_61, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X2,X3)|~r1_xboole_0(X1,k4_xboole_0(X3,X4))|~r1_xboole_0(X2,X4)), inference(spm,[status(thm)],[c_0_41, c_0_23])).
cnf(c_0_62, negated_conjecture, (r1_xboole_0(k3_subset_1(esk1_0,esk3_0),k4_xboole_0(esk2_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_13])])).
cnf(c_0_63, plain, (r1_tarski(X3,X2)|~r1_tarski(k3_subset_1(X1,X2),k3_subset_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_64, negated_conjecture, (r1_tarski(X1,k3_subset_1(esk1_0,esk2_0))|~r1_tarski(X1,esk1_0)|~r1_xboole_0(X1,esk2_0)), inference(spm,[status(thm)],[c_0_23, c_0_60])).
cnf(c_0_65, negated_conjecture, (r1_xboole_0(k3_subset_1(esk1_0,esk3_0),X1)|~r1_tarski(X1,esk2_0)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_66, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_30, c_0_56])).
cnf(c_0_67, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k3_subset_1(esk1_0,X1),esk1_0)|~r1_xboole_0(k3_subset_1(esk1_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_57])])).
cnf(c_0_68, negated_conjecture, (r1_xboole_0(k3_subset_1(esk1_0,esk3_0),X1)|~r1_tarski(X1,esk2_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_69, negated_conjecture, (r1_tarski(k3_subset_1(esk1_0,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_17, c_0_26])).
cnf(c_0_70, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_21]), c_0_13])]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 54
# Proof object formula steps           : 17
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 809
# ...of these trivial                  : 17
# ...subsumed                          : 475
# ...remaining for further processing  : 317
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 51
# Backward-rewritten                   : 1
# Generated clauses                    : 3471
# ...of the previous two non-trivial   : 3058
# Contextual simplify-reflections      : 8
# Paramodulations                      : 3467
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 247
#    Positive orientable unit clauses  : 35
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 211
# Current number of unprocessed clauses: 1952
# ...number of literals in the above   : 7463
# Current number of archived formulas  : 0
# Current number of archived clauses   : 68
# Clause-clause subsumption calls (NU) : 11949
# Rec. Clause-clause subsumption calls : 7607
# Non-unit clause-clause subsumptions  : 526
# Unit Clause-clause subsumption calls : 20
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 56636

# -------------------------------------------------
# User time                : 0.043 s
# System time              : 0.000 s
# Total time               : 0.043 s
# Maximum resident set size: 28848 pages
