# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t26_tops_2, conjecture, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)=>v3_pre_topc(k5_setfam_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/8fe5b0668e754849b102d080b02bcb63.E_input', t26_tops_2)).
fof(d1_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,X2)=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/8fe5b0668e754849b102d080b02bcb63.E_input', d1_tops_2)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/8fe5b0668e754849b102d080b02bcb63.E_input', d1_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/8fe5b0668e754849b102d080b02bcb63.E_input', d3_tarski)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/8fe5b0668e754849b102d080b02bcb63.E_input', t3_subset)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/8fe5b0668e754849b102d080b02bcb63.E_input', d5_pre_topc)).
fof(dt_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>m1_subset_1(k5_setfam_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/8fe5b0668e754849b102d080b02bcb63.E_input', dt_k5_setfam_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/8fe5b0668e754849b102d080b02bcb63.E_input', d1_pre_topc)).
fof(c_0_8, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)=>v3_pre_topc(k5_setfam_1(u1_struct_0(X1),X2),X1))))), inference(assume_negation,[status(cth)],[t26_tops_2])).
fof(c_0_9, plain, ![X21, X22, X23]:((~v1_tops_2(X22,X21)|(~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X21)))|(~r2_hidden(X23,X22)|v3_pre_topc(X23,X21)))|~m1_subset_1(X22,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X21))))|~l1_pre_topc(X21))&((m1_subset_1(esk5_2(X21,X22),k1_zfmisc_1(u1_struct_0(X21)))|v1_tops_2(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X21))))|~l1_pre_topc(X21))&((r2_hidden(esk5_2(X21,X22),X22)|v1_tops_2(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X21))))|~l1_pre_topc(X21))&(~v3_pre_topc(esk5_2(X21,X22),X21)|v1_tops_2(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X21))))|~l1_pre_topc(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tops_2])])])])])).
fof(c_0_10, negated_conjecture, ((v2_pre_topc(esk3_0)&l1_pre_topc(esk3_0))&(m1_subset_1(esk4_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk3_0))))&(v1_tops_2(esk4_0,esk3_0)&~v3_pre_topc(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_11, plain, ![X4, X5, X6, X7, X8, X9]:(((~r2_hidden(X6,X5)|r1_tarski(X6,X4)|X5!=k1_zfmisc_1(X4))&(~r1_tarski(X7,X4)|r2_hidden(X7,X5)|X5!=k1_zfmisc_1(X4)))&((~r2_hidden(esk1_2(X8,X9),X9)|~r1_tarski(esk1_2(X8,X9),X8)|X9=k1_zfmisc_1(X8))&(r2_hidden(esk1_2(X8,X9),X9)|r1_tarski(esk1_2(X8,X9),X8)|X9=k1_zfmisc_1(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_12, plain, ![X11, X12, X13, X14, X15]:((~r1_tarski(X11,X12)|(~r2_hidden(X13,X11)|r2_hidden(X13,X12)))&((r2_hidden(esk2_2(X14,X15),X14)|r1_tarski(X14,X15))&(~r2_hidden(esk2_2(X14,X15),X15)|r1_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_13, plain, ![X17, X18]:((~m1_subset_1(X17,k1_zfmisc_1(X18))|r1_tarski(X17,X18))&(~r1_tarski(X17,X18)|m1_subset_1(X17,k1_zfmisc_1(X18)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_14, plain, ![X25, X26]:((~v3_pre_topc(X26,X25)|r2_hidden(X26,u1_pre_topc(X25))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_pre_topc(X25))&(~r2_hidden(X26,u1_pre_topc(X25))|v3_pre_topc(X26,X25)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|~l1_pre_topc(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
cnf(c_0_15, plain, (v3_pre_topc(X3,X2)|~v1_tops_2(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, negated_conjecture, (v1_tops_2(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, negated_conjecture, (l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk3_0))))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (r2_hidden(X1,u1_pre_topc(X2))|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, negated_conjecture, (v3_pre_topc(X1,esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk3_0)))|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_24, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, negated_conjecture, (~v3_pre_topc(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_28, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_29, plain, ![X27, X28]:(~m1_subset_1(X28,k1_zfmisc_1(k1_zfmisc_1(X27)))|m1_subset_1(k5_setfam_1(X27,X28),k1_zfmisc_1(X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_setfam_1])])).
cnf(c_0_30, negated_conjecture, (r2_hidden(X1,u1_pre_topc(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk3_0)))|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_17])])).
cnf(c_0_31, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk3_0)))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_26, c_0_18])).
cnf(c_0_33, negated_conjecture, (~m1_subset_1(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),k1_zfmisc_1(u1_struct_0(esk3_0)))|~r2_hidden(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),u1_pre_topc(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_17])])).
cnf(c_0_34, plain, (m1_subset_1(k5_setfam_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_35, plain, ![X29, X30, X31, X32]:(((r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|~v2_pre_topc(X29)|~l1_pre_topc(X29))&((~m1_subset_1(X30,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|(~r1_tarski(X30,u1_pre_topc(X29))|r2_hidden(k5_setfam_1(u1_struct_0(X29),X30),u1_pre_topc(X29)))|~v2_pre_topc(X29)|~l1_pre_topc(X29))&(~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29)))|(~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X29)))|(~r2_hidden(X31,u1_pre_topc(X29))|~r2_hidden(X32,u1_pre_topc(X29))|r2_hidden(k9_subset_1(u1_struct_0(X29),X31,X32),u1_pre_topc(X29))))|~v2_pre_topc(X29)|~l1_pre_topc(X29))))&(((m1_subset_1(esk7_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|m1_subset_1(esk6_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&((m1_subset_1(esk8_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|m1_subset_1(esk6_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(((r2_hidden(esk7_1(X29),u1_pre_topc(X29))|m1_subset_1(esk6_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(r2_hidden(esk8_1(X29),u1_pre_topc(X29))|m1_subset_1(esk6_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))&(~r2_hidden(k9_subset_1(u1_struct_0(X29),esk7_1(X29),esk8_1(X29)),u1_pre_topc(X29))|m1_subset_1(esk6_1(X29),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X29))))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))))&(((m1_subset_1(esk7_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|r1_tarski(esk6_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&((m1_subset_1(esk8_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|r1_tarski(esk6_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(((r2_hidden(esk7_1(X29),u1_pre_topc(X29))|r1_tarski(esk6_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(r2_hidden(esk8_1(X29),u1_pre_topc(X29))|r1_tarski(esk6_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))&(~r2_hidden(k9_subset_1(u1_struct_0(X29),esk7_1(X29),esk8_1(X29)),u1_pre_topc(X29))|r1_tarski(esk6_1(X29),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))))&((m1_subset_1(esk7_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk6_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&((m1_subset_1(esk8_1(X29),k1_zfmisc_1(u1_struct_0(X29)))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk6_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(((r2_hidden(esk7_1(X29),u1_pre_topc(X29))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk6_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29))&(r2_hidden(esk8_1(X29),u1_pre_topc(X29))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk6_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))&(~r2_hidden(k9_subset_1(u1_struct_0(X29),esk7_1(X29),esk8_1(X29)),u1_pre_topc(X29))|~r2_hidden(k5_setfam_1(u1_struct_0(X29),esk6_1(X29)),u1_pre_topc(X29))|~r2_hidden(u1_struct_0(X29),u1_pre_topc(X29))|v2_pre_topc(X29)|~l1_pre_topc(X29)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37, negated_conjecture, (r2_hidden(X1,u1_pre_topc(esk3_0))|~r2_hidden(X1,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_38, negated_conjecture, (~r2_hidden(k5_setfam_1(u1_struct_0(esk3_0),esk4_0),u1_pre_topc(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_18])])).
cnf(c_0_39, plain, (r2_hidden(k5_setfam_1(u1_struct_0(X2),X1),u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~r1_tarski(X1,u1_pre_topc(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40, negated_conjecture, (v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_41, negated_conjecture, (r1_tarski(X1,u1_pre_topc(esk3_0))|~r2_hidden(esk2_2(X1,u1_pre_topc(esk3_0)),esk4_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_43, negated_conjecture, (~r1_tarski(esk4_0,u1_pre_topc(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_17]), c_0_40]), c_0_18])])).
cnf(c_0_44, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 28
# Proof object formula steps           : 17
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 104
# ...of these trivial                  : 1
# ...subsumed                          : 5
# ...remaining for further processing  : 98
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 80
# ...of the previous two non-trivial   : 60
# Contextual simplify-reflections      : 1
# Paramodulations                      : 78
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 54
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 34
# ...number of literals in the above   : 154
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 844
# Rec. Clause-clause subsumption calls : 105
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 22
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5098

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28516 pages
