# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t32_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X4)&r2_hidden(X3,X1))=>(X2=k1_xboole_0|k1_funct_1(k2_funct_1(X4),k1_funct_1(X4,X3))=X3))), file('data/MPTP/proofs/9042375f5c6342db9f6318af7bb24436.E_input', t32_funct_2)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/9042375f5c6342db9f6318af7bb24436.E_input', d1_funct_2)).
fof(t54_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k2_funct_1(X1)<=>(k1_relat_1(X2)=k2_relat_1(X1)&![X3, X4]:(((r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))=>(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4)))&((r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))=>(r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))))))))), file('data/MPTP/proofs/9042375f5c6342db9f6318af7bb24436.E_input', t54_funct_1)).
fof(dt_k2_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k2_funct_1(X1))&v1_funct_1(k2_funct_1(X1)))), file('data/MPTP/proofs/9042375f5c6342db9f6318af7bb24436.E_input', dt_k2_funct_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/9042375f5c6342db9f6318af7bb24436.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/9042375f5c6342db9f6318af7bb24436.E_input', fc6_relat_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/9042375f5c6342db9f6318af7bb24436.E_input', redefinition_k1_relset_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((v2_funct_1(X4)&r2_hidden(X3,X1))=>(X2=k1_xboole_0|k1_funct_1(k2_funct_1(X4),k1_funct_1(X4,X3))=X3)))), inference(assume_negation,[status(cth)],[t32_funct_2])).
fof(c_0_8, plain, ![X9, X10, X11]:((((~v1_funct_2(X11,X9,X10)|X9=k1_relset_1(X9,X10,X11)|X10=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))&(X9!=k1_relset_1(X9,X10,X11)|v1_funct_2(X11,X9,X10)|X10=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))))&((~v1_funct_2(X11,X9,X10)|X9=k1_relset_1(X9,X10,X11)|X9!=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))&(X9!=k1_relset_1(X9,X10,X11)|v1_funct_2(X11,X9,X10)|X9!=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))))&((~v1_funct_2(X11,X9,X10)|X11=k1_xboole_0|X9=k1_xboole_0|X10!=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10))))&(X11!=k1_xboole_0|v1_funct_2(X11,X9,X10)|X9=k1_xboole_0|X10!=k1_xboole_0|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_9, negated_conjecture, ((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,esk2_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&((v2_funct_1(esk4_0)&r2_hidden(esk3_0,esk1_0))&(esk2_0!=k1_xboole_0&k1_funct_1(k2_funct_1(esk4_0),k1_funct_1(esk4_0,esk3_0))!=esk3_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_10, plain, ![X12, X13, X14, X15, X16, X17]:(((k1_relat_1(X13)=k2_relat_1(X12)|X13!=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12)))&(((r2_hidden(X15,k1_relat_1(X12))|(~r2_hidden(X14,k2_relat_1(X12))|X15!=k1_funct_1(X13,X14))|X13!=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12)))&(X14=k1_funct_1(X12,X15)|(~r2_hidden(X14,k2_relat_1(X12))|X15!=k1_funct_1(X13,X14))|X13!=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12))))&((r2_hidden(X16,k2_relat_1(X12))|(~r2_hidden(X17,k1_relat_1(X12))|X16!=k1_funct_1(X12,X17))|X13!=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12)))&(X17=k1_funct_1(X13,X16)|(~r2_hidden(X17,k1_relat_1(X12))|X16!=k1_funct_1(X12,X17))|X13!=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12))))))&(((((r2_hidden(esk8_2(X12,X13),k1_relat_1(X12))|r2_hidden(esk5_2(X12,X13),k2_relat_1(X12))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12)))&(esk7_2(X12,X13)=k1_funct_1(X12,esk8_2(X12,X13))|r2_hidden(esk5_2(X12,X13),k2_relat_1(X12))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12))))&(~r2_hidden(esk7_2(X12,X13),k2_relat_1(X12))|esk8_2(X12,X13)!=k1_funct_1(X13,esk7_2(X12,X13))|r2_hidden(esk5_2(X12,X13),k2_relat_1(X12))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12))))&(((r2_hidden(esk8_2(X12,X13),k1_relat_1(X12))|esk6_2(X12,X13)=k1_funct_1(X13,esk5_2(X12,X13))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12)))&(esk7_2(X12,X13)=k1_funct_1(X12,esk8_2(X12,X13))|esk6_2(X12,X13)=k1_funct_1(X13,esk5_2(X12,X13))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12))))&(~r2_hidden(esk7_2(X12,X13),k2_relat_1(X12))|esk8_2(X12,X13)!=k1_funct_1(X13,esk7_2(X12,X13))|esk6_2(X12,X13)=k1_funct_1(X13,esk5_2(X12,X13))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12)))))&(((r2_hidden(esk8_2(X12,X13),k1_relat_1(X12))|(~r2_hidden(esk6_2(X12,X13),k1_relat_1(X12))|esk5_2(X12,X13)!=k1_funct_1(X12,esk6_2(X12,X13)))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12)))&(esk7_2(X12,X13)=k1_funct_1(X12,esk8_2(X12,X13))|(~r2_hidden(esk6_2(X12,X13),k1_relat_1(X12))|esk5_2(X12,X13)!=k1_funct_1(X12,esk6_2(X12,X13)))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12))))&(~r2_hidden(esk7_2(X12,X13),k2_relat_1(X12))|esk8_2(X12,X13)!=k1_funct_1(X13,esk7_2(X12,X13))|(~r2_hidden(esk6_2(X12,X13),k1_relat_1(X12))|esk5_2(X12,X13)!=k1_funct_1(X12,esk6_2(X12,X13)))|k1_relat_1(X13)!=k2_relat_1(X12)|X13=k2_funct_1(X12)|(~v1_relat_1(X13)|~v1_funct_1(X13))|~v2_funct_1(X12)|(~v1_relat_1(X12)|~v1_funct_1(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_funct_1])])])])])])).
fof(c_0_11, plain, ![X22]:((v1_relat_1(k2_funct_1(X22))|(~v1_relat_1(X22)|~v1_funct_1(X22)))&(v1_funct_1(k2_funct_1(X22))|(~v1_relat_1(X22)|~v1_funct_1(X22)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_1])])])).
fof(c_0_12, plain, ![X23, X24]:(~v1_relat_1(X23)|(~m1_subset_1(X24,k1_zfmisc_1(X23))|v1_relat_1(X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_13, plain, ![X25, X26]:v1_relat_1(k2_zfmisc_1(X25,X26)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_14, plain, ![X27, X28, X29]:(~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X27,X28)))|k1_relset_1(X27,X28,X29)=k1_relat_1(X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_15, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, negated_conjecture, (v1_funct_2(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(X1,k1_relat_1(X4))|X3!=k1_funct_1(X4,X1)|X2!=k2_funct_1(X4)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X4)|~v1_relat_1(X4)|~v1_funct_1(X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, plain, (v1_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk4_0)=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_26, negated_conjecture, (k1_funct_1(k2_funct_1(esk4_0),k1_funct_1(esk4_0,esk3_0))!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27, plain, (k1_funct_1(k2_funct_1(X1),k1_funct_1(X1,X2))=X2|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_19])]), c_0_20]), c_0_21])).
cnf(c_0_28, negated_conjecture, (v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_16]), c_0_23])])).
cnf(c_0_29, negated_conjecture, (k1_relat_1(esk4_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_16]), c_0_25])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31, negated_conjecture, (v2_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_33, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 34
# Proof object clause steps            : 19
# Proof object formula steps           : 15
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 50
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 50
# Other redundant clauses eliminated   : 14
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 52
# ...of the previous two non-trivial   : 46
# Contextual simplify-reflections      : 12
# Paramodulations                      : 43
# Factorizations                       : 0
# Equation resolutions                 : 14
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 23
# ...number of literals in the above   : 255
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 473
# Rec. Clause-clause subsumption calls : 34
# Non-unit clause-clause subsumptions  : 12
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4505

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28596 pages
