# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1))), file('data/MPTP/proofs/908bb4aa4b1c4c2cae345ed173ba94cd.E_input', t53_yellow_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/908bb4aa4b1c4c2cae345ed173ba94cd.E_input', d4_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/908bb4aa4b1c4c2cae345ed173ba94cd.E_input', t3_subset)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/908bb4aa4b1c4c2cae345ed173ba94cd.E_input', t17_xboole_1)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/908bb4aa4b1c4c2cae345ed173ba94cd.E_input', d7_yellow_0)).
fof(d12_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v3_lattice3(X1)<=>![X2]:?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/908bb4aa4b1c4c2cae345ed173ba94cd.E_input', d12_lattice3)).
fof(t50_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)=>r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&((r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r1_yellow_0(X1,X2))&((r2_yellow_0(X1,X2)=>r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&(r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r2_yellow_0(X1,X2)))))), file('data/MPTP/proofs/908bb4aa4b1c4c2cae345ed173ba94cd.E_input', t50_yellow_0)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_yellow_0(X1,X2))=>v3_lattice3(X1)))), inference(assume_negation,[status(cth)],[t53_yellow_0])).
fof(c_0_8, plain, ![X32, X33, X34, X35, X36, X37, X38, X39]:((((r2_hidden(X35,X32)|~r2_hidden(X35,X34)|X34!=k3_xboole_0(X32,X33))&(r2_hidden(X35,X33)|~r2_hidden(X35,X34)|X34!=k3_xboole_0(X32,X33)))&(~r2_hidden(X36,X32)|~r2_hidden(X36,X33)|r2_hidden(X36,X34)|X34!=k3_xboole_0(X32,X33)))&((~r2_hidden(esk9_3(X37,X38,X39),X39)|(~r2_hidden(esk9_3(X37,X38,X39),X37)|~r2_hidden(esk9_3(X37,X38,X39),X38))|X39=k3_xboole_0(X37,X38))&((r2_hidden(esk9_3(X37,X38,X39),X37)|r2_hidden(esk9_3(X37,X38,X39),X39)|X39=k3_xboole_0(X37,X38))&(r2_hidden(esk9_3(X37,X38,X39),X38)|r2_hidden(esk9_3(X37,X38,X39),X39)|X39=k3_xboole_0(X37,X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_9, negated_conjecture, ![X7]:((~v2_struct_0(esk1_0)&l1_orders_2(esk1_0))&((~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_yellow_0(esk1_0,X7))&~v3_lattice3(esk1_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])])).
fof(c_0_10, plain, ![X19, X20]:((~m1_subset_1(X19,k1_zfmisc_1(X20))|r1_tarski(X19,X20))&(~r1_tarski(X19,X20)|m1_subset_1(X19,k1_zfmisc_1(X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_11, plain, (r2_hidden(esk9_3(X1,X2,X3),X2)|r2_hidden(esk9_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (r1_yellow_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X17, X18]:r1_tarski(k3_xboole_0(X17,X18),X17), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_16, plain, (X3=k3_xboole_0(X1,X2)|~r2_hidden(esk9_3(X1,X2,X3),X3)|~r2_hidden(esk9_3(X1,X2,X3),X1)|~r2_hidden(esk9_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, plain, (k3_xboole_0(X1,X2)=X2|r2_hidden(esk9_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_12])).
fof(c_0_19, plain, ![X21, X22, X24, X25, X27, X28, X31]:(((m1_subset_1(esk5_2(X21,X22),u1_struct_0(X21))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&((r2_lattice3(X21,X22,esk5_2(X21,X22))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&((~m1_subset_1(X24,u1_struct_0(X21))|(~r2_lattice3(X21,X22,X24)|r1_orders_2(X21,esk5_2(X21,X22),X24))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&((m1_subset_1(esk6_3(X21,X22,X25),u1_struct_0(X21))|~r2_lattice3(X21,X22,X25)|X25=esk5_2(X21,X22)|~m1_subset_1(X25,u1_struct_0(X21))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&((r2_lattice3(X21,X22,esk6_3(X21,X22,X25))|~r2_lattice3(X21,X22,X25)|X25=esk5_2(X21,X22)|~m1_subset_1(X25,u1_struct_0(X21))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21))&(~r1_orders_2(X21,X25,esk6_3(X21,X22,X25))|~r2_lattice3(X21,X22,X25)|X25=esk5_2(X21,X22)|~m1_subset_1(X25,u1_struct_0(X21))|~r1_yellow_0(X21,X22)|~l1_orders_2(X21)))))))&(((m1_subset_1(esk8_3(X21,X27,X28),u1_struct_0(X21))|m1_subset_1(esk7_3(X21,X27,X28),u1_struct_0(X21))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(((r2_lattice3(X21,X27,esk8_3(X21,X27,X28))|m1_subset_1(esk7_3(X21,X27,X28),u1_struct_0(X21))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(~m1_subset_1(X31,u1_struct_0(X21))|(~r2_lattice3(X21,X27,X31)|r1_orders_2(X21,esk8_3(X21,X27,X28),X31))|m1_subset_1(esk7_3(X21,X27,X28),u1_struct_0(X21))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21)))&(esk8_3(X21,X27,X28)!=X28|m1_subset_1(esk7_3(X21,X27,X28),u1_struct_0(X21))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))))&(((m1_subset_1(esk8_3(X21,X27,X28),u1_struct_0(X21))|r2_lattice3(X21,X27,esk7_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(((r2_lattice3(X21,X27,esk8_3(X21,X27,X28))|r2_lattice3(X21,X27,esk7_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(~m1_subset_1(X31,u1_struct_0(X21))|(~r2_lattice3(X21,X27,X31)|r1_orders_2(X21,esk8_3(X21,X27,X28),X31))|r2_lattice3(X21,X27,esk7_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21)))&(esk8_3(X21,X27,X28)!=X28|r2_lattice3(X21,X27,esk7_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))))&((m1_subset_1(esk8_3(X21,X27,X28),u1_struct_0(X21))|~r1_orders_2(X21,X28,esk7_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(((r2_lattice3(X21,X27,esk8_3(X21,X27,X28))|~r1_orders_2(X21,X28,esk7_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))&(~m1_subset_1(X31,u1_struct_0(X21))|(~r2_lattice3(X21,X27,X31)|r1_orders_2(X21,esk8_3(X21,X27,X28),X31))|~r1_orders_2(X21,X28,esk7_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21)))&(esk8_3(X21,X27,X28)!=X28|~r1_orders_2(X21,X28,esk7_3(X21,X27,X28))|~r2_lattice3(X21,X27,X28)|~m1_subset_1(X28,u1_struct_0(X21))|r1_yellow_0(X21,X27)|~l1_orders_2(X21))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
fof(c_0_20, plain, ![X8, X9, X11, X13]:(((m1_subset_1(esk2_2(X8,X9),u1_struct_0(X8))|~v3_lattice3(X8)|~l1_orders_2(X8))&((r2_lattice3(X8,X9,esk2_2(X8,X9))|~v3_lattice3(X8)|~l1_orders_2(X8))&(~m1_subset_1(X11,u1_struct_0(X8))|(~r2_lattice3(X8,X9,X11)|r1_orders_2(X8,esk2_2(X8,X9),X11))|~v3_lattice3(X8)|~l1_orders_2(X8))))&((m1_subset_1(esk4_2(X8,X13),u1_struct_0(X8))|~r2_lattice3(X8,esk3_1(X8),X13)|~m1_subset_1(X13,u1_struct_0(X8))|v3_lattice3(X8)|~l1_orders_2(X8))&((r2_lattice3(X8,esk3_1(X8),esk4_2(X8,X13))|~r2_lattice3(X8,esk3_1(X8),X13)|~m1_subset_1(X13,u1_struct_0(X8))|v3_lattice3(X8)|~l1_orders_2(X8))&(~r1_orders_2(X8,X13,esk4_2(X8,X13))|~r2_lattice3(X8,esk3_1(X8),X13)|~m1_subset_1(X13,u1_struct_0(X8))|v3_lattice3(X8)|~l1_orders_2(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_lattice3])])])])])).
cnf(c_0_21, negated_conjecture, (r1_yellow_0(esk1_0,X1)|~r1_tarski(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_22, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k3_xboole_0(X1,X2)=X2|~r2_hidden(esk9_3(X1,X2,X2),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_17])).
cnf(c_0_24, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X2,X3)|r2_hidden(esk9_3(X1,k3_xboole_0(X2,X3),k3_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_25, plain, (r1_orders_2(X2,esk5_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v3_lattice3(X1)|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_27, plain, ![X15, X16]:((~r1_yellow_0(X15,X16)|r1_yellow_0(X15,k3_xboole_0(X16,u1_struct_0(X15)))|(v2_struct_0(X15)|~l1_orders_2(X15)))&((~r1_yellow_0(X15,k3_xboole_0(X16,u1_struct_0(X15)))|r1_yellow_0(X15,X16)|(v2_struct_0(X15)|~l1_orders_2(X15)))&((~r2_yellow_0(X15,X16)|r2_yellow_0(X15,k3_xboole_0(X16,u1_struct_0(X15)))|(v2_struct_0(X15)|~l1_orders_2(X15)))&(~r2_yellow_0(X15,k3_xboole_0(X16,u1_struct_0(X15)))|r2_yellow_0(X15,X16)|(v2_struct_0(X15)|~l1_orders_2(X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_yellow_0])])])])])).
cnf(c_0_28, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(u1_struct_0(esk1_0),X1))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, plain, (r1_orders_2(X1,esk5_2(X1,X2),esk4_2(X1,X3))|v3_lattice3(X1)|~r2_lattice3(X1,X2,esk4_2(X1,X3))|~r2_lattice3(X1,esk3_1(X1),X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (r2_lattice3(X1,esk3_1(X1),esk4_2(X1,X2))|v3_lattice3(X1)|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_33, plain, (r1_yellow_0(X1,X2)|v2_struct_0(X1)|~r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (r1_yellow_0(esk1_0,k3_xboole_0(X1,u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37, plain, (v3_lattice3(X1)|~r1_orders_2(X1,X2,esk4_2(X1,X2))|~r2_lattice3(X1,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_38, plain, (r1_orders_2(X1,esk5_2(X1,X2),esk4_2(X1,esk5_2(X1,X3)))|v3_lattice3(X1)|~r2_lattice3(X1,X2,esk4_2(X1,esk5_2(X1,X3)))|~r2_lattice3(X1,esk3_1(X1),esk5_2(X1,X3))|~r1_yellow_0(X1,X2)|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, plain, (r2_lattice3(X1,esk3_1(X1),esk4_2(X1,esk5_2(X1,X2)))|v3_lattice3(X1)|~r2_lattice3(X1,esk3_1(X1),esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_32, c_0_31])).
cnf(c_0_40, negated_conjecture, (r1_yellow_0(esk1_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])]), c_0_36])).
cnf(c_0_41, negated_conjecture, (~v3_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_42, plain, (r2_lattice3(X1,X2,esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_43, plain, (v3_lattice3(X1)|~r2_lattice3(X1,X2,esk4_2(X1,esk5_2(X1,X2)))|~r2_lattice3(X1,esk3_1(X1),esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_31])).
cnf(c_0_44, negated_conjecture, (r2_lattice3(esk1_0,esk3_1(esk1_0),esk4_2(esk1_0,esk5_2(esk1_0,X1)))|~r2_lattice3(esk1_0,esk3_1(esk1_0),esk5_2(esk1_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_35])]), c_0_41])).
cnf(c_0_45, negated_conjecture, (r2_lattice3(esk1_0,X1,esk5_2(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_40]), c_0_35])])).
cnf(c_0_46, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_40]), c_0_35])]), c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 32
# Proof object formula steps           : 15
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 41
# Processed clauses                    : 5706
# ...of these trivial                  : 1808
# ...subsumed                          : 2879
# ...remaining for further processing  : 1019
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 11
# Generated clauses                    : 66676
# ...of the previous two non-trivial   : 41583
# Contextual simplify-reflections      : 19
# Paramodulations                      : 66534
# Factorizations                       : 108
# Equation resolutions                 : 34
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 961
#    Positive orientable unit clauses  : 389
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 569
# Current number of unprocessed clauses: 35918
# ...number of literals in the above   : 87116
# Current number of archived formulas  : 0
# Current number of archived clauses   : 58
# Clause-clause subsumption calls (NU) : 166613
# Rec. Clause-clause subsumption calls : 82723
# Non-unit clause-clause subsumptions  : 2831
# Unit Clause-clause subsumption calls : 28718
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14944
# BW rewrite match successes           : 49
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 878386

# -------------------------------------------------
# User time                : 0.446 s
# System time              : 0.010 s
# Total time               : 0.456 s
# Maximum resident set size: 32448 pages
