# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/90b411c8ebf542999e83b967bd9cdf3b.E_input', t23_relat_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/90b411c8ebf542999e83b967bd9cdf3b.E_input', t69_enumset1)).
fof(t35_zfmisc_1, axiom, ![X1, X2]:k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2)), file('data/MPTP/proofs/90b411c8ebf542999e83b967bd9cdf3b.E_input', t35_zfmisc_1)).
fof(t129_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k1_tarski(X4)))<=>(r2_hidden(X1,X3)&X2=X4)), file('data/MPTP/proofs/90b411c8ebf542999e83b967bd9cdf3b.E_input', t129_zfmisc_1)).
fof(t34_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),k1_tarski(X4)))<=>(X1=X3&X2=X4)), file('data/MPTP/proofs/90b411c8ebf542999e83b967bd9cdf3b.E_input', t34_zfmisc_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/90b411c8ebf542999e83b967bd9cdf3b.E_input', d1_tarski)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/90b411c8ebf542999e83b967bd9cdf3b.E_input', d5_relat_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/90b411c8ebf542999e83b967bd9cdf3b.E_input', d4_relat_1)).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t23_relat_1])).
fof(c_0_9, negated_conjecture, (v1_relat_1(esk3_0)&(esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))&(k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
fof(c_0_10, plain, ![X12]:k2_tarski(X12,X12)=k1_tarski(X12), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_11, plain, ![X24, X25]:k2_zfmisc_1(k1_tarski(X24),k1_tarski(X25))=k1_tarski(k4_tarski(X24,X25)), inference(variable_rename,[status(thm)],[t35_zfmisc_1])).
fof(c_0_12, plain, ![X8, X9, X10, X11]:(((r2_hidden(X8,X10)|~r2_hidden(k4_tarski(X8,X9),k2_zfmisc_1(X10,k1_tarski(X11))))&(X9=X11|~r2_hidden(k4_tarski(X8,X9),k2_zfmisc_1(X10,k1_tarski(X11)))))&(~r2_hidden(X8,X10)|X9!=X11|r2_hidden(k4_tarski(X8,X9),k2_zfmisc_1(X10,k1_tarski(X11))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t129_zfmisc_1])])])).
cnf(c_0_13, negated_conjecture, (esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_16, plain, ![X13, X14, X15, X16]:(((X13=X15|~r2_hidden(k4_tarski(X13,X14),k2_zfmisc_1(k1_tarski(X15),k1_tarski(X16))))&(X14=X16|~r2_hidden(k4_tarski(X13,X14),k2_zfmisc_1(k1_tarski(X15),k1_tarski(X16)))))&(X13!=X15|X14!=X16|r2_hidden(k4_tarski(X13,X14),k2_zfmisc_1(k1_tarski(X15),k1_tarski(X16))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_zfmisc_1])])])).
cnf(c_0_17, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (esk3_0=k2_tarski(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0))), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19, plain, (k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X2))=k2_zfmisc_1(k2_tarski(X1,X1),k2_tarski(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_14]), c_0_14]), c_0_14])).
fof(c_0_20, plain, ![X17, X18, X19, X20, X21, X22]:(((~r2_hidden(X19,X18)|X19=X17|X18!=k1_tarski(X17))&(X20!=X17|r2_hidden(X20,X18)|X18!=k1_tarski(X17)))&((~r2_hidden(esk4_2(X21,X22),X22)|esk4_2(X21,X22)!=X21|X22=k1_tarski(X21))&(r2_hidden(esk4_2(X21,X22),X22)|esk4_2(X21,X22)=X21|X22=k1_tarski(X21)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X4)))|X1!=X2|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k2_tarski(X2,X2)))), inference(rw,[status(thm)],[c_0_17, c_0_14])).
cnf(c_0_23, negated_conjecture, (k2_zfmisc_1(k2_tarski(esk1_0,esk1_0),k2_tarski(esk2_0,esk2_0))=esk3_0), inference(rw,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_24, plain, ![X37, X38, X39, X41, X42, X43, X44, X46]:(((~r2_hidden(X39,X38)|r2_hidden(k4_tarski(esk8_3(X37,X38,X39),X39),X37)|X38!=k2_relat_1(X37))&(~r2_hidden(k4_tarski(X42,X41),X37)|r2_hidden(X41,X38)|X38!=k2_relat_1(X37)))&((~r2_hidden(esk9_2(X43,X44),X44)|~r2_hidden(k4_tarski(X46,esk9_2(X43,X44)),X43)|X44=k2_relat_1(X43))&(r2_hidden(esk9_2(X43,X44),X44)|r2_hidden(k4_tarski(esk10_2(X43,X44),esk9_2(X43,X44)),X43)|X44=k2_relat_1(X43)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_25, negated_conjecture, (k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26, plain, (r2_hidden(esk4_2(X1,X2),X2)|esk4_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_tarski(X2,X2),k2_tarski(X4,X4)))|X3!=X4|X1!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_14]), c_0_14])).
cnf(c_0_28, negated_conjecture, (X1=esk2_0|~r2_hidden(k4_tarski(X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(esk8_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (k1_relat_1(esk3_0)!=k2_tarski(esk1_0,esk1_0)|k2_relat_1(esk3_0)!=k2_tarski(esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_14]), c_0_14])).
cnf(c_0_31, plain, (X2=k2_tarski(X1,X1)|esk4_2(X1,X2)=X1|r2_hidden(esk4_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_26, c_0_14])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k2_tarski(X1,X1),k2_tarski(X2,X2)))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_27])])).
cnf(c_0_33, negated_conjecture, (X1=esk2_0|X2!=k2_relat_1(esk3_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, negated_conjecture, (esk4_2(X1,k2_relat_1(esk3_0))=X1|r2_hidden(esk4_2(X1,k2_relat_1(esk3_0)),k2_relat_1(esk3_0))|k1_relat_1(esk3_0)!=k2_tarski(esk1_0,esk1_0)|k2_tarski(X1,X1)!=k2_tarski(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_23])).
cnf(c_0_37, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X4)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_38, plain, (X2=k1_tarski(X1)|~r2_hidden(esk4_2(X1,X2),X2)|esk4_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_39, negated_conjecture, (X1=esk2_0|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_40, negated_conjecture, (esk4_2(esk2_0,k2_relat_1(esk3_0))=esk2_0|r2_hidden(esk4_2(esk2_0,k2_relat_1(esk3_0)),k2_relat_1(esk3_0))|k1_relat_1(esk3_0)!=k2_tarski(esk1_0,esk1_0)), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk2_0,X1)|X1!=k2_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_tarski(X2,X2),k2_tarski(X4,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_14]), c_0_14])).
fof(c_0_43, plain, ![X26, X27, X28, X30, X31, X32, X33, X35]:(((~r2_hidden(X28,X27)|r2_hidden(k4_tarski(X28,esk5_3(X26,X27,X28)),X26)|X27!=k1_relat_1(X26))&(~r2_hidden(k4_tarski(X30,X31),X26)|r2_hidden(X30,X27)|X27!=k1_relat_1(X26)))&((~r2_hidden(esk6_2(X32,X33),X33)|~r2_hidden(k4_tarski(esk6_2(X32,X33),X35),X32)|X33=k1_relat_1(X32))&(r2_hidden(esk6_2(X32,X33),X33)|r2_hidden(k4_tarski(esk6_2(X32,X33),esk7_2(X32,X33)),X32)|X33=k1_relat_1(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_44, plain, (X2=k2_tarski(X1,X1)|esk4_2(X1,X2)!=X1|~r2_hidden(esk4_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_38, c_0_14])).
cnf(c_0_45, negated_conjecture, (esk4_2(esk2_0,k2_relat_1(esk3_0))=esk2_0|k1_relat_1(esk3_0)!=k2_tarski(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk2_0,k2_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_47, negated_conjecture, (X1=esk1_0|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_23])).
cnf(c_0_48, plain, (r2_hidden(k4_tarski(X1,esk5_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, negated_conjecture, (k1_relat_1(esk3_0)!=k2_tarski(esk1_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])]), c_0_30])).
cnf(c_0_50, negated_conjecture, (X1=esk1_0|X2!=k1_relat_1(esk3_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_51, negated_conjecture, (esk4_2(X1,k1_relat_1(esk3_0))=X1|r2_hidden(esk4_2(X1,k1_relat_1(esk3_0)),k1_relat_1(esk3_0))|k2_tarski(X1,X1)!=k2_tarski(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_31])).
cnf(c_0_52, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, negated_conjecture, (X1=esk1_0|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_54, negated_conjecture, (esk4_2(esk1_0,k1_relat_1(esk3_0))=esk1_0|r2_hidden(esk4_2(esk1_0,k1_relat_1(esk3_0)),k1_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_51])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk1_0,X1)|X1!=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_52, c_0_36])).
cnf(c_0_56, negated_conjecture, (esk4_2(esk1_0,k1_relat_1(esk3_0))=esk1_0), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_58, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_56]), c_0_57])]), c_0_49]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 59
# Proof object clause steps            : 42
# Proof object formula steps           : 17
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 89
# ...of these trivial                  : 1
# ...subsumed                          : 5
# ...remaining for further processing  : 83
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 270
# ...of the previous two non-trivial   : 254
# Contextual simplify-reflections      : 1
# Paramodulations                      : 246
# Factorizations                       : 0
# Equation resolutions                 : 25
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 56
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 46
# Current number of unprocessed clauses: 208
# ...number of literals in the above   : 802
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 252
# Rec. Clause-clause subsumption calls : 163
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5216

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 26200 pages
