# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t132_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>((X1=k1_tsep_1(X1,X4,X5)&r4_tsep_1(X1,X4,X5))=>((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))<=>((v1_funct_1(k2_tmap_1(X1,X2,X3,X4))&(v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X3,X4),X4,X2)&m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))&(v1_funct_1(k2_tmap_1(X1,X2,X3,X5))&(v1_funct_2(k2_tmap_1(X1,X2,X3,X5),u1_struct_0(X5),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X3,X5),X5,X2)&m1_subset_1(k2_tmap_1(X1,X2,X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))))))))))))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t132_tmap_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', dt_m1_pre_topc)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t2_tsep_1)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', cc1_pre_topc)).
fof(t25_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>k1_tsep_1(X1,X2,X2)=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t25_tmap_1)).
fof(t39_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t39_pre_topc)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t1_tsep_1)).
fof(dt_k1_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k1_tsep_1(X1,X2,X3)))&(v1_pre_topc(k1_tsep_1(X1,X2,X3))&m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', dt_k1_tsep_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', d10_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t3_subset)).
fof(t7_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X2)=>m1_pre_topc(X3,X1)))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t7_tsep_1)).
fof(t22_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>m1_pre_topc(X2,k1_tsep_1(X1,X2,X3))))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t22_tsep_1)).
fof(redefinition_k2_partfun1, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k2_partfun1(X1,X2,X3,X4)=k7_relat_1(X3,X4)), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', redefinition_k2_partfun1)).
fof(d4_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_pre_topc(X4,X1)=>k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4)))))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', d4_tmap_1)).
fof(fc2_tmap_1, axiom, ![X1, X2, X3, X4]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))&(~(v2_struct_0(X4))&m1_pre_topc(X4,X1)))))=>(v1_funct_1(k2_tmap_1(X1,X2,X3,X4))&(v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))&v5_pre_topc(k2_tmap_1(X1,X2,X3,X4),X4,X2)))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', fc2_tmap_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', fc6_relat_1)).
fof(t129_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(r4_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((v1_funct_1(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)))&(v1_funct_2(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k2_tmap_1(X1,X2,X5,X3))&(v1_funct_2(k2_tmap_1(X1,X2,X5,X3),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,X3),X3,X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k2_tmap_1(X1,X2,X5,X4))&(v1_funct_2(k2_tmap_1(X1,X2,X5,X4),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,X4),X4,X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))))))))))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t129_tmap_1)).
fof(t209_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>X2=k7_relat_1(X2,X1)), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', t209_relat_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/90d7944ad6244cbb9fc8040e5cee1d0f.E_input', cc2_relset_1)).
fof(c_0_20, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>((X1=k1_tsep_1(X1,X4,X5)&r4_tsep_1(X1,X4,X5))=>((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&(v5_pre_topc(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))))<=>((v1_funct_1(k2_tmap_1(X1,X2,X3,X4))&(v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X3,X4),X4,X2)&m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))&(v1_funct_1(k2_tmap_1(X1,X2,X3,X5))&(v1_funct_2(k2_tmap_1(X1,X2,X3,X5),u1_struct_0(X5),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X3,X5),X5,X2)&m1_subset_1(k2_tmap_1(X1,X2,X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2)))))))))))))))), inference(assume_negation,[status(cth)],[t132_tmap_1])).
fof(c_0_21, plain, ![X11, X12]:(~l1_pre_topc(X11)|(~m1_pre_topc(X12,X11)|l1_pre_topc(X12))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_22, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&((~v2_struct_0(esk5_0)&m1_pre_topc(esk5_0,esk1_0))&((esk1_0=k1_tsep_1(esk1_0,esk4_0,esk5_0)&r4_tsep_1(esk1_0,esk4_0,esk5_0))&((~v1_funct_1(esk3_0)|(~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))|(~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))|(~v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))|(~v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|(~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))))|(~v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))|(~v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|(~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0)))))))))&((((v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))|v1_funct_1(esk3_0))&((v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|v1_funct_1(esk3_0))&((v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|v1_funct_1(esk3_0))&(m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|v1_funct_1(esk3_0)))))&((v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))|v1_funct_1(esk3_0))&((v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|v1_funct_1(esk3_0))&((v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|v1_funct_1(esk3_0))&(m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))|v1_funct_1(esk3_0))))))&((((v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))|v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))&((v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))&((v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))&(m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&((v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))|v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))&((v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))&((v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))&(m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))|v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))))&((((v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0))&((v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0))&((v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|v5_pre_topc(esk3_0,esk1_0,esk2_0))&(m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|v5_pre_topc(esk3_0,esk1_0,esk2_0)))))&((v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0))&((v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0))&((v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|v5_pre_topc(esk3_0,esk1_0,esk2_0))&(m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))|v5_pre_topc(esk3_0,esk1_0,esk2_0))))))&(((v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))&((v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))&((v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))&(m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))))&((v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))&((v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))&((v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))&(m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))|m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))))))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])])).
fof(c_0_23, plain, ![X13]:(~l1_pre_topc(X13)|m1_pre_topc(X13,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
cnf(c_0_24, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, negated_conjecture, (m1_pre_topc(esk5_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, ![X14, X15]:(~v2_pre_topc(X14)|~l1_pre_topc(X14)|(~m1_pre_topc(X15,X14)|v2_pre_topc(X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
cnf(c_0_28, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X16, X17]:(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))|(v2_struct_0(X17)|~m1_pre_topc(X17,X16)|k1_tsep_1(X16,X17,X17)=g1_pre_topc(u1_struct_0(X17),u1_pre_topc(X17)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_tmap_1])])])])).
cnf(c_0_30, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, negated_conjecture, (l1_pre_topc(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
cnf(c_0_32, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, negated_conjecture, (l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_28]), c_0_26])])).
cnf(c_0_35, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k1_tsep_1(X1,X2,X2)=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (m1_pre_topc(esk5_0,esk5_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, negated_conjecture, (v2_pre_topc(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_25]), c_0_26]), c_0_33])])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_39, negated_conjecture, (m1_pre_topc(esk4_0,esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_34])).
cnf(c_0_40, negated_conjecture, (v2_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_28]), c_0_26]), c_0_33])])).
cnf(c_0_41, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_42, plain, ![X18, X19, X20]:(~l1_pre_topc(X18)|(~m1_pre_topc(X19,X18)|(~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t39_pre_topc])])])).
fof(c_0_43, plain, ![X21, X22]:(~l1_pre_topc(X21)|(~m1_pre_topc(X22,X21)|m1_subset_1(u1_struct_0(X22),k1_zfmisc_1(u1_struct_0(X21))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_44, plain, ![X23, X24, X25]:((~v2_struct_0(k1_tsep_1(X23,X24,X25))|(v2_struct_0(X23)|~l1_pre_topc(X23)|(v2_struct_0(X24)|~m1_pre_topc(X24,X23)|(v2_struct_0(X25)|~m1_pre_topc(X25,X23)))))&((v1_pre_topc(k1_tsep_1(X23,X24,X25))|(v2_struct_0(X23)|~l1_pre_topc(X23)|(v2_struct_0(X24)|~m1_pre_topc(X24,X23)|(v2_struct_0(X25)|~m1_pre_topc(X25,X23)))))&(m1_pre_topc(k1_tsep_1(X23,X24,X25),X23)|(v2_struct_0(X23)|~l1_pre_topc(X23)|(v2_struct_0(X24)|~m1_pre_topc(X24,X23)|(v2_struct_0(X25)|~m1_pre_topc(X25,X23))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tsep_1])])])])).
cnf(c_0_45, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_46, negated_conjecture, (g1_pre_topc(u1_struct_0(esk5_0),u1_pre_topc(esk5_0))=k1_tsep_1(esk5_0,esk5_0,esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_31]), c_0_37])]), c_0_38])).
cnf(c_0_47, negated_conjecture, (g1_pre_topc(u1_struct_0(esk4_0),u1_pre_topc(esk4_0))=k1_tsep_1(esk4_0,esk4_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_39]), c_0_34]), c_0_40])]), c_0_41])).
fof(c_0_48, plain, ![X26, X27]:(((r1_tarski(X26,X27)|X26!=X27)&(r1_tarski(X27,X26)|X26!=X27))&(~r1_tarski(X26,X27)|~r1_tarski(X27,X26)|X26=X27)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_49, plain, ![X28, X29]:((~m1_subset_1(X28,k1_zfmisc_1(X29))|r1_tarski(X28,X29))&(~r1_tarski(X28,X29)|m1_subset_1(X28,k1_zfmisc_1(X29)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_50, plain, (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (m1_pre_topc(k1_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_53, negated_conjecture, (k1_tsep_1(esk5_0,esk5_0,esk5_0)=k1_tsep_1(esk1_0,esk5_0,esk5_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_25]), c_0_26]), c_0_33])]), c_0_38]), c_0_45]), c_0_46])).
fof(c_0_54, plain, ![X30, X31, X32]:(~v2_pre_topc(X30)|~l1_pre_topc(X30)|(~m1_pre_topc(X31,X30)|(~m1_pre_topc(X32,X31)|m1_pre_topc(X32,X30)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tsep_1])])])).
fof(c_0_55, plain, ![X33, X34, X35]:(v2_struct_0(X33)|(~v2_pre_topc(X33)|~l1_pre_topc(X33))|(v2_struct_0(X34)|~m1_pre_topc(X34,X33)|(v2_struct_0(X35)|~m1_pre_topc(X35,X33)|m1_pre_topc(X34,k1_tsep_1(X33,X34,X35))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t22_tsep_1])])])])).
cnf(c_0_56, negated_conjecture, (k1_tsep_1(esk4_0,esk4_0,esk4_0)=k1_tsep_1(esk1_0,esk4_0,esk4_0)), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_28]), c_0_26]), c_0_33])]), c_0_41]), c_0_45]), c_0_47])).
cnf(c_0_57, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_24])).
cnf(c_0_60, negated_conjecture, (m1_pre_topc(k1_tsep_1(esk1_0,esk5_0,esk5_0),esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_36]), c_0_31])]), c_0_38])).
cnf(c_0_61, plain, (m1_pre_topc(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|m1_pre_topc(X2,k1_tsep_1(X1,X2,X3))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, negated_conjecture, (m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_56]), c_0_39]), c_0_34])]), c_0_41])).
fof(c_0_64, plain, ![X36, X37, X38, X39]:(~v1_funct_1(X38)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X36,X37)))|k2_partfun1(X36,X37,X38,X39)=k7_relat_1(X38,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_partfun1])])).
cnf(c_0_65, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_66, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk5_0)))|~m1_pre_topc(X1,k1_tsep_1(esk1_0,esk5_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_31])])).
cnf(c_0_67, negated_conjecture, (m1_pre_topc(k1_tsep_1(esk1_0,esk5_0,esk5_0),X1)|~m1_pre_topc(esk5_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_61, c_0_60])).
cnf(c_0_68, negated_conjecture, (m1_pre_topc(esk5_0,k1_tsep_1(esk1_0,esk5_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_53]), c_0_36]), c_0_31]), c_0_37])]), c_0_38])).
cnf(c_0_69, negated_conjecture, (l1_pre_topc(k1_tsep_1(esk1_0,esk5_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_60]), c_0_31])])).
cnf(c_0_70, negated_conjecture, (v2_pre_topc(k1_tsep_1(esk1_0,esk5_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_60]), c_0_31]), c_0_37])])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk4_0)))|~m1_pre_topc(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_63]), c_0_34])])).
cnf(c_0_72, negated_conjecture, (m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_61, c_0_63])).
cnf(c_0_73, negated_conjecture, (m1_pre_topc(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_56]), c_0_39]), c_0_34]), c_0_40])]), c_0_41])).
cnf(c_0_74, negated_conjecture, (l1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_63]), c_0_34])])).
cnf(c_0_75, negated_conjecture, (v2_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_63]), c_0_34]), c_0_40])])).
fof(c_0_76, plain, ![X40, X41, X42, X43]:(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))|(v2_struct_0(X41)|(~v2_pre_topc(X41)|~l1_pre_topc(X41))|(~v1_funct_1(X42)|(~v1_funct_2(X42,u1_struct_0(X40),u1_struct_0(X41))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X40),u1_struct_0(X41)))))|(~m1_pre_topc(X43,X40)|k2_tmap_1(X40,X41,X42,X43)=k2_partfun1(u1_struct_0(X40),u1_struct_0(X41),X42,u1_struct_0(X43)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tmap_1])])])])).
cnf(c_0_77, plain, (k2_partfun1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_79, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_80, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_65, c_0_58])).
cnf(c_0_81, negated_conjecture, (m1_subset_1(u1_struct_0(k1_tsep_1(esk1_0,esk5_0,esk5_0)),k1_zfmisc_1(u1_struct_0(esk5_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68]), c_0_69]), c_0_70])])).
cnf(c_0_82, negated_conjecture, (m1_subset_1(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_74]), c_0_75])])).
cnf(c_0_83, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k2_tmap_1(X1,X2,X3,X4)=k2_partfun1(u1_struct_0(X1),u1_struct_0(X2),X3,u1_struct_0(X4))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_84, negated_conjecture, (k2_partfun1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1)=k7_relat_1(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])])).
cnf(c_0_85, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_86, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_87, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_88, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_89, negated_conjecture, (u1_struct_0(k1_tsep_1(esk1_0,esk5_0,esk5_0))=u1_struct_0(esk5_0)|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk5_0,esk5_0))))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_90, negated_conjecture, (u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))=u1_struct_0(esk4_0)|~m1_subset_1(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))))), inference(spm,[status(thm)],[c_0_80, c_0_82])).
cnf(c_0_91, negated_conjecture, (k7_relat_1(esk3_0,u1_struct_0(X1))=k2_tmap_1(esk1_0,esk2_0,esk3_0,X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_78]), c_0_85]), c_0_79]), c_0_86]), c_0_26]), c_0_87]), c_0_33])]), c_0_88]), c_0_45])).
cnf(c_0_92, negated_conjecture, (u1_struct_0(k1_tsep_1(esk1_0,esk5_0,esk5_0))=u1_struct_0(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_51]), c_0_68]), c_0_69])])).
cnf(c_0_93, negated_conjecture, (u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))=u1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_51]), c_0_73]), c_0_74])])).
fof(c_0_94, plain, ![X5, X4, X3, X2, X1]:(epred1_5(X1,X2,X4,X5,X3)<=>((v1_funct_1(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)))&(v1_funct_2(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k2_tmap_1(X1,X2,X5,X3))&(v1_funct_2(k2_tmap_1(X1,X2,X5,X3),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,X3),X3,X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k2_tmap_1(X1,X2,X5,X4))&(v1_funct_2(k2_tmap_1(X1,X2,X5,X4),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,X4),X4,X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))), introduced(definition)).
fof(c_0_95, plain, ![X44, X45, X46, X47]:((v1_funct_1(k2_tmap_1(X44,X45,X46,X47))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|(~v5_pre_topc(X46,X44,X45)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(v2_struct_0(X47)|~m1_pre_topc(X47,X44))))))&((v1_funct_2(k2_tmap_1(X44,X45,X46,X47),u1_struct_0(X47),u1_struct_0(X45))|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|(~v5_pre_topc(X46,X44,X45)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(v2_struct_0(X47)|~m1_pre_topc(X47,X44))))))&(v5_pre_topc(k2_tmap_1(X44,X45,X46,X47),X47,X45)|(v2_struct_0(X44)|(~v2_pre_topc(X44)|~l1_pre_topc(X44))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))|(~v1_funct_1(X46)|(~v1_funct_2(X46,u1_struct_0(X44),u1_struct_0(X45))|(~v5_pre_topc(X46,X44,X45)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X44),u1_struct_0(X45))))))|(v2_struct_0(X47)|~m1_pre_topc(X47,X44)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_tmap_1])])])])).
cnf(c_0_96, negated_conjecture, (k7_relat_1(esk3_0,u1_struct_0(esk5_0))=k2_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk5_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk5_0,esk5_0),esk1_0)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_97, negated_conjecture, (k7_relat_1(esk3_0,u1_struct_0(esk4_0))=k2_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk1_0)), inference(spm,[status(thm)],[c_0_91, c_0_93])).
fof(c_0_98, plain, ![X48, X49]:(~v1_relat_1(X48)|(~m1_subset_1(X49,k1_zfmisc_1(X48))|v1_relat_1(X49))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_99, plain, ![X50, X51]:v1_relat_1(k2_zfmisc_1(X50,X51)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_100, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(r4_tsep_1(X1,X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>epred1_5(X1,X2,X4,X5,X3))))))), inference(apply_def,[status(thm)],[t129_tmap_1, c_0_94])).
cnf(c_0_101, plain, (v1_funct_1(k2_tmap_1(X1,X2,X3,X4))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v5_pre_topc(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_102, plain, (v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v5_pre_topc(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_103, negated_conjecture, (k7_relat_1(esk3_0,u1_struct_0(esk5_0))=k2_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk5_0,esk5_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_52]), c_0_25]), c_0_26])]), c_0_38]), c_0_45])).
cnf(c_0_104, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(k1_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_105, negated_conjecture, (k7_relat_1(esk3_0,u1_struct_0(esk4_0))=k2_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_52]), c_0_28]), c_0_26])]), c_0_41]), c_0_45])).
fof(c_0_106, plain, ![X5, X4, X3, X2, X1]:(epred1_5(X1,X2,X4,X5,X3)=>((v1_funct_1(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)))&(v1_funct_2(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,k1_tsep_1(X1,X3,X4)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))<=>((v1_funct_1(k2_tmap_1(X1,X2,X5,X3))&(v1_funct_2(k2_tmap_1(X1,X2,X5,X3),u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,X3),X3,X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))&(v1_funct_1(k2_tmap_1(X1,X2,X5,X4))&(v1_funct_2(k2_tmap_1(X1,X2,X5,X4),u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(k2_tmap_1(X1,X2,X5,X4),X4,X2)&m1_subset_1(k2_tmap_1(X1,X2,X5,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))), inference(split_equiv,[status(thm)],[c_0_94])).
fof(c_0_107, plain, ![X57, X58]:(~v1_relat_1(X58)|~v4_relat_1(X58,X57)|X58=k7_relat_1(X58,X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t209_relat_1])])).
cnf(c_0_108, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_109, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_99])).
fof(c_0_110, plain, ![X59, X60, X61]:((v4_relat_1(X61,X59)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))&(v5_relat_1(X61,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_111, plain, ![X52, X53, X54, X55, X56]:(v2_struct_0(X52)|(~v2_pre_topc(X52)|~l1_pre_topc(X52))|(v2_struct_0(X53)|(~v2_pre_topc(X53)|~l1_pre_topc(X53))|(v2_struct_0(X54)|~m1_pre_topc(X54,X52)|(v2_struct_0(X55)|~m1_pre_topc(X55,X52)|(~r4_tsep_1(X52,X54,X55)|(~v1_funct_1(X56)|(~v1_funct_2(X56,u1_struct_0(X52),u1_struct_0(X53))|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X52),u1_struct_0(X53)))))|epred1_5(X52,X53,X55,X56,X54))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_100])])])])).
cnf(c_0_112, negated_conjecture, (~v1_funct_1(esk3_0)|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))|~v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))|~v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))|~v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_113, negated_conjecture, (v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,X1))|v2_struct_0(X1)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_78]), c_0_85]), c_0_79]), c_0_86]), c_0_26]), c_0_87]), c_0_33])]), c_0_88]), c_0_45])).
cnf(c_0_114, negated_conjecture, (v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_115, negated_conjecture, (v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_116, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,X1),u1_struct_0(X1),u1_struct_0(esk2_0))|v2_struct_0(X1)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_78]), c_0_85]), c_0_79]), c_0_86]), c_0_26]), c_0_87]), c_0_33])]), c_0_88]), c_0_45])).
cnf(c_0_117, negated_conjecture, (k2_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk5_0,esk5_0))=k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_103]), c_0_25])])).
cnf(c_0_118, negated_conjecture, (~v2_struct_0(k1_tsep_1(esk1_0,esk5_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_53]), c_0_36]), c_0_31])]), c_0_38])).
cnf(c_0_119, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_120, negated_conjecture, (k2_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))=k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_105]), c_0_28])])).
cnf(c_0_121, negated_conjecture, (~v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_56]), c_0_39]), c_0_34])]), c_0_41])).
cnf(c_0_122, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_123, plain, ![X62, X63, X64, X65, X66]:((((v1_funct_1(k2_tmap_1(X66,X65,X62,X64))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65)))))))|~epred1_5(X66,X65,X63,X62,X64))&((v1_funct_2(k2_tmap_1(X66,X65,X62,X64),u1_struct_0(X64),u1_struct_0(X65))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65)))))))|~epred1_5(X66,X65,X63,X62,X64))&((v5_pre_topc(k2_tmap_1(X66,X65,X62,X64),X64,X65)|(~v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65)))))))|~epred1_5(X66,X65,X63,X62,X64))&(m1_subset_1(k2_tmap_1(X66,X65,X62,X64),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X64),u1_struct_0(X65))))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65)))))))|~epred1_5(X66,X65,X63,X62,X64)))))&((v1_funct_1(k2_tmap_1(X66,X65,X62,X63))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65)))))))|~epred1_5(X66,X65,X63,X62,X64))&((v1_funct_2(k2_tmap_1(X66,X65,X62,X63),u1_struct_0(X63),u1_struct_0(X65))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65)))))))|~epred1_5(X66,X65,X63,X62,X64))&((v5_pre_topc(k2_tmap_1(X66,X65,X62,X63),X63,X65)|(~v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65)))))))|~epred1_5(X66,X65,X63,X62,X64))&(m1_subset_1(k2_tmap_1(X66,X65,X62,X63),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X63),u1_struct_0(X65))))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65)))))))|~epred1_5(X66,X65,X63,X62,X64))))))&((v1_funct_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,X64))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,X64),u1_struct_0(X64),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,X64),X64,X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,X64),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X64),u1_struct_0(X65))))))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,X63))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,X63),u1_struct_0(X63),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,X63),X63,X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,X63),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X63),u1_struct_0(X65))))))))|~epred1_5(X66,X65,X63,X62,X64))&((v1_funct_2(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,X64))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,X64),u1_struct_0(X64),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,X64),X64,X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,X64),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X64),u1_struct_0(X65))))))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,X63))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,X63),u1_struct_0(X63),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,X63),X63,X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,X63),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X63),u1_struct_0(X65))))))))|~epred1_5(X66,X65,X63,X62,X64))&((v5_pre_topc(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_tsep_1(X66,X64,X63),X65)|(~v1_funct_1(k2_tmap_1(X66,X65,X62,X64))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,X64),u1_struct_0(X64),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,X64),X64,X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,X64),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X64),u1_struct_0(X65))))))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,X63))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,X63),u1_struct_0(X63),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,X63),X63,X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,X63),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X63),u1_struct_0(X65))))))))|~epred1_5(X66,X65,X63,X62,X64))&(m1_subset_1(k2_tmap_1(X66,X65,X62,k1_tsep_1(X66,X64,X63)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X66,X64,X63)),u1_struct_0(X65))))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,X64))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,X64),u1_struct_0(X64),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,X64),X64,X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,X64),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X64),u1_struct_0(X65))))))|(~v1_funct_1(k2_tmap_1(X66,X65,X62,X63))|(~v1_funct_2(k2_tmap_1(X66,X65,X62,X63),u1_struct_0(X63),u1_struct_0(X65))|(~v5_pre_topc(k2_tmap_1(X66,X65,X62,X63),X63,X65)|~m1_subset_1(k2_tmap_1(X66,X65,X62,X63),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X63),u1_struct_0(X65))))))))|~epred1_5(X66,X65,X63,X62,X64)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_106])])])).
cnf(c_0_124, plain, (X1=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_125, negated_conjecture, (v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_78]), c_0_109])])).
cnf(c_0_126, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_127, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|epred1_5(X1,X2,X4,X5,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~r4_tsep_1(X1,X3,X4)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_128, negated_conjecture, (~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))|~v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|~v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))|~v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_79]), c_0_85]), c_0_78])])).
cnf(c_0_129, negated_conjecture, (v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_25]), c_0_38]), c_0_114])).
cnf(c_0_130, negated_conjecture, (v1_funct_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_28]), c_0_41]), c_0_115])).
cnf(c_0_131, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk5_0,esk5_0),esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_92]), c_0_117]), c_0_118]), c_0_119])).
cnf(c_0_132, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_120]), c_0_93]), c_0_121]), c_0_122])).
cnf(c_0_133, plain, (v5_pre_topc(k2_tmap_1(X1,X2,X3,X4),X4,X2)|~v1_funct_1(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)))|~v1_funct_2(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)),u1_struct_0(k1_tsep_1(X1,X4,X5)),u1_struct_0(X2))|~v5_pre_topc(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)),k1_tsep_1(X1,X4,X5),X2)|~m1_subset_1(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X4,X5)),u1_struct_0(X2))))|~epred1_5(X1,X2,X5,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_134, negated_conjecture, (esk1_0=k1_tsep_1(esk1_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_135, negated_conjecture, (k2_tmap_1(esk1_0,esk2_0,esk3_0,X1)=esk3_0|~v4_relat_1(esk3_0,u1_struct_0(X1))|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_91]), c_0_125])])).
cnf(c_0_136, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_126, c_0_78])).
cnf(c_0_137, negated_conjecture, (m1_pre_topc(esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_26])).
cnf(c_0_138, negated_conjecture, (epred1_5(esk1_0,esk2_0,X1,esk3_0,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~r4_tsep_1(esk1_0,X2,X1)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_78]), c_0_85]), c_0_79]), c_0_86]), c_0_26]), c_0_87]), c_0_33])]), c_0_88]), c_0_45])).
cnf(c_0_139, negated_conjecture, (r4_tsep_1(esk1_0,esk4_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_140, negated_conjecture, (~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))|~v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_129])]), c_0_130])])).
cnf(c_0_141, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),u1_struct_0(esk5_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_52]), c_0_25]), c_0_26])]), c_0_38]), c_0_45])).
cnf(c_0_142, negated_conjecture, (v1_funct_2(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_52]), c_0_28]), c_0_26])]), c_0_41]), c_0_45])).
cnf(c_0_143, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk4_0),esk4_0,X1)|~epred1_5(esk1_0,X1,esk5_0,X2,esk4_0)|~v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk1_0),esk1_0,X1)|~m1_subset_1(k2_tmap_1(esk1_0,X1,X2,esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(k2_tmap_1(esk1_0,X1,X2,esk1_0),u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(k2_tmap_1(esk1_0,X1,X2,esk1_0))), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_144, negated_conjecture, (k2_tmap_1(esk1_0,esk2_0,esk3_0,esk1_0)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136]), c_0_137])])).
cnf(c_0_145, negated_conjecture, (epred1_5(esk1_0,esk2_0,esk5_0,esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_25]), c_0_28])]), c_0_41]), c_0_38])).
cnf(c_0_146, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_147, plain, (v5_pre_topc(k2_tmap_1(X1,X2,X3,X4),X4,X2)|~v1_funct_1(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X5,X4)))|~v1_funct_2(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X5,X4)),u1_struct_0(k1_tsep_1(X1,X5,X4)),u1_struct_0(X2))|~v5_pre_topc(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X5,X4)),k1_tsep_1(X1,X5,X4),X2)|~m1_subset_1(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X5,X4)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X5,X4)),u1_struct_0(X2))))|~epred1_5(X1,X2,X4,X3,X5)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_148, negated_conjecture, (~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)|~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_140, c_0_141])]), c_0_142])])).
cnf(c_0_149, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),esk4_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_145]), c_0_78]), c_0_85]), c_0_79])]), c_0_146])).
cnf(c_0_150, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk5_0),esk5_0,X1)|~epred1_5(esk1_0,X1,esk5_0,X2,esk4_0)|~v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk1_0),esk1_0,X1)|~m1_subset_1(k2_tmap_1(esk1_0,X1,X2,esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(k2_tmap_1(esk1_0,X1,X2,esk1_0),u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(k2_tmap_1(esk1_0,X1,X2,esk1_0))), inference(spm,[status(thm)],[c_0_147, c_0_134])).
cnf(c_0_151, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_152, negated_conjecture, (~v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_149])])).
cnf(c_0_153, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),esk5_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_144]), c_0_145]), c_0_78]), c_0_85]), c_0_79])]), c_0_151])).
cnf(c_0_154, plain, (m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_1(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X5,X4)))|~v1_funct_2(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X5,X4)),u1_struct_0(k1_tsep_1(X1,X5,X4)),u1_struct_0(X2))|~v5_pre_topc(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X5,X4)),k1_tsep_1(X1,X5,X4),X2)|~m1_subset_1(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X5,X4)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X5,X4)),u1_struct_0(X2))))|~epred1_5(X1,X2,X4,X3,X5)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_155, negated_conjecture, (~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_153])])).
cnf(c_0_156, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,X1,X2,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X1))))|~epred1_5(esk1_0,X1,esk5_0,X2,esk4_0)|~v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk1_0),esk1_0,X1)|~m1_subset_1(k2_tmap_1(esk1_0,X1,X2,esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(k2_tmap_1(esk1_0,X1,X2,esk1_0),u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(k2_tmap_1(esk1_0,X1,X2,esk1_0))), inference(spm,[status(thm)],[c_0_154, c_0_134])).
cnf(c_0_157, plain, (m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_1(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)))|~v1_funct_2(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)),u1_struct_0(k1_tsep_1(X1,X4,X5)),u1_struct_0(X2))|~v5_pre_topc(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)),k1_tsep_1(X1,X4,X5),X2)|~m1_subset_1(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X4,X5)),u1_struct_0(X2))))|~epred1_5(X1,X2,X5,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_158, negated_conjecture, (~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_145]), c_0_144]), c_0_144]), c_0_78]), c_0_144]), c_0_85]), c_0_144]), c_0_79])])).
cnf(c_0_159, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,X1,X2,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~epred1_5(esk1_0,X1,esk5_0,X2,esk4_0)|~v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk1_0),esk1_0,X1)|~m1_subset_1(k2_tmap_1(esk1_0,X1,X2,esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(k2_tmap_1(esk1_0,X1,X2,esk1_0),u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(k2_tmap_1(esk1_0,X1,X2,esk1_0))), inference(spm,[status(thm)],[c_0_157, c_0_134])).
cnf(c_0_160, plain, (v5_pre_topc(k2_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5)),k1_tsep_1(X1,X4,X5),X2)|~v1_funct_1(k2_tmap_1(X1,X2,X3,X4))|~v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))|~v5_pre_topc(k2_tmap_1(X1,X2,X3,X4),X4,X2)|~m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_1(k2_tmap_1(X1,X2,X3,X5))|~v1_funct_2(k2_tmap_1(X1,X2,X3,X5),u1_struct_0(X5),u1_struct_0(X2))|~v5_pre_topc(k2_tmap_1(X1,X2,X3,X5),X5,X2)|~m1_subset_1(k2_tmap_1(X1,X2,X3,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~epred1_5(X1,X2,X5,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_161, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_162, negated_conjecture, (~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159]), c_0_145]), c_0_144]), c_0_144]), c_0_78]), c_0_144]), c_0_85]), c_0_144]), c_0_79])])).
cnf(c_0_163, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_164, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk1_0),esk1_0,X1)|~epred1_5(esk1_0,X1,esk5_0,X2,esk4_0)|~v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk5_0),esk5_0,X1)|~v5_pre_topc(k2_tmap_1(esk1_0,X1,X2,esk4_0),esk4_0,X1)|~m1_subset_1(k2_tmap_1(esk1_0,X1,X2,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X1))))|~m1_subset_1(k2_tmap_1(esk1_0,X1,X2,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~v1_funct_2(k2_tmap_1(esk1_0,X1,X2,esk5_0),u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_2(k2_tmap_1(esk1_0,X1,X2,esk4_0),u1_struct_0(esk4_0),u1_struct_0(X1))|~v1_funct_1(k2_tmap_1(esk1_0,X1,X2,esk5_0))|~v1_funct_1(k2_tmap_1(esk1_0,X1,X2,esk4_0))), inference(spm,[status(thm)],[c_0_160, c_0_134])).
cnf(c_0_165, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_166, negated_conjecture, (m1_subset_1(k2_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[c_0_163, c_0_162])).
cnf(c_0_167, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_144]), c_0_145]), c_0_153]), c_0_149]), c_0_166]), c_0_141]), c_0_142]), c_0_129]), c_0_130])]), c_0_162]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 168
# Proof object clause steps            : 123
# Proof object formula steps           : 45
# Proof object conjectures             : 97
# Proof object clause conjectures      : 94
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 50
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 66
# Proof object simplifying inferences  : 228
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 87
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 87
# Processed clauses                    : 2756
# ...of these trivial                  : 51
# ...subsumed                          : 1851
# ...remaining for further processing  : 854
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 36
# Backward-rewritten                   : 88
# Generated clauses                    : 5939
# ...of the previous two non-trivial   : 5137
# Contextual simplify-reflections      : 45
# Paramodulations                      : 5927
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 656
#    Positive orientable unit clauses  : 99
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 548
# Current number of unprocessed clauses: 2468
# ...number of literals in the above   : 16349
# Current number of archived formulas  : 0
# Current number of archived clauses   : 196
# Clause-clause subsumption calls (NU) : 84373
# Rec. Clause-clause subsumption calls : 24368
# Non-unit clause-clause subsumptions  : 1868
# Unit Clause-clause subsumption calls : 1855
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 73
# BW rewrite match successes           : 21
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 248953

# -------------------------------------------------
# User time                : 0.145 s
# System time              : 0.000 s
# Total time               : 0.145 s
# Maximum resident set size: 28636 pages
