# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r3_orders_2(k2_yellow_1(X1),X2,X3)<=>r1_tarski(X2,X3))))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', t3_yellow_1)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', t1_yellow_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', d3_tarski)).
fof(dt_k3_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', dt_k3_yellow_0)).
fof(dt_k2_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&l1_orders_2(k2_yellow_1(X1))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', dt_k2_yellow_1)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', redefinition_r3_orders_2)).
fof(fc5_yellow_1, axiom, ![X1]:(v1_orders_2(k2_yellow_1(X1))&(v3_orders_2(k2_yellow_1(X1))&(v4_orders_2(k2_yellow_1(X1))&v5_orders_2(k2_yellow_1(X1))))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', fc5_yellow_1)).
fof(fc6_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(~(v2_struct_0(k2_yellow_1(X1)))&v1_orders_2(k2_yellow_1(X1)))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', fc6_yellow_1)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', d8_lattice3)).
fof(d4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v1_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r1_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', d4_yellow_0)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', t3_xboole_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', d2_subset_1)).
fof(t44_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,k3_yellow_0(X1),X2))), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', t44_yellow_0)).
fof(t13_yellow_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k1_xboole_0,X1)=>k3_yellow_0(k2_yellow_1(X1))=k1_xboole_0)), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', t13_yellow_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', t1_subset)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/910a2298da0a4246a3f2444732d45b49.E_input', t2_xboole_1)).
fof(c_0_17, plain, ![X37, X38, X39]:((~r3_orders_2(k2_yellow_1(X37),X38,X39)|r1_tarski(X38,X39)|~m1_subset_1(X39,u1_struct_0(k2_yellow_1(X37)))|~m1_subset_1(X38,u1_struct_0(k2_yellow_1(X37)))|v1_xboole_0(X37))&(~r1_tarski(X38,X39)|r3_orders_2(k2_yellow_1(X37),X38,X39)|~m1_subset_1(X39,u1_struct_0(k2_yellow_1(X37)))|~m1_subset_1(X38,u1_struct_0(k2_yellow_1(X37)))|v1_xboole_0(X37))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_yellow_1])])])])])).
fof(c_0_18, plain, ![X32]:(u1_struct_0(k2_yellow_1(X32))=X32&u1_orders_2(k2_yellow_1(X32))=k1_yellow_1(X32)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
fof(c_0_19, plain, ![X25, X26, X27]:((~v1_xboole_0(X25)|~r2_hidden(X26,X25))&(r2_hidden(esk4_1(X27),X27)|v1_xboole_0(X27))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_20, plain, ![X14, X15, X16, X17, X18]:((~r1_tarski(X14,X15)|(~r2_hidden(X16,X14)|r2_hidden(X16,X15)))&((r2_hidden(esk3_2(X17,X18),X17)|r1_tarski(X17,X18))&(~r2_hidden(esk3_2(X17,X18),X18)|r1_tarski(X17,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_21, plain, ![X20]:(~l1_orders_2(X20)|m1_subset_1(k3_yellow_0(X20),u1_struct_0(X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_yellow_0])])).
fof(c_0_22, plain, ![X42]:(v1_orders_2(k2_yellow_1(X42))&l1_orders_2(k2_yellow_1(X42))), inference(variable_rename,[status(thm)],[dt_k2_yellow_1])).
fof(c_0_23, plain, ![X33, X34, X35]:((~r3_orders_2(X33,X34,X35)|r1_orders_2(X33,X34,X35)|(v2_struct_0(X33)|(~v3_orders_2(X33)|~l1_orders_2(X33))|(~m1_subset_1(X34,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33)))))&(~r1_orders_2(X33,X34,X35)|r3_orders_2(X33,X34,X35)|(v2_struct_0(X33)|(~v3_orders_2(X33)|~l1_orders_2(X33))|(~m1_subset_1(X34,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_24, plain, (r3_orders_2(k2_yellow_1(X3),X1,X2)|v1_xboole_0(X3)|~r1_tarski(X1,X2)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X3)))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_26, plain, ![X21]:(v1_orders_2(k2_yellow_1(X21))&(v3_orders_2(k2_yellow_1(X21))&(v4_orders_2(k2_yellow_1(X21))&v5_orders_2(k2_yellow_1(X21))))), inference(variable_rename,[status(thm)],[fc5_yellow_1])).
fof(c_0_27, plain, ![X36]:((~v2_struct_0(k2_yellow_1(X36))|v1_xboole_0(X36))&(v1_orders_2(k2_yellow_1(X36))|v1_xboole_0(X36))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc6_yellow_1])])])])).
fof(c_0_28, plain, ![X6, X7, X8, X9]:((~r1_lattice3(X6,X7,X8)|(~m1_subset_1(X9,u1_struct_0(X6))|(~r2_hidden(X9,X7)|r1_orders_2(X6,X8,X9)))|~m1_subset_1(X8,u1_struct_0(X6))|~l1_orders_2(X6))&((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X6))|r1_lattice3(X6,X7,X8)|~m1_subset_1(X8,u1_struct_0(X6))|~l1_orders_2(X6))&((r2_hidden(esk2_3(X6,X7,X8),X7)|r1_lattice3(X6,X7,X8)|~m1_subset_1(X8,u1_struct_0(X6))|~l1_orders_2(X6))&(~r1_orders_2(X6,X8,esk2_3(X6,X7,X8))|r1_lattice3(X6,X7,X8)|~m1_subset_1(X8,u1_struct_0(X6))|~l1_orders_2(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
fof(c_0_29, plain, ![X29, X31]:(((m1_subset_1(esk5_1(X29),u1_struct_0(X29))|~v1_yellow_0(X29)|~l1_orders_2(X29))&(r1_lattice3(X29,u1_struct_0(X29),esk5_1(X29))|~v1_yellow_0(X29)|~l1_orders_2(X29)))&(~m1_subset_1(X31,u1_struct_0(X29))|~r1_lattice3(X29,u1_struct_0(X29),X31)|v1_yellow_0(X29)|~l1_orders_2(X29))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_yellow_0])])])])])).
cnf(c_0_30, plain, (r1_tarski(X2,X3)|v1_xboole_0(X1)|~r3_orders_2(k2_yellow_1(X1),X2,X3)|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_31, plain, ![X11]:(~r1_tarski(X11,k1_xboole_0)|X11=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_32, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_34, plain, ![X12, X13]:(((~m1_subset_1(X13,X12)|r2_hidden(X13,X12)|v1_xboole_0(X12))&(~r2_hidden(X13,X12)|m1_subset_1(X13,X12)|v1_xboole_0(X12)))&((~m1_subset_1(X13,X12)|v1_xboole_0(X13)|~v1_xboole_0(X12))&(~v1_xboole_0(X13)|m1_subset_1(X13,X12)|~v1_xboole_0(X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_35, plain, (m1_subset_1(k3_yellow_0(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, plain, (l1_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_37, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, plain, (r3_orders_2(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_25])).
cnf(c_0_39, plain, (v3_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_40, plain, (v1_xboole_0(X1)|~v2_struct_0(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_41, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (v1_yellow_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,u1_struct_0(X2),X1)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, plain, (r2_hidden(esk2_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_44, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r3_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_25]), c_0_25])).
cnf(c_0_45, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_46, plain, ![X23, X24]:(v2_struct_0(X23)|(~v5_orders_2(X23)|(~v1_yellow_0(X23)|~l1_orders_2(X23)))|(~m1_subset_1(X24,u1_struct_0(X23))|r1_orders_2(X23,k3_yellow_0(X23),X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t44_yellow_0])])])])).
cnf(c_0_47, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_48, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_49, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_50, plain, (m1_subset_1(k3_yellow_0(k2_yellow_1(X1)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_25]), c_0_36])])).
fof(c_0_51, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k1_xboole_0,X1)=>k3_yellow_0(k2_yellow_1(X1))=k1_xboole_0))), inference(assume_negation,[status(cth)],[t13_yellow_1])).
cnf(c_0_52, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk2_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_53, plain, (r1_orders_2(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X2,X3)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_25]), c_0_25]), c_0_36])]), c_0_40])).
cnf(c_0_54, plain, (r1_lattice3(k2_yellow_1(X1),X2,X3)|m1_subset_1(esk2_3(k2_yellow_1(X1),X2,X3),X1)|~m1_subset_1(X3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_25]), c_0_36])])).
cnf(c_0_55, plain, (v1_yellow_0(k2_yellow_1(X1))|~r1_lattice3(k2_yellow_1(X1),X1,X2)|~m1_subset_1(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_25]), c_0_36])])).
cnf(c_0_56, plain, (r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_32, c_0_43])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|v1_xboole_0(X3)|~r1_orders_2(k2_yellow_1(X3),X1,X2)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_39]), c_0_25]), c_0_25]), c_0_36])]), c_0_40])).
cnf(c_0_58, plain, (v2_struct_0(X1)|r1_orders_2(X1,k3_yellow_0(X1),X2)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_59, plain, (v5_orders_2(k2_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_60, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_61, plain, (v1_xboole_0(k3_yellow_0(k2_yellow_1(X1)))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_62, plain, ![X40, X41]:(~r2_hidden(X40,X41)|m1_subset_1(X40,X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_63, negated_conjecture, (~v1_xboole_0(esk1_0)&(r2_hidden(k1_xboole_0,esk1_0)&k3_yellow_0(k2_yellow_1(esk1_0))!=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_51])])])])).
fof(c_0_64, plain, ![X22]:r1_tarski(k1_xboole_0,X22), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_65, plain, (r1_lattice3(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~r1_tarski(X3,esk2_3(k2_yellow_1(X1),X2,X3))|~m1_subset_1(X3,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_25]), c_0_36])]), c_0_54])).
cnf(c_0_66, plain, (v1_yellow_0(k2_yellow_1(X1))|~m1_subset_1(X2,X1)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_25]), c_0_36])])).
cnf(c_0_67, plain, (r1_tarski(k3_yellow_0(k2_yellow_1(X1)),X2)|v1_xboole_0(X1)|~v1_yellow_0(k2_yellow_1(X1))|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_50]), c_0_59]), c_0_25]), c_0_36])]), c_0_40])).
cnf(c_0_68, plain, (k3_yellow_0(k2_yellow_1(X1))=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_69, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_70, negated_conjecture, (r2_hidden(k1_xboole_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_71, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_72, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73, plain, (r1_lattice3(k2_yellow_1(X1),X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_65, c_0_48])).
cnf(c_0_74, plain, (v1_yellow_0(k2_yellow_1(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_66, c_0_50])).
cnf(c_0_75, negated_conjecture, (k3_yellow_0(k2_yellow_1(esk1_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_76, plain, (k3_yellow_0(k2_yellow_1(X1))=k1_xboole_0|~v1_yellow_0(k2_yellow_1(X1))|~m1_subset_1(k1_xboole_0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_67]), c_0_68])).
cnf(c_0_77, negated_conjecture, (m1_subset_1(k1_xboole_0,esk1_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_79, plain, (r2_hidden(esk4_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_80, plain, (v1_yellow_0(k2_yellow_1(X1))|~m1_subset_1(X2,X1)|~v1_xboole_0(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_73]), c_0_25]), c_0_36])]), c_0_74])).
cnf(c_0_81, negated_conjecture, (~v1_yellow_0(k2_yellow_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])])).
cnf(c_0_82, plain, (r2_hidden(esk4_1(k1_xboole_0),X1)|v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_83, negated_conjecture, (~v1_xboole_0(k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_77]), c_0_81])).
cnf(c_0_84, plain, (r2_hidden(esk4_1(k1_xboole_0),X1)), inference(sr,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_85, plain, (r1_lattice3(k2_yellow_1(X1),X2,k1_xboole_0)|v1_xboole_0(X1)|~m1_subset_1(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_65, c_0_72])).
cnf(c_0_86, plain, (~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_32, c_0_84])).
cnf(c_0_87, plain, (v1_yellow_0(k2_yellow_1(X1))|~m1_subset_1(k1_xboole_0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_85]), c_0_25]), c_0_36])]), c_0_86])).
cnf(c_0_88, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_87]), c_0_77])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 89
# Proof object clause steps            : 54
# Proof object formula steps           : 35
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 37
# Processed clauses                    : 755
# ...of these trivial                  : 4
# ...subsumed                          : 461
# ...remaining for further processing  : 290
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 0
# Generated clauses                    : 1616
# ...of the previous two non-trivial   : 1527
# Contextual simplify-reflections      : 29
# Paramodulations                      : 1611
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 238
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 217
# Current number of unprocessed clauses: 819
# ...number of literals in the above   : 4423
# Current number of archived formulas  : 0
# Current number of archived clauses   : 53
# Clause-clause subsumption calls (NU) : 12910
# Rec. Clause-clause subsumption calls : 5722
# Non-unit clause-clause subsumptions  : 465
# Unit Clause-clause subsumption calls : 377
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 49115

# -------------------------------------------------
# User time                : 0.037 s
# System time              : 0.000 s
# Total time               : 0.037 s
# Maximum resident set size: 28792 pages
