# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_tmap_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>(r1_subset_1(X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))&(k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X4)=X6))))))))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t6_tmap_1)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', cc5_funct_2)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', cc1_relset_1)).
fof(d4_partfun1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>(v1_partfun1(X2,X1)<=>k1_relat_1(X2)=X1)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', d4_partfun1)).
fof(t187_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r1_xboole_0(X2,k1_relat_1(X1))=>k7_relat_1(X1,X2)=k1_xboole_0)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t187_relat_1)).
fof(redefinition_r1_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))=>(r1_subset_1(X1,X2)<=>r1_xboole_0(X1,X2))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', redefinition_r1_subset_1)).
fof(t151_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(k9_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k1_relat_1(X2),X1))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t151_relat_1)).
fof(t95_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(k7_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k1_relat_1(X2),X1))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t95_relat_1)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', redefinition_k9_subset_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t12_setfam_1)).
fof(t90_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k1_relat_1(k7_relat_1(X2,X1))=k3_xboole_0(k1_relat_1(X2),X1)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t90_relat_1)).
fof(redefinition_k2_partfun1, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k2_partfun1(X1,X2,X3,X4)=k7_relat_1(X3,X4)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', redefinition_k2_partfun1)).
fof(t148_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k2_relat_1(k7_relat_1(X2,X1))=k9_relat_1(X2,X1)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t148_relat_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t4_subset_1)).
fof(t65_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k1_relat_1(X1)=k1_xboole_0<=>k2_relat_1(X1)=k1_xboole_0)), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t65_relat_1)).
fof(t192_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k7_relat_1(X2,X1)=k7_relat_1(X2,k3_xboole_0(k1_relat_1(X2),X1))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', t192_relat_1)).
fof(dt_k1_tmap_1, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))&((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))&((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2))))))))))=>(v1_funct_1(k1_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k1_tmap_1(X1,X2,X3,X4,X5,X6),k4_subset_1(X1,X3,X4),X2)&m1_subset_1(k1_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', dt_k1_tmap_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', cc1_subset_1)).
fof(d1_tmap_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))=>![X7]:((v1_funct_1(X7)&(v1_funct_2(X7,k4_subset_1(X1,X3,X4),X2)&m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))))=>(X7=k1_tmap_1(X1,X2,X3,X4,X5,X6)<=>(k2_partfun1(k4_subset_1(X1,X3,X4),X2,X7,X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,X7,X4)=X6)))))))))), file('data/MPTP/proofs/914ff129d0564e0080b44a9ba1ba093a.E_input', d1_tmap_1)).
fof(c_0_20, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>(r1_subset_1(X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))&(k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X4)=X6)))))))))), inference(assume_negation,[status(cth)],[t6_tmap_1])).
fof(c_0_21, plain, ![X50, X51, X52]:((v1_funct_1(X52)|(~v1_funct_1(X52)|~v1_funct_2(X52,X50,X51))|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51)))|v1_xboole_0(X51))&(v1_partfun1(X52,X50)|(~v1_funct_1(X52)|~v1_funct_2(X52,X50,X51))|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51)))|v1_xboole_0(X51))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
fof(c_0_22, negated_conjecture, (~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&((~v1_xboole_0(esk3_0)&m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0)))&((~v1_xboole_0(esk4_0)&m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0)))&(r1_subset_1(esk3_0,esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,esk3_0,esk2_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk2_0)))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,esk4_0,esk2_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0)))))&(k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk1_0,esk3_0,esk4_0))!=k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk1_0,esk3_0,esk4_0))|(k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
fof(c_0_23, plain, ![X40, X41, X42]:((v4_relat_1(X42,X40)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X40,X41))))&(v5_relat_1(X42,X41)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X40,X41))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_24, plain, ![X16, X17, X18]:(~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X16,X17)))|v1_relat_1(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_25, plain, ![X57, X58]:((~v1_partfun1(X58,X57)|k1_relat_1(X58)=X57|(~v1_relat_1(X58)|~v4_relat_1(X58,X57)))&(k1_relat_1(X58)!=X57|v1_partfun1(X58,X57)|(~v1_relat_1(X58)|~v4_relat_1(X58,X57)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_partfun1])])])).
cnf(c_0_26, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (v1_funct_2(esk6_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_33, plain, ![X45, X46]:(~v1_relat_1(X45)|(~r1_xboole_0(X46,k1_relat_1(X45))|k7_relat_1(X45,X46)=k1_xboole_0)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t187_relat_1])])])).
cnf(c_0_34, plain, (k1_relat_1(X1)=X2|~v1_partfun1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (v1_partfun1(esk6_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_36, negated_conjecture, (v4_relat_1(esk6_0,esk4_0)), inference(spm,[status(thm)],[c_0_31, c_0_29])).
cnf(c_0_37, negated_conjecture, (v1_relat_1(esk6_0)), inference(spm,[status(thm)],[c_0_32, c_0_29])).
fof(c_0_38, plain, ![X43, X44]:((~r1_subset_1(X43,X44)|r1_xboole_0(X43,X44)|(v1_xboole_0(X43)|v1_xboole_0(X44)))&(~r1_xboole_0(X43,X44)|r1_subset_1(X43,X44)|(v1_xboole_0(X43)|v1_xboole_0(X44)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r1_subset_1])])])])).
cnf(c_0_39, plain, (k7_relat_1(X1,X2)=k1_xboole_0|~v1_relat_1(X1)|~r1_xboole_0(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, negated_conjecture, (k1_relat_1(esk6_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])])).
cnf(c_0_41, plain, (r1_xboole_0(X1,X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42, negated_conjecture, (r1_subset_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_43, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_44, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_45, negated_conjecture, (v1_funct_2(esk5_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_46, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_48, plain, ![X63, X64]:((k9_relat_1(X64,X63)!=k1_xboole_0|r1_xboole_0(k1_relat_1(X64),X63)|~v1_relat_1(X64))&(~r1_xboole_0(k1_relat_1(X64),X63)|k9_relat_1(X64,X63)=k1_xboole_0|~v1_relat_1(X64))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t151_relat_1])])])).
fof(c_0_49, plain, ![X14, X15]:((k7_relat_1(X15,X14)!=k1_xboole_0|r1_xboole_0(k1_relat_1(X15),X14)|~v1_relat_1(X15))&(~r1_xboole_0(k1_relat_1(X15),X14)|k7_relat_1(X15,X14)=k1_xboole_0|~v1_relat_1(X15))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_relat_1])])])).
cnf(c_0_50, negated_conjecture, (k7_relat_1(esk6_0,X1)=k1_xboole_0|~r1_xboole_0(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_37])])).
cnf(c_0_51, negated_conjecture, (r1_xboole_0(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44])).
cnf(c_0_52, negated_conjecture, (v1_partfun1(esk5_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_45]), c_0_46]), c_0_47])]), c_0_30])).
cnf(c_0_53, negated_conjecture, (v4_relat_1(esk5_0,esk3_0)), inference(spm,[status(thm)],[c_0_31, c_0_47])).
cnf(c_0_54, negated_conjecture, (v1_relat_1(esk5_0)), inference(spm,[status(thm)],[c_0_32, c_0_47])).
cnf(c_0_55, plain, (k9_relat_1(X1,X2)=k1_xboole_0|~r1_xboole_0(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56, plain, (r1_xboole_0(k1_relat_1(X1),X2)|k7_relat_1(X1,X2)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57, negated_conjecture, (k7_relat_1(esk6_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_58, plain, ![X36, X37, X38]:(~m1_subset_1(X38,k1_zfmisc_1(X36))|k9_subset_1(X36,X37,X38)=k3_xboole_0(X37,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
fof(c_0_59, plain, ![X21, X22]:k1_setfam_1(k2_tarski(X21,X22))=k3_xboole_0(X21,X22), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_60, plain, ![X61, X62]:(~v1_relat_1(X62)|k1_relat_1(k7_relat_1(X62,X61))=k3_xboole_0(k1_relat_1(X62),X61)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t90_relat_1])])).
fof(c_0_61, plain, ![X53, X54, X55, X56]:(~v1_funct_1(X55)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X54)))|k2_partfun1(X53,X54,X55,X56)=k7_relat_1(X55,X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_partfun1])])).
cnf(c_0_62, negated_conjecture, (k1_relat_1(esk5_0)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_52]), c_0_53]), c_0_54])])).
fof(c_0_63, plain, ![X19, X20]:(~v1_relat_1(X20)|k2_relat_1(k7_relat_1(X20,X19))=k9_relat_1(X20,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_relat_1])])).
cnf(c_0_64, negated_conjecture, (k9_relat_1(esk6_0,X1)=k1_xboole_0|~r1_xboole_0(esk4_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_40]), c_0_37])])).
cnf(c_0_65, negated_conjecture, (r1_xboole_0(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_40]), c_0_37])])).
fof(c_0_66, plain, ![X39]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X39)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_67, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_69, plain, (k1_relat_1(k7_relat_1(X1,X2))=k3_xboole_0(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70, plain, (k2_partfun1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_71, negated_conjecture, (k7_relat_1(esk5_0,X1)=k1_xboole_0|~r1_xboole_0(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_62]), c_0_54])])).
fof(c_0_72, plain, ![X47]:((k1_relat_1(X47)!=k1_xboole_0|k2_relat_1(X47)=k1_xboole_0|~v1_relat_1(X47))&(k2_relat_1(X47)!=k1_xboole_0|k1_relat_1(X47)=k1_xboole_0|~v1_relat_1(X47))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_relat_1])])])).
cnf(c_0_73, plain, (k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_74, negated_conjecture, (k9_relat_1(esk6_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_75, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_76, plain, ![X48, X49]:(~v1_relat_1(X49)|k7_relat_1(X49,X48)=k7_relat_1(X49,k3_xboole_0(k1_relat_1(X49),X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t192_relat_1])])).
fof(c_0_77, plain, ![X23, X24, X25, X26, X27, X28]:((v1_funct_1(k1_tmap_1(X23,X24,X25,X26,X27,X28))|(v1_xboole_0(X23)|(v1_xboole_0(X24)|(v1_xboole_0(X25)|~m1_subset_1(X25,k1_zfmisc_1(X23))|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(X23))|(~v1_funct_1(X27)|(~v1_funct_2(X27,X25,X24)|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,X24))))|(~v1_funct_1(X28)|(~v1_funct_2(X28,X26,X24)|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X24)))))))))))&((v1_funct_2(k1_tmap_1(X23,X24,X25,X26,X27,X28),k4_subset_1(X23,X25,X26),X24)|(v1_xboole_0(X23)|(v1_xboole_0(X24)|(v1_xboole_0(X25)|~m1_subset_1(X25,k1_zfmisc_1(X23))|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(X23))|(~v1_funct_1(X27)|(~v1_funct_2(X27,X25,X24)|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,X24))))|(~v1_funct_1(X28)|(~v1_funct_2(X28,X26,X24)|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X24)))))))))))&(m1_subset_1(k1_tmap_1(X23,X24,X25,X26,X27,X28),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X23,X25,X26),X24)))|(v1_xboole_0(X23)|(v1_xboole_0(X24)|(v1_xboole_0(X25)|~m1_subset_1(X25,k1_zfmisc_1(X23))|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(X23))|(~v1_funct_1(X27)|(~v1_funct_2(X27,X25,X24)|~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,X24))))|(~v1_funct_1(X28)|(~v1_funct_2(X28,X26,X24)|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X24))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tmap_1])])])])).
fof(c_0_78, plain, ![X59, X60]:(~v1_xboole_0(X59)|(~m1_subset_1(X60,k1_zfmisc_1(X59))|v1_xboole_0(X60))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_79, plain, (k9_subset_1(X2,X3,X1)=k1_setfam_1(k2_tarski(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_80, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_81, plain, (k1_relat_1(k7_relat_1(X1,X2))=k1_setfam_1(k2_tarski(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_69, c_0_68])).
cnf(c_0_82, negated_conjecture, (k7_relat_1(esk5_0,X1)=k2_partfun1(esk3_0,esk2_0,esk5_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_47]), c_0_46])])).
cnf(c_0_83, negated_conjecture, (k7_relat_1(esk5_0,esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_71, c_0_65])).
cnf(c_0_84, plain, (k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_85, negated_conjecture, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_57]), c_0_37])])).
cnf(c_0_86, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_32, c_0_75])).
cnf(c_0_87, plain, (k7_relat_1(X1,X2)=k7_relat_1(X1,k3_xboole_0(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
fof(c_0_88, plain, ![X29, X30, X31, X32, X33, X34, X35]:(((k2_partfun1(k4_subset_1(X29,X31,X32),X30,X35,X31)=X33|X35!=k1_tmap_1(X29,X30,X31,X32,X33,X34)|(~v1_funct_1(X35)|(~v1_funct_2(X35,k4_subset_1(X29,X31,X32),X30)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X29,X31,X32),X30)))))|k2_partfun1(X31,X30,X33,k9_subset_1(X29,X31,X32))!=k2_partfun1(X32,X30,X34,k9_subset_1(X29,X31,X32))|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X30)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X30)))))|(~v1_funct_1(X33)|(~v1_funct_2(X33,X31,X30)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X30)))))|(v1_xboole_0(X32)|~m1_subset_1(X32,k1_zfmisc_1(X29)))|(v1_xboole_0(X31)|~m1_subset_1(X31,k1_zfmisc_1(X29)))|v1_xboole_0(X30)|v1_xboole_0(X29))&(k2_partfun1(k4_subset_1(X29,X31,X32),X30,X35,X32)=X34|X35!=k1_tmap_1(X29,X30,X31,X32,X33,X34)|(~v1_funct_1(X35)|(~v1_funct_2(X35,k4_subset_1(X29,X31,X32),X30)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X29,X31,X32),X30)))))|k2_partfun1(X31,X30,X33,k9_subset_1(X29,X31,X32))!=k2_partfun1(X32,X30,X34,k9_subset_1(X29,X31,X32))|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X30)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X30)))))|(~v1_funct_1(X33)|(~v1_funct_2(X33,X31,X30)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X30)))))|(v1_xboole_0(X32)|~m1_subset_1(X32,k1_zfmisc_1(X29)))|(v1_xboole_0(X31)|~m1_subset_1(X31,k1_zfmisc_1(X29)))|v1_xboole_0(X30)|v1_xboole_0(X29)))&(k2_partfun1(k4_subset_1(X29,X31,X32),X30,X35,X31)!=X33|k2_partfun1(k4_subset_1(X29,X31,X32),X30,X35,X32)!=X34|X35=k1_tmap_1(X29,X30,X31,X32,X33,X34)|(~v1_funct_1(X35)|(~v1_funct_2(X35,k4_subset_1(X29,X31,X32),X30)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X29,X31,X32),X30)))))|k2_partfun1(X31,X30,X33,k9_subset_1(X29,X31,X32))!=k2_partfun1(X32,X30,X34,k9_subset_1(X29,X31,X32))|(~v1_funct_1(X34)|(~v1_funct_2(X34,X32,X30)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X30)))))|(~v1_funct_1(X33)|(~v1_funct_2(X33,X31,X30)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X30)))))|(v1_xboole_0(X32)|~m1_subset_1(X32,k1_zfmisc_1(X29)))|(v1_xboole_0(X31)|~m1_subset_1(X31,k1_zfmisc_1(X29)))|v1_xboole_0(X30)|v1_xboole_0(X29))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tmap_1])])])])])).
cnf(c_0_89, plain, (m1_subset_1(k1_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_90, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_91, plain, (v1_funct_1(k1_tmap_1(X1,X2,X3,X4,X5,X6))|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_92, plain, (v1_funct_2(k1_tmap_1(X1,X2,X3,X4,X5,X6),k4_subset_1(X1,X3,X4),X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_93, negated_conjecture, (k1_setfam_1(k2_tarski(X1,esk4_0))=k9_subset_1(esk1_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_94, negated_conjecture, (k1_setfam_1(k2_tarski(esk3_0,X1))=k1_relat_1(k2_partfun1(esk3_0,esk2_0,esk5_0,X1))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_62]), c_0_54])]), c_0_82])).
cnf(c_0_95, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,esk4_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_83, c_0_82])).
cnf(c_0_96, negated_conjecture, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])])).
cnf(c_0_97, plain, (k7_relat_1(X1,X2)=k7_relat_1(X1,k1_setfam_1(k2_tarski(k1_relat_1(X1),X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_87, c_0_68])).
cnf(c_0_98, negated_conjecture, (k7_relat_1(esk6_0,X1)=k2_partfun1(esk4_0,esk2_0,esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_29]), c_0_28])])).
cnf(c_0_99, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,X5,X3)=X6|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X4)|v1_xboole_0(X1)|X5!=k1_tmap_1(X1,X4,X2,X3,X7,X6)|~v1_funct_1(X5)|~v1_funct_2(X5,k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X2,X3),X4)))|k2_partfun1(X2,X4,X7,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_1(X6)|~v1_funct_2(X6,X3,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X7)|~v1_funct_2(X7,X2,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_100, plain, (m1_subset_1(k1_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_funct_2(X6,X4,X2)|~v1_funct_2(X5,X3,X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_101, plain, (v1_funct_1(k1_tmap_1(X1,X2,X3,X4,X5,X6))|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_funct_2(X6,X4,X2)|~v1_funct_2(X5,X3,X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[c_0_91, c_0_90])).
cnf(c_0_102, plain, (v1_funct_2(k1_tmap_1(X1,X2,X3,X4,X5,X6),k4_subset_1(X1,X3,X4),X2)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_funct_2(X6,X4,X2)|~v1_funct_2(X5,X3,X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[c_0_92, c_0_90])).
cnf(c_0_103, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk1_0,esk3_0,esk4_0))!=k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk1_0,esk3_0,esk4_0))|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_104, negated_conjecture, (k9_subset_1(esk1_0,esk3_0,esk4_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95]), c_0_96])).
cnf(c_0_105, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k1_relat_1(k2_partfun1(esk3_0,esk2_0,esk5_0,X1)))=k2_partfun1(esk3_0,esk2_0,esk5_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_62]), c_0_54])]), c_0_94]), c_0_82]), c_0_82])).
cnf(c_0_106, negated_conjecture, (k1_setfam_1(k2_tarski(esk4_0,X1))=k1_relat_1(k2_partfun1(esk4_0,esk2_0,esk6_0,X1))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_40]), c_0_37])]), c_0_98])).
cnf(c_0_107, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,k1_tmap_1(X1,X4,X2,X3,X5,X6),X3)=X6|v1_xboole_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|k2_partfun1(X2,X4,X5,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_2(X5,X2,X4)|~v1_funct_2(X6,X3,X4)|~v1_funct_1(X5)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_99, c_0_90])]), c_0_100]), c_0_101]), c_0_102])).
cnf(c_0_108, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_109, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,X5,X2)=X6|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X4)|v1_xboole_0(X1)|X5!=k1_tmap_1(X1,X4,X2,X3,X6,X7)|~v1_funct_1(X5)|~v1_funct_2(X5,k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X2,X3),X4)))|k2_partfun1(X2,X4,X6,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X7,k9_subset_1(X1,X2,X3))|~v1_funct_1(X7)|~v1_funct_2(X7,X3,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X6)|~v1_funct_2(X6,X2,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_110, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk3_0,esk2_0,esk5_0,k1_xboole_0)!=k2_partfun1(esk4_0,esk2_0,esk6_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104]), c_0_104])).
cnf(c_0_111, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_95]), c_0_96])).
cnf(c_0_112, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,k1_relat_1(k2_partfun1(esk4_0,esk2_0,esk6_0,X1)))=k2_partfun1(esk4_0,esk2_0,esk6_0,X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_97]), c_0_98]), c_0_40]), c_0_37])]), c_0_106])).
cnf(c_0_113, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,esk3_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_57, c_0_98])).
cnf(c_0_114, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),X1,k1_tmap_1(esk1_0,X1,esk3_0,esk4_0,X2,X3),esk4_0)=X3|v1_xboole_0(X1)|k2_partfun1(esk3_0,X1,X2,k1_xboole_0)!=k2_partfun1(esk4_0,X1,X3,k1_xboole_0)|~v1_funct_2(X2,esk3_0,X1)|~v1_funct_2(X3,esk4_0,X1)|~v1_funct_1(X2)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk4_0,X1)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_104]), c_0_80]), c_0_108])]), c_0_44]), c_0_43])).
cnf(c_0_115, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,k1_tmap_1(X1,X4,X2,X3,X5,X6),X2)=X5|v1_xboole_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|k2_partfun1(X2,X4,X5,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_2(X6,X3,X4)|~v1_funct_2(X5,X2,X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_109, c_0_90])]), c_0_100]), c_0_101]), c_0_102])).
cnf(c_0_116, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk4_0,esk2_0,esk6_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_117, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_96])).
cnf(c_0_118, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,X1),esk4_0)=X1|k2_partfun1(esk4_0,esk2_0,X1,k1_xboole_0)!=k1_xboole_0|~v1_funct_2(X1,esk4_0,esk2_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_111]), c_0_45]), c_0_46]), c_0_47])]), c_0_30])).
cnf(c_0_119, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),X1,k1_tmap_1(esk1_0,X1,esk3_0,esk4_0,X2,X3),esk3_0)=X2|v1_xboole_0(X1)|k2_partfun1(esk3_0,X1,X2,k1_xboole_0)!=k2_partfun1(esk4_0,X1,X3,k1_xboole_0)|~v1_funct_2(X3,esk4_0,X1)|~v1_funct_2(X2,esk3_0,X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk4_0,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X1)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_104]), c_0_80]), c_0_108])]), c_0_44]), c_0_43])).
cnf(c_0_120, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_117])])).
cnf(c_0_121, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_117]), c_0_27]), c_0_28]), c_0_29])])).
cnf(c_0_122, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,X1,esk6_0),esk3_0)=X1|k2_partfun1(esk3_0,esk2_0,X1,k1_xboole_0)!=k1_xboole_0|~v1_funct_2(X1,esk3_0,esk2_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_117]), c_0_27]), c_0_28]), c_0_29])]), c_0_30])).
cnf(c_0_123, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_124, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_111]), c_0_45]), c_0_46]), c_0_47])]), c_0_123]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 125
# Proof object clause steps            : 84
# Proof object formula steps           : 41
# Proof object conjectures             : 56
# Proof object clause conjectures      : 53
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 108
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 42
# Processed clauses                    : 866
# ...of these trivial                  : 61
# ...subsumed                          : 159
# ...remaining for further processing  : 646
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 95
# Generated clauses                    : 1957
# ...of the previous two non-trivial   : 1824
# Contextual simplify-reflections      : 12
# Paramodulations                      : 1950
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 505
#    Positive orientable unit clauses  : 185
#    Positive unorientable unit clauses: 6
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 309
# Current number of unprocessed clauses: 966
# ...number of literals in the above   : 4139
# Current number of archived formulas  : 0
# Current number of archived clauses   : 138
# Clause-clause subsumption calls (NU) : 20158
# Rec. Clause-clause subsumption calls : 10227
# Non-unit clause-clause subsumptions  : 153
# Unit Clause-clause subsumption calls : 7140
# Rewrite failures with RHS unbound    : 100
# BW rewrite match attempts            : 1022
# BW rewrite match successes           : 41
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 105730

# -------------------------------------------------
# User time                : 0.053 s
# System time              : 0.007 s
# Total time               : 0.060 s
# Maximum resident set size: 28792 pages
