# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k3_struct_0(X1)=k6_partfun1(u1_struct_0(X1))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', d4_struct_0)).
fof(d10_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1)))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', d10_tmap_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', dt_l1_pre_topc)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', t1_tsep_1)).
fof(t120_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(r1_tsep_1(X2,X3)=>(v1_funct_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X3))&(v1_funct_2(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X3),u1_struct_0(X3),u1_struct_0(k8_tmap_1(X1,X2)))&(v5_pre_topc(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X3),X3,k8_tmap_1(X1,X2))&m1_subset_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k8_tmap_1(X1,X2)))))))))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', t120_tmap_1)).
fof(dt_k7_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_funct_1(k7_tmap_1(X1,X2))&(v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))&m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', dt_k7_tmap_1)).
fof(fc3_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_funct_1(k3_struct_0(X1))&(v1_funct_2(k3_struct_0(X1),u1_struct_0(X1),u1_struct_0(X1))&v5_pre_topc(k3_struct_0(X1),X1,X1)))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', fc3_tmap_1)).
fof(t104_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)&u1_pre_topc(k6_tmap_1(X1,X2))=k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', t104_tmap_1)).
fof(t117_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(X1),k9_tmap_1(X1,X2),k3_struct_0(X1)))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', t117_tmap_1)).
fof(existence_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>?[X2]:m1_pre_topc(X2,X1)), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', existence_m1_pre_topc)).
fof(d11_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:((v1_pre_topc(X3)&(v2_pre_topc(X3)&l1_pre_topc(X3)))=>(X3=k8_tmap_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(X4=u1_struct_0(X2)=>X3=k6_tmap_1(X1,X4))))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', d11_tmap_1)).
fof(dt_k8_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_pre_topc(k8_tmap_1(X1,X2))&(v2_pre_topc(k8_tmap_1(X1,X2))&l1_pre_topc(k8_tmap_1(X1,X2))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', dt_k8_tmap_1)).
fof(redefinition_r1_funct_2, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X2))&(~(v1_xboole_0(X4))&((v1_funct_1(X5)&(v1_funct_2(X5,X1,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X3,X4)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))))))=>(r1_funct_2(X1,X2,X3,X4,X5,X6)<=>X5=X6)), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', redefinition_r1_funct_2)).
fof(dt_k9_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_funct_1(k9_tmap_1(X1,X2))&(v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))&m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', dt_k9_tmap_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', fc2_struct_0)).
fof(t111_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(r1_xboole_0(u1_struct_0(X3),X2)=>(v1_funct_1(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3))&(v1_funct_2(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),u1_struct_0(X3),u1_struct_0(k6_tmap_1(X1,X2)))&(v5_pre_topc(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),X3,k6_tmap_1(X1,X2))&m1_subset_1(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k6_tmap_1(X1,X2)))))))))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', t111_tmap_1)).
fof(dt_k2_tmap_1, axiom, ![X1, X2, X3, X4]:((l1_struct_0(X1)&(l1_struct_0(X2)&((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))&l1_struct_0(X4))))=>(v1_funct_1(k2_tmap_1(X1,X2,X3,X4))&(v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', dt_k2_tmap_1)).
fof(symmetry_r1_tsep_1, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_struct_0(X2))=>(r1_tsep_1(X1,X2)=>r1_tsep_1(X2,X1))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', symmetry_r1_tsep_1)).
fof(d3_tsep_1, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>(r1_tsep_1(X1,X2)<=>r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', d3_tsep_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/91ad5549498f4d2ab0fcbcd16d18ad78.E_input', dt_m1_pre_topc)).
fof(c_0_20, plain, ![X26]:(~l1_struct_0(X26)|k3_struct_0(X26)=k6_partfun1(u1_struct_0(X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_struct_0])])).
fof(c_0_21, plain, ![X49, X50]:(v2_struct_0(X49)|(~v2_pre_topc(X49)|~l1_pre_topc(X49))|(~m1_subset_1(X50,k1_zfmisc_1(u1_struct_0(X49)))|k7_tmap_1(X49,X50)=k6_partfun1(u1_struct_0(X49)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_tmap_1])])])])).
fof(c_0_22, plain, ![X56]:(~l1_pre_topc(X56)|l1_struct_0(X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_23, plain, (k3_struct_0(X1)=k6_partfun1(u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (v2_struct_0(X1)|k7_tmap_1(X1,X2)=k6_partfun1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (k7_tmap_1(X1,X2)=k3_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
fof(c_0_27, plain, ![X24, X25]:(~l1_pre_topc(X24)|(~m1_pre_topc(X25,X24)|m1_subset_1(u1_struct_0(X25),k1_zfmisc_1(u1_struct_0(X24))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_28, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(r1_tsep_1(X2,X3)=>(v1_funct_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X3))&(v1_funct_2(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X3),u1_struct_0(X3),u1_struct_0(k8_tmap_1(X1,X2)))&(v5_pre_topc(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X3),X3,k8_tmap_1(X1,X2))&m1_subset_1(k2_tmap_1(X1,k8_tmap_1(X1,X2),k9_tmap_1(X1,X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k8_tmap_1(X1,X2))))))))))))), inference(assume_negation,[status(cth)],[t120_tmap_1])).
fof(c_0_29, plain, ![X54, X55]:((v1_funct_1(k7_tmap_1(X54,X55))|(v2_struct_0(X54)|(~v2_pre_topc(X54)|~l1_pre_topc(X54))|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(X54)))))&((v1_funct_2(k7_tmap_1(X54,X55),u1_struct_0(X54),u1_struct_0(k6_tmap_1(X54,X55)))|(v2_struct_0(X54)|(~v2_pre_topc(X54)|~l1_pre_topc(X54))|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(X54)))))&(m1_subset_1(k7_tmap_1(X54,X55),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X54),u1_struct_0(k6_tmap_1(X54,X55)))))|(v2_struct_0(X54)|(~v2_pre_topc(X54)|~l1_pre_topc(X54))|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(X54))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_tmap_1])])])])).
cnf(c_0_30, plain, (k6_partfun1(u1_struct_0(X1))=k3_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_24, c_0_26])).
cnf(c_0_31, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_32, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&(r1_tsep_1(esk2_0,esk3_0)&(~v1_funct_1(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0))|(~v1_funct_2(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk3_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|(~v5_pre_topc(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0),esk3_0,k8_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_28])])])])).
cnf(c_0_33, plain, (v1_funct_1(k7_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (k6_partfun1(u1_struct_0(X1))=k3_struct_0(X1)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_36, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (v1_funct_1(k6_partfun1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_33, c_0_24])).
cnf(c_0_40, negated_conjecture, (k6_partfun1(u1_struct_0(esk1_0))=k3_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_41, plain, ![X16]:((v1_funct_1(k3_struct_0(X16))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&((v1_funct_2(k3_struct_0(X16),u1_struct_0(X16),u1_struct_0(X16))|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16))))&(v5_pre_topc(k3_struct_0(X16),X16,X16)|(v2_struct_0(X16)|(~v2_pre_topc(X16)|~l1_pre_topc(X16)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_tmap_1])])])])).
cnf(c_0_42, plain, (v1_funct_1(k6_partfun1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_39, c_0_31])).
cnf(c_0_43, plain, (m1_subset_1(k7_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44, negated_conjecture, (k7_tmap_1(esk1_0,X1)=k3_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_40]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_45, plain, ![X20, X21]:((u1_struct_0(k6_tmap_1(X20,X21))=u1_struct_0(X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))))&(u1_pre_topc(k6_tmap_1(X20,X21))=k5_tmap_1(X20,X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t104_tmap_1])])])])])).
cnf(c_0_46, plain, (v1_funct_2(k3_struct_0(X1),u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47, negated_conjecture, (v1_funct_1(k6_partfun1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_35]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_48, plain, ![X37, X38]:(v2_struct_0(X37)|(~v2_pre_topc(X37)|~l1_pre_topc(X37))|(v2_struct_0(X38)|~m1_pre_topc(X38,X37)|r1_funct_2(u1_struct_0(X37),u1_struct_0(k8_tmap_1(X37,X38)),u1_struct_0(X37),u1_struct_0(X37),k9_tmap_1(X37,X38),k3_struct_0(X37)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_tmap_1])])])])).
cnf(c_0_49, negated_conjecture, (m1_subset_1(k3_struct_0(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k6_tmap_1(esk1_0,X1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_50, plain, (u1_struct_0(k6_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, negated_conjecture, (v1_funct_2(k7_tmap_1(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_44]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_52, negated_conjecture, (v1_funct_1(k7_tmap_1(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_24]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_53, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_funct_2(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)),u1_struct_0(X1),u1_struct_0(X1),k9_tmap_1(X1,X2),k3_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(k3_struct_0(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_55, plain, ![X33]:(~l1_pre_topc(X33)|m1_pre_topc(esk4_1(X33),X33)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[existence_m1_pre_topc])])])).
cnf(c_0_56, negated_conjecture, (v1_funct_2(k3_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_24]), c_0_40]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_57, negated_conjecture, (v1_funct_1(k3_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_26]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_58, plain, ![X39, X40, X41, X42]:((X41!=k8_tmap_1(X39,X40)|(~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X39)))|(X42!=u1_struct_0(X40)|X41=k6_tmap_1(X39,X42)))|(~v1_pre_topc(X41)|(~v2_pre_topc(X41)|~l1_pre_topc(X41)))|~m1_pre_topc(X40,X39)|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))))&((m1_subset_1(esk5_3(X39,X40,X41),k1_zfmisc_1(u1_struct_0(X39)))|X41=k8_tmap_1(X39,X40)|(~v1_pre_topc(X41)|(~v2_pre_topc(X41)|~l1_pre_topc(X41)))|~m1_pre_topc(X40,X39)|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))))&((esk5_3(X39,X40,X41)=u1_struct_0(X40)|X41=k8_tmap_1(X39,X40)|(~v1_pre_topc(X41)|(~v2_pre_topc(X41)|~l1_pre_topc(X41)))|~m1_pre_topc(X40,X39)|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))))&(X41!=k6_tmap_1(X39,esk5_3(X39,X40,X41))|X41=k8_tmap_1(X39,X40)|(~v1_pre_topc(X41)|(~v2_pre_topc(X41)|~l1_pre_topc(X41)))|~m1_pre_topc(X40,X39)|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_tmap_1])])])])])])).
fof(c_0_59, plain, ![X35, X36]:((v1_pre_topc(k8_tmap_1(X35,X36))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|~m1_pre_topc(X36,X35)))&((v2_pre_topc(k8_tmap_1(X35,X36))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|~m1_pre_topc(X36,X35)))&(l1_pre_topc(k8_tmap_1(X35,X36))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))|~m1_pre_topc(X36,X35))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_tmap_1])])])])).
cnf(c_0_60, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,X1)),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k9_tmap_1(esk1_0,X1),k7_tmap_1(esk1_0,X2))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_44]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_61, plain, ![X10, X11, X12, X13, X14, X15]:((~r1_funct_2(X10,X11,X12,X13,X14,X15)|X14=X15|(v1_xboole_0(X11)|(v1_xboole_0(X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,X10,X11)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X10,X11))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,X12,X13)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X12,X13)))))))))&(X14!=X15|r1_funct_2(X10,X11,X12,X13,X14,X15)|(v1_xboole_0(X11)|(v1_xboole_0(X13)|(~v1_funct_1(X14)|(~v1_funct_2(X14,X10,X11)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X10,X11))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,X12,X13)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X12,X13)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r1_funct_2])])])])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(k3_struct_0(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_31]), c_0_36])])).
cnf(c_0_63, plain, (m1_pre_topc(esk4_1(X1),X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64, negated_conjecture, (v1_funct_2(k3_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_31]), c_0_36])])).
cnf(c_0_65, negated_conjecture, (v1_funct_1(k3_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_31]), c_0_36])])).
cnf(c_0_66, plain, (X1=k6_tmap_1(X2,X4)|v2_struct_0(X2)|X1!=k8_tmap_1(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X4!=u1_struct_0(X3)|~v1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_67, plain, (v2_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_68, plain, (l1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_69, plain, (v1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,X1)),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k9_tmap_1(esk1_0,X1),k3_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_24]), c_0_40]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_71, plain, (X5=X6|v1_xboole_0(X2)|v1_xboole_0(X4)|~r1_funct_2(X1,X2,X3,X4,X5,X6)|~v1_funct_1(X5)|~v1_funct_2(X5,X1,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X3,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_72, negated_conjecture, (m1_subset_1(k3_struct_0(esk1_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_36])])).
cnf(c_0_73, negated_conjecture, (v1_funct_2(k3_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_63]), c_0_36])])).
cnf(c_0_74, negated_conjecture, (v1_funct_1(k3_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_63]), c_0_36])])).
cnf(c_0_75, plain, (k6_tmap_1(X1,u1_struct_0(X2))=k8_tmap_1(X1,X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_66])]), c_0_67]), c_0_68]), c_0_31]), c_0_69])).
cnf(c_0_76, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,X1)),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k9_tmap_1(esk1_0,X1),k3_struct_0(esk1_0))|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_31]), c_0_36])])).
cnf(c_0_77, negated_conjecture, (X1=k3_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X2)|~r1_funct_2(X3,X2,u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,k3_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_2(X1,X3,X2)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_74])])).
cnf(c_0_78, plain, (u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_75]), c_0_31])).
fof(c_0_79, plain, ![X47, X48]:((v1_funct_1(k9_tmap_1(X47,X48))|(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47))|~m1_pre_topc(X48,X47)))&((v1_funct_2(k9_tmap_1(X47,X48),u1_struct_0(X47),u1_struct_0(k8_tmap_1(X47,X48)))|(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47))|~m1_pre_topc(X48,X47)))&(m1_subset_1(k9_tmap_1(X47,X48),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(k8_tmap_1(X47,X48)))))|(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47))|~m1_pre_topc(X48,X47))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k9_tmap_1])])])])).
cnf(c_0_80, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,X1)),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k9_tmap_1(esk1_0,X1),k3_struct_0(esk1_0))|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_63]), c_0_36])])).
cnf(c_0_81, negated_conjecture, (X1=k7_tmap_1(esk1_0,X2)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X3)|~r1_funct_2(X4,X3,u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1,k7_tmap_1(esk1_0,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(X1,X4,X3)|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_26]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_82, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k9_tmap_1(esk1_0,X1),k7_tmap_1(esk1_0,X2))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_78]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_83, plain, (m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_84, negated_conjecture, (r1_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),u1_struct_0(esk1_0),k9_tmap_1(esk1_0,X1),k3_struct_0(esk1_0))|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_78]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_85, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k7_tmap_1(esk1_0,X2)|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(k9_tmap_1(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(k9_tmap_1(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(k9_tmap_1(esk1_0,X1))|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_86, plain, (m1_subset_1(k9_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_83, c_0_78])).
cnf(c_0_87, plain, (v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_88, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k3_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(k9_tmap_1(esk1_0,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(k9_tmap_1(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(k9_tmap_1(esk1_0,X1))|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_77, c_0_84])).
cnf(c_0_89, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k7_tmap_1(esk1_0,X2)|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_2(k9_tmap_1(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(k9_tmap_1(esk1_0,X1))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_90, plain, (v1_funct_2(k9_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_87, c_0_78])).
cnf(c_0_91, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k3_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~v1_funct_2(k9_tmap_1(esk1_0,X1),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(k9_tmap_1(esk1_0,X1))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_86]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_92, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k7_tmap_1(esk1_0,X2)|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_funct_1(k9_tmap_1(esk1_0,X1))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_93, plain, (v1_funct_1(k9_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_94, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k3_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~v1_funct_1(k9_tmap_1(esk1_0,X1))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_90]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_95, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k7_tmap_1(esk1_0,X2)|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_96, plain, ![X44]:(v2_struct_0(X44)|~l1_struct_0(X44)|~v1_xboole_0(u1_struct_0(X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_97, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k3_struct_0(esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_93]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_98, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_99, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_100, negated_conjecture, (k9_tmap_1(esk1_0,X1)=k7_tmap_1(esk1_0,u1_struct_0(X2))|v1_xboole_0(u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_31]), c_0_36])])).
fof(c_0_101, plain, ![X51, X52, X53]:((v1_funct_1(k2_tmap_1(X51,k6_tmap_1(X51,X52),k7_tmap_1(X51,X52),X53))|~r1_xboole_0(u1_struct_0(X53),X52)|(v2_struct_0(X53)|~m1_pre_topc(X53,X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(X51)))|(v2_struct_0(X51)|(~v2_pre_topc(X51)|~l1_pre_topc(X51))))&((v1_funct_2(k2_tmap_1(X51,k6_tmap_1(X51,X52),k7_tmap_1(X51,X52),X53),u1_struct_0(X53),u1_struct_0(k6_tmap_1(X51,X52)))|~r1_xboole_0(u1_struct_0(X53),X52)|(v2_struct_0(X53)|~m1_pre_topc(X53,X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(X51)))|(v2_struct_0(X51)|(~v2_pre_topc(X51)|~l1_pre_topc(X51))))&((v5_pre_topc(k2_tmap_1(X51,k6_tmap_1(X51,X52),k7_tmap_1(X51,X52),X53),X53,k6_tmap_1(X51,X52))|~r1_xboole_0(u1_struct_0(X53),X52)|(v2_struct_0(X53)|~m1_pre_topc(X53,X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(X51)))|(v2_struct_0(X51)|(~v2_pre_topc(X51)|~l1_pre_topc(X51))))&(m1_subset_1(k2_tmap_1(X51,k6_tmap_1(X51,X52),k7_tmap_1(X51,X52),X53),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X53),u1_struct_0(k6_tmap_1(X51,X52)))))|~r1_xboole_0(u1_struct_0(X53),X52)|(v2_struct_0(X53)|~m1_pre_topc(X53,X51))|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(X51)))|(v2_struct_0(X51)|(~v2_pre_topc(X51)|~l1_pre_topc(X51))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t111_tmap_1])])])])])).
cnf(c_0_102, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_103, negated_conjecture, (k3_struct_0(esk1_0)=k9_tmap_1(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])).
cnf(c_0_104, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(X1))=k9_tmap_1(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_98]), c_0_99])).
fof(c_0_105, plain, ![X29, X30, X31, X32]:((v1_funct_1(k2_tmap_1(X29,X30,X31,X32))|(~l1_struct_0(X29)|(~l1_struct_0(X30)|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30)))))|~l1_struct_0(X32)))))&((v1_funct_2(k2_tmap_1(X29,X30,X31,X32),u1_struct_0(X32),u1_struct_0(X30))|(~l1_struct_0(X29)|(~l1_struct_0(X30)|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30)))))|~l1_struct_0(X32)))))&(m1_subset_1(k2_tmap_1(X29,X30,X31,X32),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X32),u1_struct_0(X30))))|(~l1_struct_0(X29)|(~l1_struct_0(X30)|(~v1_funct_1(X31)|(~v1_funct_2(X31,u1_struct_0(X29),u1_struct_0(X30))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X29),u1_struct_0(X30)))))|~l1_struct_0(X32))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_tmap_1])])])).
cnf(c_0_106, plain, (v5_pre_topc(k2_tmap_1(X1,k6_tmap_1(X1,X2),k7_tmap_1(X1,X2),X3),X3,k6_tmap_1(X1,X2))|v2_struct_0(X3)|v2_struct_0(X1)|~r1_xboole_0(u1_struct_0(X3),X2)|~m1_pre_topc(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_107, negated_conjecture, (k3_struct_0(esk1_0)=k9_tmap_1(esk1_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_38])).
cnf(c_0_108, plain, (v1_funct_2(k7_tmap_1(X1,X2),u1_struct_0(X1),u1_struct_0(k6_tmap_1(X1,X2)))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_109, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk2_0))=k9_tmap_1(esk1_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_104, c_0_98])).
cnf(c_0_110, negated_conjecture, (~v1_funct_1(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0))|~v1_funct_2(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk3_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~v5_pre_topc(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0),esk3_0,k8_tmap_1(esk1_0,esk2_0))|~m1_subset_1(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_111, plain, (m1_subset_1(k2_tmap_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_struct_0(X4)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_112, plain, (v1_funct_1(k2_tmap_1(X1,X2,X3,X4))|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_struct_0(X4)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_113, plain, (v1_funct_2(k2_tmap_1(X1,X2,X3,X4),u1_struct_0(X4),u1_struct_0(X2))|~l1_struct_0(X1)|~l1_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_struct_0(X4)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_114, plain, (v5_pre_topc(k2_tmap_1(X1,k8_tmap_1(X1,X2),k7_tmap_1(X1,u1_struct_0(X2)),X3),X3,k8_tmap_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X3)|~r1_xboole_0(u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_75]), c_0_31])).
cnf(c_0_115, negated_conjecture, (k7_tmap_1(esk1_0,X1)=k9_tmap_1(esk1_0,esk2_0)|~l1_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_107]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_116, plain, (v1_funct_1(k3_struct_0(X1))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_117, plain, (v1_funct_2(k7_tmap_1(X1,u1_struct_0(X2)),u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_75]), c_0_31])).
cnf(c_0_118, negated_conjecture, (k7_tmap_1(esk1_0,u1_struct_0(esk2_0))=k9_tmap_1(esk1_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_109]), c_0_38])).
cnf(c_0_119, plain, (m1_subset_1(k7_tmap_1(X1,u1_struct_0(X2)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k8_tmap_1(X1,X2)))))|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_75]), c_0_31])).
fof(c_0_120, plain, ![X22, X23]:(~l1_struct_0(X22)|~l1_struct_0(X23)|(~r1_tsep_1(X22,X23)|r1_tsep_1(X23,X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_tsep_1])])).
cnf(c_0_121, negated_conjecture, (~l1_struct_0(k8_tmap_1(esk1_0,esk2_0))|~l1_struct_0(esk3_0)|~l1_struct_0(esk1_0)|~m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|~v5_pre_topc(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,esk2_0),k9_tmap_1(esk1_0,esk2_0),esk3_0),esk3_0,k8_tmap_1(esk1_0,esk2_0))|~v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~v1_funct_1(k9_tmap_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112]), c_0_113])).
cnf(c_0_122, negated_conjecture, (v5_pre_topc(k2_tmap_1(esk1_0,k8_tmap_1(esk1_0,X1),k9_tmap_1(esk1_0,esk2_0),X2),X2,k8_tmap_1(esk1_0,X1))|v2_struct_0(X2)|~r1_xboole_0(u1_struct_0(X2),u1_struct_0(X1))|~l1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_pre_topc(X2,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_123, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_124, negated_conjecture, (v1_funct_1(k9_tmap_1(esk1_0,esk2_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_107]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_125, negated_conjecture, (v1_funct_2(k9_tmap_1(esk1_0,esk2_0),u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_98]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_126, negated_conjecture, (m1_subset_1(k9_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(k8_tmap_1(esk1_0,esk2_0)))))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_118]), c_0_98]), c_0_36]), c_0_37])]), c_0_38])).
fof(c_0_127, plain, ![X27, X28]:((~r1_tsep_1(X27,X28)|r1_xboole_0(u1_struct_0(X27),u1_struct_0(X28))|~l1_struct_0(X28)|~l1_struct_0(X27))&(~r1_xboole_0(u1_struct_0(X27),u1_struct_0(X28))|r1_tsep_1(X27,X28)|~l1_struct_0(X28)|~l1_struct_0(X27))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tsep_1])])])])).
cnf(c_0_128, plain, (r1_tsep_1(X2,X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~r1_tsep_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_129, negated_conjecture, (r1_tsep_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_130, negated_conjecture, (~r1_xboole_0(u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~l1_struct_0(k8_tmap_1(esk1_0,esk2_0))|~l1_struct_0(esk3_0)|~l1_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_35]), c_0_98])]), c_0_123]), c_0_124]), c_0_125]), c_0_126])).
cnf(c_0_131, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~r1_tsep_1(X1,X2)|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_132, negated_conjecture, (r1_tsep_1(esk3_0,esk2_0)|~l1_struct_0(esk3_0)|~l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_128, c_0_129])).
cnf(c_0_133, negated_conjecture, (~l1_struct_0(k8_tmap_1(esk1_0,esk2_0))|~l1_struct_0(esk3_0)|~l1_struct_0(esk1_0)|~l1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_132])).
cnf(c_0_134, negated_conjecture, (~l1_struct_0(esk3_0)|~l1_struct_0(esk1_0)|~l1_struct_0(esk2_0)|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_133, c_0_25])).
fof(c_0_135, plain, ![X45, X46]:(~l1_pre_topc(X45)|(~m1_pre_topc(X46,X45)|l1_pre_topc(X46))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_136, negated_conjecture, (~l1_struct_0(esk3_0)|~l1_struct_0(esk1_0)|~l1_struct_0(esk2_0)|~l1_pre_topc(k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_31]), c_0_98]), c_0_36])])).
cnf(c_0_137, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_138, negated_conjecture, (~l1_struct_0(esk3_0)|~l1_struct_0(esk1_0)|~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_68]), c_0_98]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_139, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_35]), c_0_36])])).
cnf(c_0_140, negated_conjecture, (~l1_struct_0(esk1_0)|~l1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_25]), c_0_139])])).
cnf(c_0_141, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_98]), c_0_36])])).
cnf(c_0_142, negated_conjecture, (~l1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_25]), c_0_141])])).
cnf(c_0_143, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_25]), c_0_36])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 144
# Proof object clause steps            : 103
# Proof object formula steps           : 41
# Proof object conjectures             : 66
# Proof object clause conjectures      : 63
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 37
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 65
# Proof object simplifying inferences  : 168
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 48
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 48
# Processed clauses                    : 7018
# ...of these trivial                  : 32
# ...subsumed                          : 5188
# ...remaining for further processing  : 1798
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 494
# Backward-rewritten                   : 11
# Generated clauses                    : 24298
# ...of the previous two non-trivial   : 24036
# Contextual simplify-reflections      : 289
# Paramodulations                      : 24295
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1243
#    Positive orientable unit clauses  : 36
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 1203
# Current number of unprocessed clauses: 15322
# ...number of literals in the above   : 138299
# Current number of archived formulas  : 0
# Current number of archived clauses   : 553
# Clause-clause subsumption calls (NU) : 794044
# Rec. Clause-clause subsumption calls : 61190
# Non-unit clause-clause subsumptions  : 5970
# Unit Clause-clause subsumption calls : 6228
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 374
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 904243

# -------------------------------------------------
# User time                : 0.624 s
# System time              : 0.013 s
# Total time               : 0.637 s
# Maximum resident set size: 28792 pages
