# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_S0e
# and selection function SelectLargestNegLit.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_yellow_0(X1,k6_waybel_0(X1,X2))&k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2))), file('data/MPTP/proofs/9208558ad61a4053b0fc30de8b7627c2.E_input', t39_waybel_0)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/9208558ad61a4053b0fc30de8b7627c2.E_input', t35_orders_2)).
fof(d18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k6_waybel_0(X1,X2)=k4_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2)))), file('data/MPTP/proofs/9208558ad61a4053b0fc30de8b7627c2.E_input', d18_waybel_0)).
fof(t38_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))&r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/9208558ad61a4053b0fc30de8b7627c2.E_input', t38_yellow_0)).
fof(t37_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_yellow_0(X1,X2)<=>r2_yellow_0(X1,k4_waybel_0(X1,X2))))), file('data/MPTP/proofs/9208558ad61a4053b0fc30de8b7627c2.E_input', t37_waybel_0)).
fof(t39_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/9208558ad61a4053b0fc30de8b7627c2.E_input', t39_yellow_0)).
fof(t38_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_yellow_0(X1,X2)=>k2_yellow_0(X1,X2)=k2_yellow_0(X1,k4_waybel_0(X1,X2))))), file('data/MPTP/proofs/9208558ad61a4053b0fc30de8b7627c2.E_input', t38_waybel_0)).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_yellow_0(X1,k6_waybel_0(X1,X2))&k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2)))), inference(assume_negation,[status(cth)],[t39_waybel_0])).
fof(c_0_8, plain, ![X11, X12]:((v6_orders_2(k6_domain_1(u1_struct_0(X11),X12),X11)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v3_orders_2(X11)|~l1_orders_2(X11))))&(m1_subset_1(k6_domain_1(u1_struct_0(X11),X12),k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v3_orders_2(X11)|~l1_orders_2(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
fof(c_0_9, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~r2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))|k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])).
fof(c_0_10, plain, ![X9, X10]:(v2_struct_0(X9)|~l1_orders_2(X9)|(~m1_subset_1(X10,u1_struct_0(X9))|k6_waybel_0(X9,X10)=k4_waybel_0(X9,k6_domain_1(u1_struct_0(X9),X10)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_waybel_0])])])])).
fof(c_0_11, plain, ![X15, X16]:((r1_yellow_0(X15,k6_domain_1(u1_struct_0(X15),X16))|~m1_subset_1(X16,u1_struct_0(X15))|(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15)))))&(r2_yellow_0(X15,k6_domain_1(u1_struct_0(X15),X16))|~m1_subset_1(X16,u1_struct_0(X15))|(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_yellow_0])])])])])).
fof(c_0_12, plain, ![X13, X14]:((~r2_yellow_0(X13,X14)|r2_yellow_0(X13,k4_waybel_0(X13,X14))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|(v2_struct_0(X13)|(~v3_orders_2(X13)|(~v4_orders_2(X13)|~l1_orders_2(X13)))))&(~r2_yellow_0(X13,k4_waybel_0(X13,X14))|r2_yellow_0(X13,X14)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|(v2_struct_0(X13)|(~v3_orders_2(X13)|(~v4_orders_2(X13)|~l1_orders_2(X13)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t37_waybel_0])])])])])).
cnf(c_0_13, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18, plain, (v2_struct_0(X1)|k6_waybel_0(X1,X2)=k4_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, plain, (r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_21, plain, ![X5, X6]:((k1_yellow_0(X5,k6_domain_1(u1_struct_0(X5),X6))=X6|~m1_subset_1(X6,u1_struct_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))&(k2_yellow_0(X5,k6_domain_1(u1_struct_0(X5),X6))=X6|~m1_subset_1(X6,u1_struct_0(X5))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_yellow_0])])])])])).
cnf(c_0_22, plain, (r2_yellow_0(X1,k4_waybel_0(X1,X2))|v2_struct_0(X1)|~r2_yellow_0(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])]), c_0_17])).
cnf(c_0_24, negated_conjecture, (k4_waybel_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=k6_waybel_0(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_14]), c_0_15])]), c_0_17])).
cnf(c_0_25, negated_conjecture, (r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_14]), c_0_15]), c_0_20]), c_0_16])]), c_0_17])).
cnf(c_0_26, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_27, plain, ![X7, X8]:(v2_struct_0(X7)|(~v3_orders_2(X7)|(~v4_orders_2(X7)|~l1_orders_2(X7)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|(~r2_yellow_0(X7,X8)|k2_yellow_0(X7,X8)=k2_yellow_0(X7,k4_waybel_0(X7,X8))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_waybel_0])])])])).
cnf(c_0_28, plain, (k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, negated_conjecture, (~r2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))|k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_30, negated_conjecture, (r2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25]), c_0_15]), c_0_26]), c_0_16])]), c_0_17])).
cnf(c_0_31, plain, (v2_struct_0(X1)|k2_yellow_0(X1,X2)=k2_yellow_0(X1,k4_waybel_0(X1,X2))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r2_yellow_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, negated_conjecture, (k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_14]), c_0_15]), c_0_20]), c_0_16])]), c_0_17])).
cnf(c_0_33, negated_conjecture, (k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30])])).
cnf(c_0_34, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_23]), c_0_24]), c_0_32]), c_0_25]), c_0_15]), c_0_26]), c_0_16])]), c_0_17]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 20
# Proof object formula steps           : 15
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 27
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 26
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 10
# ...of the previous two non-trivial   : 10
# Contextual simplify-reflections      : 0
# Paramodulations                      : 10
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 25
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1833

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28792 pages
