# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t67_tops_2, conjecture, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&v2_funct_1(X3))=>k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)=k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4)))))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t67_tops_2)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', redefinition_k2_relset_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', d3_struct_0)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', cc2_relset_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', fc6_relat_1)).
fof(t148_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k2_relat_1(k7_relat_1(X2,X1))=k9_relat_1(X2,X1)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t148_relat_1)).
fof(t154_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(v2_funct_1(X2)=>k9_relat_1(X2,X1)=k10_relat_1(k2_funct_1(X2),X1))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t154_funct_1)).
fof(t209_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>X2=k7_relat_1(X2,X1)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t209_relat_1)).
fof(t21_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t21_relat_1)).
fof(d4_tops_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((k2_relset_1(X1,X2,X3)=X2&v2_funct_1(X3))=>k2_tops_2(X1,X2,X3)=k2_funct_1(X3))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', d4_tops_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t3_subset)).
fof(t55_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))&k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t55_funct_1)).
fof(d9_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>k2_funct_1(X1)=k4_relat_1(X1))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', d9_funct_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', l13_xboole_0)).
fof(fc14_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>(v1_xboole_0(k4_relat_1(X1))&v1_relat_1(k4_relat_1(X1)))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', fc14_relat_1)).
fof(cc2_funct_1, axiom, ![X1]:((v1_xboole_0(X1)&(v1_relat_1(X1)&v1_funct_1(X1)))=>(v1_relat_1(X1)&(v1_funct_1(X1)&v2_funct_1(X1)))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', cc2_funct_1)).
fof(cc1_funct_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_funct_1(X1)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', cc1_funct_1)).
fof(cc1_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', cc1_relat_1)).
fof(t142_funct_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k2_relat_1(X2))<=>~(k10_relat_1(X2,k1_tarski(X1))=k1_xboole_0))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t142_funct_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t75_enumset1)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k10_relat_1(X3,X4)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', redefinition_k8_relset_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t4_subset_1)).
fof(t172_relat_1, axiom, ![X1]:k10_relat_1(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t172_relat_1)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t60_relat_1)).
fof(t45_ordinal1, axiom, ![X1]:(v1_relat_1(k1_xboole_0)&(v5_relat_1(k1_xboole_0,X1)&(v1_funct_1(k1_xboole_0)&v5_ordinal1(k1_xboole_0)))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t45_ordinal1)).
fof(t169_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t169_relat_1)).
fof(dt_k2_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k2_funct_1(X1))&v1_funct_1(k2_funct_1(X1)))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', dt_k2_funct_1)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', redefinition_k7_relset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', d3_tarski)).
fof(t132_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|v1_partfun1(X3,X1)))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t132_funct_2)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', d10_xboole_0)).
fof(d4_partfun1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>(v1_partfun1(X2,X1)<=>k1_relat_1(X2)=X1)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', d4_partfun1)).
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t69_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t3_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/92691f7c1d06471bb3f0e71ebf0754df.E_input', t113_zfmisc_1)).
fof(c_0_42, negated_conjecture, ~(![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((k2_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3)=k2_struct_0(X2)&v2_funct_1(X3))=>k7_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)=k8_relset_1(u1_struct_0(X2),u1_struct_0(X1),k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X4))))))), inference(assume_negation,[status(cth)],[t67_tops_2])).
fof(c_0_43, plain, ![X12, X13, X14]:(~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X12,X13)))|k2_relset_1(X12,X13,X14)=k2_relat_1(X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_44, negated_conjecture, (l1_struct_0(esk1_0)&(l1_struct_0(esk2_0)&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)&v2_funct_1(esk3_0))&k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0)!=k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk4_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_42])])])).
fof(c_0_45, plain, ![X15]:(~l1_struct_0(X15)|k2_struct_0(X15)=u1_struct_0(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_46, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_48, negated_conjecture, (k2_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_49, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_50, negated_conjecture, (k2_struct_0(esk2_0)=k2_relat_1(esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_51, negated_conjecture, (l1_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_52, plain, ![X16, X17, X18]:((v4_relat_1(X18,X16)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X16,X17))))&(v5_relat_1(X18,X17)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X16,X17))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_53, negated_conjecture, (u1_struct_0(esk2_0)=k2_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])])).
fof(c_0_54, plain, ![X19, X20]:(~v1_relat_1(X19)|(~m1_subset_1(X20,k1_zfmisc_1(X19))|v1_relat_1(X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_55, plain, ![X21, X22]:v1_relat_1(k2_zfmisc_1(X21,X22)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_56, plain, ![X23, X24]:(~v1_relat_1(X24)|k2_relat_1(k7_relat_1(X24,X23))=k9_relat_1(X24,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_relat_1])])).
fof(c_0_57, plain, ![X25, X26]:(~v1_relat_1(X26)|~v1_funct_1(X26)|(~v2_funct_1(X26)|k9_relat_1(X26,X25)=k10_relat_1(k2_funct_1(X26),X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t154_funct_1])])).
fof(c_0_58, plain, ![X27, X28]:(~v1_relat_1(X28)|~v4_relat_1(X28,X27)|X28=k7_relat_1(X28,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t209_relat_1])])).
cnf(c_0_59, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0))))), inference(rw,[status(thm)],[c_0_47, c_0_53])).
cnf(c_0_61, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_63, plain, ![X66]:(~v1_relat_1(X66)|r1_tarski(X66,k2_zfmisc_1(k1_relat_1(X66),k2_relat_1(X66)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_relat_1])])).
cnf(c_0_64, plain, (k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_65, plain, (k9_relat_1(X1,X2)=k10_relat_1(k2_funct_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_66, plain, (X1=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_67, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_68, negated_conjecture, (v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_60]), c_0_62])])).
cnf(c_0_69, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_70, plain, (k2_relat_1(k7_relat_1(X1,X2))=k10_relat_1(k2_funct_1(X1),X2)|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_71, negated_conjecture, (k7_relat_1(esk3_0,u1_struct_0(esk1_0))=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])])).
cnf(c_0_72, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_73, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_74, plain, ![X83, X84, X85]:(~v1_funct_1(X85)|(~v1_funct_2(X85,X83,X84)|~m1_subset_1(X85,k1_zfmisc_1(k2_zfmisc_1(X83,X84))))|(k2_relset_1(X83,X84,X85)!=X84|~v2_funct_1(X85)|k2_tops_2(X83,X84,X85)=k2_funct_1(X85))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tops_2])])).
cnf(c_0_75, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_76, plain, ![X71, X72]:((~m1_subset_1(X71,k1_zfmisc_1(X72))|r1_tarski(X71,X72))&(~r1_tarski(X71,X72)|m1_subset_1(X71,k1_zfmisc_1(X72)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_77, plain, (r1_tarski(k7_relat_1(X1,X2),k2_zfmisc_1(k1_relat_1(k7_relat_1(X1,X2)),k10_relat_1(k2_funct_1(X1),X2)))|~v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, negated_conjecture, (k10_relat_1(k2_funct_1(esk3_0),u1_struct_0(esk1_0))=k2_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_79, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk4_0)!=k8_relset_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_80, plain, (k2_tops_2(X2,X3,X1)=k2_funct_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|k2_relset_1(X2,X3,X1)!=X3|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_81, negated_conjecture, (k2_relset_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0)=k2_relat_1(esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_50]), c_0_53])).
cnf(c_0_82, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),k2_relat_1(esk3_0))), inference(rw,[status(thm)],[c_0_75, c_0_53])).
fof(c_0_83, plain, ![X29]:((k2_relat_1(X29)=k1_relat_1(k2_funct_1(X29))|~v2_funct_1(X29)|(~v1_relat_1(X29)|~v1_funct_1(X29)))&(k1_relat_1(X29)=k2_relat_1(k2_funct_1(X29))|~v2_funct_1(X29)|(~v1_relat_1(X29)|~v1_funct_1(X29)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_funct_1])])])).
fof(c_0_84, plain, ![X30]:(~v1_relat_1(X30)|~v1_funct_1(X30)|(~v2_funct_1(X30)|k2_funct_1(X30)=k4_relat_1(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_funct_1])])).
fof(c_0_85, plain, ![X31]:(~v1_xboole_0(X31)|X31=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
fof(c_0_86, plain, ![X32]:((v1_xboole_0(k4_relat_1(X32))|~v1_xboole_0(X32))&(v1_relat_1(k4_relat_1(X32))|~v1_xboole_0(X32))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc14_relat_1])])])).
fof(c_0_87, plain, ![X33]:((v1_relat_1(X33)|(~v1_xboole_0(X33)|(~v1_relat_1(X33)|~v1_funct_1(X33))))&((v1_funct_1(X33)|(~v1_xboole_0(X33)|(~v1_relat_1(X33)|~v1_funct_1(X33))))&(v2_funct_1(X33)|(~v1_xboole_0(X33)|(~v1_relat_1(X33)|~v1_funct_1(X33)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_funct_1])])])).
fof(c_0_88, plain, ![X34]:(~v1_xboole_0(X34)|v1_funct_1(X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_1])])).
fof(c_0_89, plain, ![X35]:(~v1_xboole_0(X35)|v1_relat_1(X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
cnf(c_0_90, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_91, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(k1_relat_1(esk3_0),k2_relat_1(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_71]), c_0_78]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_92, negated_conjecture, (k8_relset_1(k2_relat_1(esk3_0),u1_struct_0(esk1_0),k2_tops_2(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0),esk4_0)!=k7_relset_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_53]), c_0_53]), c_0_53])).
cnf(c_0_93, negated_conjecture, (k2_tops_2(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0)=k2_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_72]), c_0_73])]), c_0_60]), c_0_82])])).
cnf(c_0_94, plain, (k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_95, plain, (k2_funct_1(X1)=k4_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_96, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_97, plain, (v1_xboole_0(k4_relat_1(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_98, plain, (v2_funct_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_99, plain, (v1_funct_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_100, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_101, plain, ![X36, X37]:((~r2_hidden(X36,k2_relat_1(X37))|k10_relat_1(X37,k1_tarski(X36))!=k1_xboole_0|~v1_relat_1(X37))&(k10_relat_1(X37,k1_tarski(X36))=k1_xboole_0|r2_hidden(X36,k2_relat_1(X37))|~v1_relat_1(X37))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t142_funct_1])])])])).
fof(c_0_102, plain, ![X38]:k2_tarski(X38,X38)=k1_tarski(X38), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_103, plain, ![X39, X40]:k1_enumset1(X39,X39,X40)=k2_tarski(X39,X40), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_104, plain, ![X41, X42, X43]:k2_enumset1(X41,X41,X42,X43)=k1_enumset1(X41,X42,X43), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_105, plain, ![X44, X45, X46, X47]:k3_enumset1(X44,X44,X45,X46,X47)=k2_enumset1(X44,X45,X46,X47), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_106, plain, ![X48, X49, X50, X51, X52]:k4_enumset1(X48,X48,X49,X50,X51,X52)=k3_enumset1(X48,X49,X50,X51,X52), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_107, plain, ![X53, X54, X55, X56, X57, X58]:k5_enumset1(X53,X53,X54,X55,X56,X57,X58)=k4_enumset1(X53,X54,X55,X56,X57,X58), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_108, plain, ![X59, X60, X61, X62, X63, X64, X65]:k6_enumset1(X59,X59,X60,X61,X62,X63,X64,X65)=k5_enumset1(X59,X60,X61,X62,X63,X64,X65), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_109, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),k2_relat_1(esk3_0))))), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_110, negated_conjecture, (k8_relset_1(k2_relat_1(esk3_0),u1_struct_0(esk1_0),k2_funct_1(esk3_0),esk4_0)!=k7_relset_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_92, c_0_93])).
fof(c_0_111, plain, ![X99, X100, X101, X102]:(~m1_subset_1(X101,k1_zfmisc_1(k2_zfmisc_1(X99,X100)))|k8_relset_1(X99,X100,X101,X102)=k10_relat_1(X101,X102)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
fof(c_0_112, plain, ![X103]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X103)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_113, plain, ![X69]:k10_relat_1(k1_xboole_0,X69)=k1_xboole_0, inference(variable_rename,[status(thm)],[t172_relat_1])).
cnf(c_0_114, plain, (k1_relat_1(k4_relat_1(X1))=k2_relat_1(X1)|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_115, plain, (k4_relat_1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_116, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_117, plain, (v2_funct_1(X1)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_98, c_0_99]), c_0_100])).
cnf(c_0_118, plain, (~r2_hidden(X1,k2_relat_1(X2))|k10_relat_1(X2,k1_tarski(X1))!=k1_xboole_0|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_119, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_120, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_121, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_122, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_123, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_124, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_125, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_108])).
fof(c_0_126, plain, ![X70]:(v1_relat_1(k1_xboole_0)&(v5_relat_1(k1_xboole_0,X70)&(v1_funct_1(k1_xboole_0)&v5_ordinal1(k1_xboole_0)))), inference(variable_rename,[status(thm)],[t45_ordinal1])).
fof(c_0_127, plain, ![X94]:(~v1_relat_1(X94)|k10_relat_1(X94,k2_relat_1(X94))=k1_relat_1(X94)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t169_relat_1])])).
fof(c_0_128, plain, ![X87]:((v1_relat_1(k2_funct_1(X87))|(~v1_relat_1(X87)|~v1_funct_1(X87)))&(v1_funct_1(k2_funct_1(X87))|(~v1_relat_1(X87)|~v1_funct_1(X87)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_1])])])).
cnf(c_0_129, negated_conjecture, (v4_relat_1(esk3_0,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_59, c_0_109])).
cnf(c_0_130, negated_conjecture, (k8_relset_1(k2_relat_1(esk3_0),u1_struct_0(esk1_0),k4_relat_1(esk3_0),esk4_0)!=k7_relset_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_95]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_131, plain, (k8_relset_1(X2,X3,X1,X4)=k10_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_132, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_133, plain, (k10_relat_1(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_134, plain, ![X95, X96, X97, X98]:(~m1_subset_1(X97,k1_zfmisc_1(k2_zfmisc_1(X95,X96)))|k7_relset_1(X95,X96,X97,X98)=k9_relat_1(X97,X98)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
cnf(c_0_135, plain, (k10_relat_1(X1,k1_tarski(X2))=k1_xboole_0|r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_136, plain, (k2_relat_1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_116]), c_0_99]), c_0_117]), c_0_100])).
cnf(c_0_137, plain, (k10_relat_1(X2,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))!=k1_xboole_0|~v1_relat_1(X2)|~r2_hidden(X1,k2_relat_1(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_118, c_0_119]), c_0_120]), c_0_121]), c_0_122]), c_0_123]), c_0_124]), c_0_125])).
cnf(c_0_138, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_139, plain, (v1_relat_1(k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_140, plain, (k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_141, plain, (k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_142, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_143, negated_conjecture, (k7_relat_1(esk3_0,k1_relat_1(esk3_0))=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_129]), c_0_68])])).
fof(c_0_144, plain, ![X88, X89, X90, X91, X92]:((~r1_tarski(X88,X89)|(~r2_hidden(X90,X88)|r2_hidden(X90,X89)))&((r2_hidden(esk6_2(X91,X92),X91)|r1_tarski(X91,X92))&(~r2_hidden(esk6_2(X91,X92),X92)|r1_tarski(X91,X92)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_145, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_146, negated_conjecture, (k8_relset_1(k2_relat_1(esk3_0),u1_struct_0(esk1_0),k1_xboole_0,esk4_0)!=k7_relset_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0,esk4_0)|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_130, c_0_115])).
cnf(c_0_147, plain, (k8_relset_1(X1,X2,k1_xboole_0,X3)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_133])).
cnf(c_0_148, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_149, plain, (k10_relat_1(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))=k1_xboole_0|r2_hidden(X2,k2_relat_1(X1))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_135, c_0_119]), c_0_120]), c_0_121]), c_0_122]), c_0_123]), c_0_124]), c_0_125])).
cnf(c_0_150, plain, (k2_relat_1(k4_relat_1(X1))=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_136, c_0_97])).
cnf(c_0_151, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_133]), c_0_138]), c_0_139])])).
cnf(c_0_152, plain, (v1_relat_1(k4_relat_1(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_153, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_90, c_0_69])).
cnf(c_0_154, plain, (k10_relat_1(k2_funct_1(X1),k1_relat_1(X1))=k1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_142])).
cnf(c_0_155, negated_conjecture, (k10_relat_1(k2_funct_1(esk3_0),k1_relat_1(esk3_0))=k2_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_143]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_156, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_157, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_145, c_0_60])).
cnf(c_0_158, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0,esk4_0)!=k1_xboole_0|~v1_xboole_0(esk3_0)), inference(rw,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_159, negated_conjecture, (k7_relset_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0),esk3_0,X1)=k9_relat_1(esk3_0,X1)), inference(spm,[status(thm)],[c_0_148, c_0_60])).
cnf(c_0_160, plain, (k10_relat_1(k4_relat_1(X1),k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))=k1_xboole_0|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_151]), c_0_152])).
fof(c_0_161, plain, ![X104, X105, X106]:((X105=k1_xboole_0|v1_partfun1(X106,X104)|(~v1_funct_1(X106)|(~v1_funct_2(X106,X104,X105)|~m1_subset_1(X106,k1_zfmisc_1(k2_zfmisc_1(X104,X105)))))|(~v1_funct_1(X106)|~m1_subset_1(X106,k1_zfmisc_1(k2_zfmisc_1(X104,X105)))))&(X104!=k1_xboole_0|v1_partfun1(X106,X104)|(~v1_funct_1(X106)|(~v1_funct_2(X106,X104,X105)|~m1_subset_1(X106,k1_zfmisc_1(k2_zfmisc_1(X104,X105)))))|(~v1_funct_1(X106)|~m1_subset_1(X106,k1_zfmisc_1(k2_zfmisc_1(X104,X105)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t132_funct_2])])])])).
cnf(c_0_162, plain, (m1_subset_1(k2_funct_1(X1),k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(k2_funct_1(X1)),k1_relat_1(X1))))|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_141]), c_0_142])).
cnf(c_0_163, negated_conjecture, (k1_relat_1(k2_funct_1(esk3_0))=k2_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_164, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_165, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_166, negated_conjecture, (k9_relat_1(esk3_0,esk4_0)!=k1_xboole_0|~v1_xboole_0(esk3_0)), inference(rw,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_167, plain, (~r2_hidden(X1,k2_relat_1(k4_relat_1(X2)))|~v1_xboole_0(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_160]), c_0_152])).
cnf(c_0_168, plain, (k2_relat_1(k4_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_141, c_0_95])).
cnf(c_0_169, plain, (X1=k1_xboole_0|v1_partfun1(X2,X3)|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_170, negated_conjecture, (m1_subset_1(k2_funct_1(esk3_0),k1_zfmisc_1(k2_zfmisc_1(k2_relat_1(esk3_0),k1_relat_1(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_171, plain, (r2_hidden(X1,k2_zfmisc_1(k1_relat_1(X2),k2_relat_1(X2)))|~r2_hidden(X1,X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_156, c_0_69])).
cnf(c_0_172, negated_conjecture, (r1_tarski(esk3_0,X1)|r2_hidden(esk6_2(esk3_0,X1),k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_173, negated_conjecture, (k10_relat_1(k2_funct_1(esk3_0),esk4_0)!=k1_xboole_0|~v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_65]), c_0_68]), c_0_72]), c_0_73])])).
fof(c_0_174, plain, ![X67, X68]:(((r1_tarski(X67,X68)|X67!=X68)&(r1_tarski(X68,X67)|X67!=X68))&(~r1_tarski(X67,X68)|~r1_tarski(X68,X67)|X67=X68)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_175, plain, (~r2_hidden(X1,k1_relat_1(X2))|~v1_xboole_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_99]), c_0_117]), c_0_100])).
fof(c_0_176, plain, ![X107, X108]:((~v1_partfun1(X108,X107)|k1_relat_1(X108)=X107|(~v1_relat_1(X108)|~v4_relat_1(X108,X107)))&(k1_relat_1(X108)!=X107|v1_partfun1(X108,X107)|(~v1_relat_1(X108)|~v4_relat_1(X108,X107)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_partfun1])])])).
cnf(c_0_177, plain, (X1=k1_xboole_0|v1_partfun1(X2,X3)|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(cn,[status(thm)],[c_0_169])).
cnf(c_0_178, plain, (k2_relset_1(k1_relat_1(X1),k2_relat_1(X1),X1)=k2_relat_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_46, c_0_153])).
cnf(c_0_179, negated_conjecture, (v1_relat_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_170]), c_0_62])])).
fof(c_0_180, plain, ![X75, X76]:(v1_xboole_0(X76)|(~r1_tarski(X76,X75)|~r1_xboole_0(X76,X75))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
cnf(c_0_181, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_182, negated_conjecture, (r1_tarski(esk3_0,X1)|r2_hidden(esk6_2(esk3_0,X1),k2_zfmisc_1(k1_relat_1(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0))),k2_relat_1(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172]), c_0_62])])).
cnf(c_0_183, negated_conjecture, (k10_relat_1(k4_relat_1(esk3_0),esk4_0)!=k1_xboole_0|~v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_95]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_184, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_185, plain, (r1_tarski(k1_relat_1(X1),X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_175, c_0_165])).
fof(c_0_186, plain, ![X77, X78, X80, X81, X82]:(((r2_hidden(esk5_2(X77,X78),X77)|r1_xboole_0(X77,X78))&(r2_hidden(esk5_2(X77,X78),X78)|r1_xboole_0(X77,X78)))&(~r2_hidden(X82,X80)|~r2_hidden(X82,X81)|~r1_xboole_0(X80,X81))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_187, plain, (k1_relat_1(X1)=X2|~v1_partfun1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_188, negated_conjecture, (k2_relat_1(esk3_0)=k1_xboole_0|v1_partfun1(esk3_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_82]), c_0_73])]), c_0_60])])).
cnf(c_0_189, negated_conjecture, (k2_relset_1(k2_relat_1(esk3_0),k2_relat_1(k2_funct_1(esk3_0)),k2_funct_1(esk3_0))=k2_relat_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_163]), c_0_179])])).
cnf(c_0_190, negated_conjecture, (k2_relset_1(k2_relat_1(esk3_0),k1_relat_1(esk3_0),k2_funct_1(esk3_0))=k2_relat_1(k2_funct_1(esk3_0))), inference(spm,[status(thm)],[c_0_46, c_0_170])).
cnf(c_0_191, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_180])).
cnf(c_0_192, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(k1_relat_1(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0))),k2_relat_1(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))))), inference(spm,[status(thm)],[c_0_181, c_0_182])).
cnf(c_0_193, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_115]), c_0_133])])).
cnf(c_0_194, plain, (X1=k1_relat_1(X2)|~r1_tarski(X1,k1_relat_1(X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_184, c_0_185])).
fof(c_0_195, plain, ![X73, X74]:((k2_zfmisc_1(X73,X74)!=k1_xboole_0|(X73=k1_xboole_0|X74=k1_xboole_0))&((X73!=k1_xboole_0|k2_zfmisc_1(X73,X74)=k1_xboole_0)&(X74!=k1_xboole_0|k2_zfmisc_1(X73,X74)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_196, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_151, c_0_165])).
cnf(c_0_197, plain, (r2_hidden(esk5_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_198, negated_conjecture, (k8_relset_1(k2_relat_1(esk3_0),u1_struct_0(esk1_0),k2_funct_1(esk3_0),esk4_0)!=k9_relat_1(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_110, c_0_159])).
cnf(c_0_199, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|k2_relat_1(esk3_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_68]), c_0_67])])).
cnf(c_0_200, plain, (k8_relset_1(k1_relat_1(X1),k2_relat_1(X1),X1,X2)=k10_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_131, c_0_153])).
cnf(c_0_201, negated_conjecture, (k2_relat_1(k2_funct_1(esk3_0))=k1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_141]), c_0_190]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_202, negated_conjecture, (~r1_xboole_0(esk3_0,k2_zfmisc_1(k1_relat_1(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0))),k2_relat_1(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192]), c_0_193])).
cnf(c_0_203, plain, (k1_relat_1(X1)=k1_relat_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_194, c_0_185])).
cnf(c_0_204, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_195])).
cnf(c_0_205, plain, (v1_xboole_0(k1_xboole_0)|~r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_191, c_0_196])).
cnf(c_0_206, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_151, c_0_197])).
cnf(c_0_207, negated_conjecture, (k2_relat_1(esk3_0)=k1_xboole_0|k8_relset_1(k2_relat_1(esk3_0),k1_relat_1(esk3_0),k2_funct_1(esk3_0),esk4_0)!=k9_relat_1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_198, c_0_199])).
cnf(c_0_208, negated_conjecture, (k8_relset_1(k2_relat_1(esk3_0),k1_relat_1(esk3_0),k2_funct_1(esk3_0),X1)=k10_relat_1(k2_funct_1(esk3_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_163]), c_0_201]), c_0_179])])).
cnf(c_0_209, negated_conjecture, (~r1_xboole_0(esk3_0,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_210, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_204])).
cnf(c_0_211, plain, (v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_205, c_0_206])).
cnf(c_0_212, negated_conjecture, (k2_relat_1(esk3_0)=k1_xboole_0|k9_relat_1(esk3_0,esk4_0)!=k10_relat_1(k2_funct_1(esk3_0),esk4_0)), inference(rw,[status(thm)],[c_0_207, c_0_208])).
cnf(c_0_213, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_195])).
cnf(c_0_214, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),k2_relat_1(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_116]), c_0_210]), c_0_206]), c_0_211])])).
cnf(c_0_215, negated_conjecture, (k2_relat_1(esk3_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_65]), c_0_68]), c_0_72]), c_0_73])])).
cnf(c_0_216, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_213])).
cnf(c_0_217, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_214, c_0_215]), c_0_216]), c_0_211])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 218
# Proof object clause steps            : 134
# Proof object formula steps           : 84
# Proof object conjectures             : 57
# Proof object clause conjectures      : 54
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 56
# Proof object initial formulas used   : 42
# Proof object generating inferences   : 62
# Proof object simplifying inferences  : 131
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 43
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 74
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 65
# Processed clauses                    : 22006
# ...of these trivial                  : 27
# ...subsumed                          : 17682
# ...remaining for further processing  : 4297
# Other redundant clauses eliminated   : 26
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 368
# Backward-rewritten                   : 694
# Generated clauses                    : 237103
# ...of the previous two non-trivial   : 164649
# Contextual simplify-reflections      : 1035
# Paramodulations                      : 237077
# Factorizations                       : 0
# Equation resolutions                 : 26
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3229
#    Positive orientable unit clauses  : 50
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 3177
# Current number of unprocessed clauses: 142029
# ...number of literals in the above   : 561611
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1069
# Clause-clause subsumption calls (NU) : 2264920
# Rec. Clause-clause subsumption calls : 1537983
# Non-unit clause-clause subsumptions  : 7270
# Unit Clause-clause subsumption calls : 608
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 107
# BW rewrite match successes           : 48
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3969146

# -------------------------------------------------
# User time                : 2.117 s
# System time              : 0.023 s
# Total time               : 2.140 s
# Maximum resident set size: 109528 pages
