# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/92883a567a264d108a3e4ecdd4a688da.E_input', t36_waybel_0)).
fof(t9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X3,X4)=>r1_lattice3(X1,X2,X4))&(r2_lattice3(X1,X3,X4)=>r2_lattice3(X1,X2,X4)))))), file('data/MPTP/proofs/92883a567a264d108a3e4ecdd4a688da.E_input', t9_yellow_0)).
fof(t16_waybel_0, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,k3_waybel_0(X1,X2))&r1_tarski(X2,k4_waybel_0(X1,X2))))), file('data/MPTP/proofs/92883a567a264d108a3e4ecdd4a688da.E_input', t16_waybel_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/92883a567a264d108a3e4ecdd4a688da.E_input', d8_lattice3)).
fof(fraenkel_a_2_1_waybel_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X5,X4)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/92883a567a264d108a3e4ecdd4a688da.E_input', fraenkel_a_2_1_waybel_0)).
fof(t4_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>![X4]:((r1_lattice3(X1,X4,X3)=>r1_lattice3(X1,X4,X2))&(r2_lattice3(X1,X4,X2)=>r2_lattice3(X1,X4,X3))))))), file('data/MPTP/proofs/92883a567a264d108a3e4ecdd4a688da.E_input', t4_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/92883a567a264d108a3e4ecdd4a688da.E_input', t7_yellow_0)).
fof(t15_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2))), file('data/MPTP/proofs/92883a567a264d108a3e4ecdd4a688da.E_input', t15_waybel_0)).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t36_waybel_0])).
fof(c_0_9, plain, ![X12, X13, X14, X15]:((~r1_lattice3(X12,X14,X15)|r1_lattice3(X12,X13,X15)|~m1_subset_1(X15,u1_struct_0(X12))|~r1_tarski(X13,X14)|~l1_orders_2(X12))&(~r2_lattice3(X12,X14,X15)|r2_lattice3(X12,X13,X15)|~m1_subset_1(X15,u1_struct_0(X12))|~r1_tarski(X13,X14)|~l1_orders_2(X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_yellow_0])])])])).
fof(c_0_10, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0))&(r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
fof(c_0_11, plain, ![X30, X31]:((r1_tarski(X31,k3_waybel_0(X30,X31))|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(~v3_orders_2(X30)|~l1_orders_2(X30)))&(r1_tarski(X31,k4_waybel_0(X30,X31))|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X30)))|(~v3_orders_2(X30)|~l1_orders_2(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_waybel_0])])])])).
cnf(c_0_12, plain, (r1_lattice3(X1,X4,X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_tarski(X4,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r1_tarski(X1,k4_waybel_0(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_18, plain, ![X23, X24, X25, X26]:((~r1_lattice3(X23,X24,X25)|(~m1_subset_1(X26,u1_struct_0(X23))|(~r2_hidden(X26,X24)|r1_orders_2(X23,X25,X26)))|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&((m1_subset_1(esk6_3(X23,X24,X25),u1_struct_0(X23))|r1_lattice3(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&((r2_hidden(esk6_3(X23,X24,X25),X24)|r1_lattice3(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))&(~r1_orders_2(X23,X25,esk6_3(X23,X24,X25))|r1_lattice3(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~l1_orders_2(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_19, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)|r1_lattice3(esk1_0,X1,esk3_0)|~r1_tarski(X1,k4_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])])).
cnf(c_0_20, negated_conjecture, (r1_tarski(X1,k4_waybel_0(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_15])])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_22, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
fof(c_0_24, plain, ![X16, X17, X18, X21, X22]:(((m1_subset_1(esk4_3(X16,X17,X18),u1_struct_0(X17))|~r2_hidden(X16,a_2_1_waybel_0(X17,X18))|(v2_struct_0(X17)|~l1_orders_2(X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))))&((X16=esk4_3(X16,X17,X18)|~r2_hidden(X16,a_2_1_waybel_0(X17,X18))|(v2_struct_0(X17)|~l1_orders_2(X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))))&((m1_subset_1(esk5_3(X16,X17,X18),u1_struct_0(X17))|~r2_hidden(X16,a_2_1_waybel_0(X17,X18))|(v2_struct_0(X17)|~l1_orders_2(X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))))&((r1_orders_2(X17,esk5_3(X16,X17,X18),esk4_3(X16,X17,X18))|~r2_hidden(X16,a_2_1_waybel_0(X17,X18))|(v2_struct_0(X17)|~l1_orders_2(X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))))&(r2_hidden(esk5_3(X16,X17,X18),X18)|~r2_hidden(X16,a_2_1_waybel_0(X17,X18))|(v2_struct_0(X17)|~l1_orders_2(X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))))))))&(~m1_subset_1(X21,u1_struct_0(X17))|(X16!=X21|(~m1_subset_1(X22,u1_struct_0(X17))|(~r1_orders_2(X17,X22,X21)|~r2_hidden(X22,X18))))|r2_hidden(X16,a_2_1_waybel_0(X17,X18))|(v2_struct_0(X17)|~l1_orders_2(X17)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_waybel_0])])])])])])).
fof(c_0_25, plain, ![X32, X33, X34, X35]:((~r1_lattice3(X32,X35,X34)|r1_lattice3(X32,X35,X33)|~r1_orders_2(X32,X33,X34)|~m1_subset_1(X34,u1_struct_0(X32))|~m1_subset_1(X33,u1_struct_0(X32))|(~v4_orders_2(X32)|~l1_orders_2(X32)))&(~r2_lattice3(X32,X35,X33)|r2_lattice3(X32,X35,X34)|~r1_orders_2(X32,X33,X34)|~m1_subset_1(X34,u1_struct_0(X32))|~m1_subset_1(X33,u1_struct_0(X32))|(~v4_orders_2(X32)|~l1_orders_2(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_yellow_0])])])])).
cnf(c_0_26, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_14]), c_0_15])])).
cnf(c_0_27, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_28, plain, ![X9, X10, X11]:((~r1_lattice3(X9,k1_tarski(X11),X10)|r1_orders_2(X9,X10,X11)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))&((~r1_orders_2(X9,X10,X11)|r1_lattice3(X9,k1_tarski(X11),X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))&((~r2_lattice3(X9,k1_tarski(X11),X10)|r1_orders_2(X9,X11,X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))&(~r1_orders_2(X9,X11,X10)|r2_lattice3(X9,k1_tarski(X11),X10)|~m1_subset_1(X11,u1_struct_0(X9))|~m1_subset_1(X10,u1_struct_0(X9))|~l1_orders_2(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_29, plain, (r1_lattice3(X1,X2,X4)|~r1_lattice3(X1,X2,X3)|~r1_orders_2(X1,X4,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk5_3(X1,X2,esk2_0))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,esk2_0))|~m1_subset_1(esk5_3(X1,X2,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_32, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r1_orders_2(X1,esk5_3(X2,X1,X3),esk4_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_waybel_0(X1,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, negated_conjecture, (r1_lattice3(esk1_0,X1,esk3_0)|v2_struct_0(X2)|~r2_hidden(X3,a_2_1_waybel_0(X2,esk2_0))|~r1_lattice3(esk1_0,X1,esk5_3(X3,X2,esk2_0))|~m1_subset_1(esk5_3(X3,X2,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_14]), c_0_15]), c_0_31])])).
cnf(c_0_37, plain, (r1_lattice3(X1,k1_tarski(esk4_3(X2,X1,X3)),esk5_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_39, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(esk4_3(X1,esk1_0,esk2_0)),esk3_0)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(esk5_3(X1,esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_21]), c_0_15])]), c_0_38])).
cnf(c_0_40, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(esk4_3(X1,esk1_0,esk2_0)),esk3_0)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_34]), c_0_21]), c_0_15])]), c_0_38])).
cnf(c_0_41, plain, (X1=esk4_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_42, plain, ![X28, X29]:(v2_struct_0(X28)|~l1_orders_2(X28)|(~m1_subset_1(X29,k1_zfmisc_1(u1_struct_0(X28)))|k4_waybel_0(X28,X29)=a_2_1_waybel_0(X28,X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_waybel_0])])])])).
cnf(c_0_43, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_44, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(X1),esk3_0)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_21]), c_0_15])]), c_0_38])).
cnf(c_0_45, negated_conjecture, (~r1_lattice3(esk1_0,esk2_0,esk3_0)|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_46, plain, (v2_struct_0(X1)|k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_47, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,a_2_1_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_14]), c_0_15])])).
cnf(c_0_48, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_49, negated_conjecture, (~r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)|~r1_lattice3(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_21]), c_0_15])]), c_0_38])).
cnf(c_0_50, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_51, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,esk6_3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2))|r1_lattice3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2)|~m1_subset_1(esk6_3(X1,a_2_1_waybel_0(esk1_0,esk2_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52, negated_conjecture, (~r1_lattice3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_23])])).
cnf(c_0_53, negated_conjecture, (~m1_subset_1(esk6_3(esk1_0,a_2_1_waybel_0(esk1_0,esk2_0),esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_14]), c_0_15])]), c_0_52])).
cnf(c_0_54, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_55, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_14]), c_0_15])]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 39
# Proof object formula steps           : 17
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 261
# ...of these trivial                  : 0
# ...subsumed                          : 67
# ...remaining for further processing  : 194
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 5
# Generated clauses                    : 369
# ...of the previous two non-trivial   : 357
# Contextual simplify-reflections      : 15
# Paramodulations                      : 368
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 145
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 135
# Current number of unprocessed clauses: 148
# ...number of literals in the above   : 1157
# Current number of archived formulas  : 0
# Current number of archived clauses   : 48
# Clause-clause subsumption calls (NU) : 6823
# Rec. Clause-clause subsumption calls : 687
# Non-unit clause-clause subsumptions  : 95
# Unit Clause-clause subsumption calls : 41
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 18170

# -------------------------------------------------
# User time                : 0.022 s
# System time              : 0.000 s
# Total time               : 0.022 s
# Maximum resident set size: 28624 pages
