# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/932c7ad8c20b4ae69927b38dd13e2b00.E_input', d4_relat_1)).
fof(t22_relset_1, conjecture, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X4,X5),X3))))<=>k1_relset_1(X2,X1,X3)=X2)), file('data/MPTP/proofs/932c7ad8c20b4ae69927b38dd13e2b00.E_input', t22_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/932c7ad8c20b4ae69927b38dd13e2b00.E_input', redefinition_k1_relset_1)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/932c7ad8c20b4ae69927b38dd13e2b00.E_input', t6_xboole_0)).
fof(dt_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k1_relset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/932c7ad8c20b4ae69927b38dd13e2b00.E_input', dt_k1_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/932c7ad8c20b4ae69927b38dd13e2b00.E_input', t3_subset)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/932c7ad8c20b4ae69927b38dd13e2b00.E_input', d8_xboole_0)).
fof(c_0_7, plain, ![X6, X7, X8, X10, X11, X12, X13, X15]:(((~r2_hidden(X8,X7)|r2_hidden(k4_tarski(X8,esk1_3(X6,X7,X8)),X6)|X7!=k1_relat_1(X6))&(~r2_hidden(k4_tarski(X10,X11),X6)|r2_hidden(X10,X7)|X7!=k1_relat_1(X6)))&((~r2_hidden(esk2_2(X12,X13),X13)|~r2_hidden(k4_tarski(esk2_2(X12,X13),X15),X12)|X13=k1_relat_1(X12))&(r2_hidden(esk2_2(X12,X13),X13)|r2_hidden(k4_tarski(esk2_2(X12,X13),esk3_2(X12,X13)),X12)|X13=k1_relat_1(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X4,X5),X3))))<=>k1_relset_1(X2,X1,X3)=X2))), inference(assume_negation,[status(cth)],[t22_relset_1])).
cnf(c_0_9, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, negated_conjecture, ![X21, X22]:(m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk5_0,esk4_0)))&(((r2_hidden(esk7_0,esk5_0)|k1_relset_1(esk5_0,esk4_0,esk6_0)!=esk5_0)&(~r2_hidden(k4_tarski(esk7_0,X21),esk6_0)|k1_relset_1(esk5_0,esk4_0,esk6_0)!=esk5_0))&(~r2_hidden(X22,esk5_0)|r2_hidden(k4_tarski(X22,esk8_1(X22)),esk6_0)|k1_relset_1(esk5_0,esk4_0,esk6_0)=esk5_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
fof(c_0_11, plain, ![X24, X25, X26]:(~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))|k1_relset_1(X24,X25,X26)=k1_relat_1(X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_12, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (r2_hidden(k4_tarski(X1,esk8_1(X1)),esk6_0)|k1_relset_1(esk5_0,esk4_0,esk6_0)=esk5_0|~r2_hidden(X1,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X27, X28]:((r2_hidden(esk9_2(X27,X28),X28)|~r2_xboole_0(X27,X28))&(~r2_hidden(esk9_2(X27,X28),X27)|~r2_xboole_0(X27,X28))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
fof(c_0_15, plain, ![X30, X31, X32]:(~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))|m1_subset_1(k1_relset_1(X30,X31,X32),k1_zfmisc_1(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
cnf(c_0_16, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk5_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, negated_conjecture, (k1_relset_1(esk5_0,esk4_0,esk6_0)=esk5_0|r2_hidden(X1,k1_relat_1(esk6_0))|~r2_hidden(X1,esk5_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, plain, (r2_hidden(esk9_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_20, plain, ![X33, X34]:((~m1_subset_1(X33,k1_zfmisc_1(X34))|r1_tarski(X33,X34))&(~r1_tarski(X33,X34)|m1_subset_1(X33,k1_zfmisc_1(X34)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_21, plain, (m1_subset_1(k1_relset_1(X2,X3,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, negated_conjecture, (k1_relset_1(esk5_0,esk4_0,esk6_0)=k1_relat_1(esk6_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23, plain, (~r2_hidden(esk9_2(X1,X2),X1)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, negated_conjecture, (k1_relset_1(esk5_0,esk4_0,esk6_0)=esk5_0|r2_hidden(esk9_2(X1,esk5_0),k1_relat_1(esk6_0))|~r2_xboole_0(X1,esk5_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, negated_conjecture, (~r2_hidden(k4_tarski(esk7_0,X1),esk6_0)|k1_relset_1(esk5_0,esk4_0,esk6_0)!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, plain, (r2_hidden(k4_tarski(X1,esk1_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_27, plain, ![X35, X36]:(((r1_tarski(X35,X36)|~r2_xboole_0(X35,X36))&(X35!=X36|~r2_xboole_0(X35,X36)))&(~r1_tarski(X35,X36)|X35=X36|r2_xboole_0(X35,X36))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(k1_relat_1(esk6_0),k1_zfmisc_1(esk5_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_17]), c_0_22])).
cnf(c_0_30, negated_conjecture, (k1_relset_1(esk5_0,esk4_0,esk6_0)=esk5_0|~r2_xboole_0(k1_relat_1(esk6_0),esk5_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, negated_conjecture, (k1_relat_1(esk6_0)!=esk5_0|~r2_hidden(k4_tarski(esk7_0,X1),esk6_0)), inference(rw,[status(thm)],[c_0_25, c_0_22])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,esk1_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (r1_tarski(k1_relat_1(esk6_0),esk5_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, negated_conjecture, (k1_relat_1(esk6_0)=esk5_0|~r2_xboole_0(k1_relat_1(esk6_0),esk5_0)), inference(rw,[status(thm)],[c_0_30, c_0_22])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk7_0,esk5_0)|k1_relset_1(esk5_0,esk4_0,esk6_0)!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_37, negated_conjecture, (k1_relat_1(esk6_0)!=esk5_0|~r2_hidden(esk7_0,k1_relat_1(esk6_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, negated_conjecture, (k1_relat_1(esk6_0)=esk5_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk7_0,esk5_0)|k1_relat_1(esk6_0)!=esk5_0), inference(rw,[status(thm)],[c_0_36, c_0_22])).
cnf(c_0_40, negated_conjecture, (~r2_hidden(esk7_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_38])])).
cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_38])]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 27
# Proof object formula steps           : 15
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 56
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 13
# Generated clauses                    : 24
# ...of the previous two non-trivial   : 32
# Contextual simplify-reflections      : 1
# Paramodulations                      : 21
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 20
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 87
# Rec. Clause-clause subsumption calls : 80
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1573

# -------------------------------------------------
# User time                : 0.005 s
# System time              : 0.004 s
# Total time               : 0.009 s
# Maximum resident set size: 28656 pages
