# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t4_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', commutativity_k3_xboole_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t7_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t1_boole)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t28_xboole_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t106_xboole_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', d1_xboole_0)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t63_xboole_1)).
fof(t138_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t138_zfmisc_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t3_boole)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', l32_xboole_1)).
fof(t79_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k4_xboole_0(X1,X2),X2), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t79_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', d7_xboole_0)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t123_zfmisc_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', l13_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t113_zfmisc_1)).
fof(t118_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>(r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&r1_tarski(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2)))), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t118_zfmisc_1)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t117_zfmisc_1)).
fof(t122_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k3_xboole_0(X1,X2))=k3_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t122_zfmisc_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t2_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t2_boole)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/93349d1f0cb64a98bef7ef8e9a1a58b2.E_input', t65_xboole_1)).
fof(c_0_22, plain, ![X44, X45, X47, X48, X49]:((r1_xboole_0(X44,X45)|r2_hidden(esk5_2(X44,X45),k3_xboole_0(X44,X45)))&(~r2_hidden(X49,k3_xboole_0(X47,X48))|~r1_xboole_0(X47,X48))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_23, plain, ![X17, X18]:k3_xboole_0(X17,X18)=k3_xboole_0(X18,X17), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_24, plain, ![X39, X40]:r1_tarski(X39,k2_xboole_0(X39,X40)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_25, plain, ![X19]:k2_xboole_0(X19,k1_xboole_0)=X19, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_26, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_28, plain, ![X50, X51]:(~r1_tarski(X50,X51)|k3_xboole_0(X50,X51)=X50), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_29, plain, ![X29, X30, X31]:((r1_tarski(X29,X30)|~r1_tarski(X29,k4_xboole_0(X30,X31)))&(r1_xboole_0(X29,X31)|~r1_tarski(X29,k4_xboole_0(X30,X31)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
cnf(c_0_30, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_34, plain, ![X57, X58, X59]:((~v1_xboole_0(X57)|~r2_hidden(X58,X57))&(r2_hidden(esk6_1(X59),X59)|v1_xboole_0(X59))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_35, plain, ![X54, X55, X56]:(~r1_tarski(X54,X55)|~r1_xboole_0(X55,X56)|r1_xboole_0(X54,X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X3,X2)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, plain, (r2_hidden(esk6_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_42, negated_conjecture, ~(![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4))))), inference(assume_negation,[status(cth)],[t138_zfmisc_1])).
cnf(c_0_43, plain, (v1_xboole_0(X1)|~r1_xboole_0(X2,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X3)|~r1_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_45, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))&(k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0&(~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_42])])])).
fof(c_0_46, plain, ![X43]:k4_xboole_0(X43,k1_xboole_0)=X43, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_47, plain, ![X41, X42]:((k4_xboole_0(X41,X42)!=k1_xboole_0|r1_tarski(X41,X42))&(~r1_tarski(X41,X42)|k4_xboole_0(X41,X42)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_48, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_26, c_0_33])).
fof(c_0_49, plain, ![X32, X33]:r1_xboole_0(k4_xboole_0(X32,X33),X33), inference(variable_rename,[status(thm)],[t79_xboole_1])).
cnf(c_0_50, plain, (v1_xboole_0(X1)|~r1_xboole_0(X2,X1)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_52, plain, ![X37, X38]:((~r1_xboole_0(X37,X38)|k3_xboole_0(X37,X38)=k1_xboole_0)&(k3_xboole_0(X37,X38)!=k1_xboole_0|r1_xboole_0(X37,X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_53, plain, ![X20, X21, X22, X23]:k2_zfmisc_1(k3_xboole_0(X20,X21),k3_xboole_0(X22,X23))=k3_xboole_0(k2_zfmisc_1(X20,X22),k2_zfmisc_1(X21,X23)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_54, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (v1_xboole_0(X1)|~r1_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_48, c_0_39])).
cnf(c_0_57, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_58, plain, ![X52]:(~v1_xboole_0(X52)|X52=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_59, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_26, c_0_39])).
fof(c_0_60, plain, ![X27, X28]:((k2_zfmisc_1(X27,X28)!=k1_xboole_0|(X27=k1_xboole_0|X28=k1_xboole_0))&((X27!=k1_xboole_0|k2_zfmisc_1(X27,X28)=k1_xboole_0)&(X28!=k1_xboole_0|k2_zfmisc_1(X27,X28)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_61, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~r1_xboole_0(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_50, c_0_37])).
cnf(c_0_62, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),X1)|~r1_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X1)), inference(spm,[status(thm)],[c_0_40, c_0_51])).
cnf(c_0_63, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (k1_xboole_0=X1|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_66, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~r1_tarski(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_67, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_59, c_0_27])).
cnf(c_0_69, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_70, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (v1_xboole_0(X1)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_61, c_0_54])).
cnf(c_0_72, negated_conjecture, (r1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),X1)|k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_73, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X3,X2),k2_zfmisc_1(X1,X4))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_27]), c_0_64])).
cnf(c_0_74, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_27]), c_0_64])).
cnf(c_0_75, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_76, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_65, c_0_41])).
cnf(c_0_77, plain, (v1_xboole_0(k4_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_66, c_0_41])).
fof(c_0_78, plain, ![X14, X15, X16]:((r1_tarski(k2_zfmisc_1(X14,X16),k2_zfmisc_1(X15,X16))|~r1_tarski(X14,X15))&(r1_tarski(k2_zfmisc_1(X16,X14),k2_zfmisc_1(X16,X15))|~r1_tarski(X14,X15))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t118_zfmisc_1])])])).
cnf(c_0_79, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_80, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_81, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_57]), c_0_27])).
cnf(c_0_82, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_83, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))|k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_74])).
cnf(c_0_84, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_85, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_86, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_67, c_0_77])).
cnf(c_0_87, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_88, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_79])).
fof(c_0_89, plain, ![X34, X35, X36]:((~r1_tarski(k2_zfmisc_1(X35,X34),k2_zfmisc_1(X36,X34))|r1_tarski(X35,X36)|X34=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X34,X35),k2_zfmisc_1(X34,X36))|r1_tarski(X35,X36)|X34=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
fof(c_0_90, plain, ![X9, X10, X11]:(k2_zfmisc_1(k3_xboole_0(X9,X10),X11)=k3_xboole_0(k2_zfmisc_1(X9,X11),k2_zfmisc_1(X10,X11))&k2_zfmisc_1(X11,k3_xboole_0(X9,X10))=k3_xboole_0(k2_zfmisc_1(X11,X9),k2_zfmisc_1(X11,X10))), inference(variable_rename,[status(thm)],[t122_zfmisc_1])).
cnf(c_0_91, negated_conjecture, (k3_xboole_0(X1,k2_zfmisc_1(esk1_0,esk2_0))=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_72])).
cnf(c_0_92, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k4_xboole_0(X3,X1),X4))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_81]), c_0_82])).
cnf(c_0_93, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_83]), c_0_84])).
cnf(c_0_94, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_69, c_0_85])).
cnf(c_0_95, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_36, c_0_86])).
cnf(c_0_96, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_97, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_98, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),X3)=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_99, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,X1),k2_zfmisc_1(k4_xboole_0(X2,esk3_0),esk2_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_73])).
fof(c_0_100, plain, ![X13]:r1_tarski(k1_xboole_0,X13), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_101, negated_conjecture, (~r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_102, plain, (r1_tarski(k2_zfmisc_1(X1,X2),X3)|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
fof(c_0_103, plain, ![X12]:k3_xboole_0(X12,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_104, plain, ![X53]:r1_xboole_0(X53,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_105, plain, (X1=k1_xboole_0|r1_tarski(X2,k1_xboole_0)|~r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_97, c_0_88])).
cnf(c_0_106, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0)),esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_107, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_108, negated_conjecture, (~r1_tarski(esk2_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_109, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_110, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_111, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_112, negated_conjecture, (k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])]), c_0_108])).
cnf(c_0_113, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_109]), c_0_110])])).
cnf(c_0_114, negated_conjecture, (r1_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113])).
cnf(c_0_115, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk1_0,esk3_0))), inference(spm,[status(thm)],[c_0_61, c_0_114])).
cnf(c_0_116, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_117, negated_conjecture, (k4_xboole_0(esk1_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_67, c_0_115])).
cnf(c_0_118, negated_conjecture, (r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_119, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_120, negated_conjecture, (r1_xboole_0(esk1_0,X1)|~r1_xboole_0(esk3_0,X1)), inference(spm,[status(thm)],[c_0_40, c_0_118])).
cnf(c_0_121, plain, (k3_xboole_0(X1,X2)=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_33, c_0_27])).
cnf(c_0_122, plain, (r1_xboole_0(X1,X2)|~v1_xboole_0(k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_119, c_0_111])).
cnf(c_0_123, negated_conjecture, (r1_xboole_0(esk1_0,X1)|k3_xboole_0(esk3_0,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_120, c_0_63])).
cnf(c_0_124, negated_conjecture, (k4_xboole_0(X1,esk3_0)=k1_xboole_0|~r1_tarski(k4_xboole_0(X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_121, c_0_112])).
cnf(c_0_125, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_36, c_0_41])).
cnf(c_0_126, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),X1)=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_69, c_0_72])).
cnf(c_0_127, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k2_zfmisc_1(X1,k3_xboole_0(X2,X4))|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_64, c_0_33])).
cnf(c_0_128, plain, (r1_xboole_0(X1,X2)|~v1_xboole_0(X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_122, c_0_33])).
cnf(c_0_129, negated_conjecture, (v1_xboole_0(esk1_0)|k3_xboole_0(esk1_0,esk3_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_123]), c_0_27])).
cnf(c_0_130, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk1_0,X1),esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_131, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),X1)=k1_xboole_0|k3_xboole_0(X1,k2_zfmisc_1(esk3_0,esk4_0))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_126, c_0_27])).
cnf(c_0_132, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,esk4_0))!=k1_xboole_0|~r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_93, c_0_127])).
cnf(c_0_133, negated_conjecture, (r1_xboole_0(esk1_0,X1)|k3_xboole_0(esk1_0,esk3_0)!=k1_xboole_0|~r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_128, c_0_129])).
cnf(c_0_134, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~r1_xboole_0(X1,X3)|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_56, c_0_44])).
cnf(c_0_135, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,X1),esk3_0)), inference(spm,[status(thm)],[c_0_116, c_0_130])).
cnf(c_0_136, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(X1,esk4_0))=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_131, c_0_98])).
cnf(c_0_137, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,esk4_0))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_132, c_0_118])])).
cnf(c_0_138, negated_conjecture, (k3_xboole_0(esk1_0,X1)=k1_xboole_0|k3_xboole_0(esk1_0,esk3_0)!=k1_xboole_0|~r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_69, c_0_133])).
cnf(c_0_139, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk1_0,X1))|~r1_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_134, c_0_135])).
cnf(c_0_140, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)!=k1_xboole_0|~r1_tarski(esk1_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_136]), c_0_137])).
cnf(c_0_141, negated_conjecture, (k3_xboole_0(esk1_0,X1)=k1_xboole_0|esk1_0!=k1_xboole_0|~r1_tarski(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_33]), c_0_118])])).
cnf(c_0_142, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_143, negated_conjecture, (v1_xboole_0(esk1_0)|~r1_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_139, c_0_54])).
cnf(c_0_144, negated_conjecture, (esk1_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_82]), c_0_37]), c_0_118])])).
cnf(c_0_145, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|k3_xboole_0(X3,X4)=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_142, c_0_64])).
cnf(c_0_146, plain, (k2_zfmisc_1(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_147, negated_conjecture, (~r1_xboole_0(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_143]), c_0_144])).
cnf(c_0_148, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|k3_xboole_0(X3,X4)=k1_xboole_0|k3_xboole_0(k2_zfmisc_1(X4,X1),k2_zfmisc_1(X3,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_145, c_0_73])).
cnf(c_0_149, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,X1))=k1_xboole_0|k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_126, c_0_146])).
cnf(c_0_150, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_147, c_0_63])).
cnf(c_0_151, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,X1))!=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_27]), c_0_150])).
cnf(c_0_152, negated_conjecture, (k3_xboole_0(esk2_0,k4_xboole_0(X1,esk4_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_81]), c_0_88])])).
cnf(c_0_153, negated_conjecture, (r1_xboole_0(esk2_0,k4_xboole_0(X1,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_152]), c_0_113])).
cnf(c_0_154, negated_conjecture, (v1_xboole_0(k4_xboole_0(esk2_0,esk4_0))), inference(spm,[status(thm)],[c_0_61, c_0_153])).
cnf(c_0_155, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_156, negated_conjecture, (k4_xboole_0(esk2_0,esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_67, c_0_154])).
cnf(c_0_157, negated_conjecture, (~r1_tarski(esk2_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_155, c_0_118])])).
cnf(c_0_158, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_156]), c_0_157]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 159
# Proof object clause steps            : 114
# Proof object formula steps           : 45
# Proof object conjectures             : 48
# Proof object clause conjectures      : 45
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 78
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 35
# Processed clauses                    : 7635
# ...of these trivial                  : 108
# ...subsumed                          : 6378
# ...remaining for further processing  : 1149
# Other redundant clauses eliminated   : 38
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 46
# Backward-rewritten                   : 25
# Generated clauses                    : 165920
# ...of the previous two non-trivial   : 148303
# Contextual simplify-reflections      : 3
# Paramodulations                      : 165882
# Factorizations                       : 0
# Equation resolutions                 : 38
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1041
#    Positive orientable unit clauses  : 97
#    Positive unorientable unit clauses: 17
#    Negative unit clauses             : 47
#    Non-unit-clauses                  : 880
# Current number of unprocessed clauses: 140431
# ...number of literals in the above   : 352709
# Current number of archived formulas  : 0
# Current number of archived clauses   : 106
# Clause-clause subsumption calls (NU) : 77311
# Rec. Clause-clause subsumption calls : 53568
# Non-unit clause-clause subsumptions  : 2369
# Unit Clause-clause subsumption calls : 6780
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 525
# BW rewrite match successes           : 308
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3033865

# -------------------------------------------------
# User time                : 1.613 s
# System time              : 0.020 s
# Total time               : 1.633 s
# Maximum resident set size: 117204 pages
