# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t34_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2))))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', t34_lattice3)).
fof(t45_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k15_lattice3(X1,a_2_3_lattice3(X1,X2))&X2=k16_lattice3(X1,a_2_4_lattice3(X1,X2))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', t45_lattice3)).
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', d16_lattice3)).
fof(fraenkel_a_2_4_lattice3, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))&m1_subset_1(X3,u1_struct_0(X2)))=>(r2_hidden(X1,a_2_4_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattices(X2,X3,X4))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', fraenkel_a_2_4_lattice3)).
fof(t50_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', t50_lattice3)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', t2_xboole_1)).
fof(t21_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattices(X1,X2,X3)<=>k2_lattices(X1,X2,X3)=X2)))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', t21_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', cc1_lattices)).
fof(t48_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r3_lattices(X1,X4,X5)&r2_hidden(X5,X3)))))))=>r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3)))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', t48_lattice3)).
fof(fraenkel_a_2_5_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_5_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r3_lattices(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', fraenkel_a_2_5_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', dt_k15_lattice3)).
fof(t47_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))&k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2)))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', t47_lattice3)).
fof(d16_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k5_lattices(X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2)))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', d16_lattices)).
fof(dt_k5_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>m1_subset_1(k5_lattices(X1),u1_struct_0(X1))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', dt_k5_lattices)).
fof(d6_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v6_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', d6_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', dt_l3_lattices)).
fof(d13_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l1_lattices(X1))=>(v13_lattices(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k2_lattices(X1,X2,X3)=X2&k2_lattices(X1,X3,X2)=X2))))), file('data/MPTP/proofs/937fd81d90344375b1f948d1e843f41f.E_input', d13_lattices)).
fof(c_0_18, plain, ![X6, X7, X8, X9, X10]:(((r3_lattice3(X6,X7,X8)|X7!=k16_lattice3(X6,X8)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&(~m1_subset_1(X9,u1_struct_0(X6))|(~r3_lattice3(X6,X9,X8)|r3_lattices(X6,X9,X7))|X7!=k16_lattice3(X6,X8)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6))))))&((m1_subset_1(esk1_3(X6,X7,X10),u1_struct_0(X6))|~r3_lattice3(X6,X7,X10)|X7=k16_lattice3(X6,X10)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&((r3_lattice3(X6,esk1_3(X6,X7,X10),X10)|~r3_lattice3(X6,X7,X10)|X7=k16_lattice3(X6,X10)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&(~r3_lattices(X6,esk1_3(X6,X7,X10),X7)|~r3_lattice3(X6,X7,X10)|X7=k16_lattice3(X6,X10)|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])).
cnf(c_0_19, plain, (r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|X2!=k16_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_20, plain, ![X12, X13]:((X13=k15_lattice3(X12,a_2_3_lattice3(X12,X13))|~m1_subset_1(X13,u1_struct_0(X12))|(v2_struct_0(X12)|(~v10_lattices(X12)|(~v4_lattice3(X12)|~l3_lattices(X12)))))&(X13=k16_lattice3(X12,a_2_4_lattice3(X12,X13))|~m1_subset_1(X13,u1_struct_0(X12))|(v2_struct_0(X12)|(~v10_lattices(X12)|(~v4_lattice3(X12)|~l3_lattices(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t45_lattice3])])])])])).
fof(c_0_21, plain, ![X14, X15, X16, X17, X18]:((~r3_lattice3(X14,X15,X16)|(~m1_subset_1(X17,u1_struct_0(X14))|(~r2_hidden(X17,X16)|r1_lattices(X14,X15,X17)))|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l3_lattices(X14)))&((m1_subset_1(esk2_3(X14,X15,X18),u1_struct_0(X14))|r3_lattice3(X14,X15,X18)|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l3_lattices(X14)))&((r2_hidden(esk2_3(X14,X15,X18),X18)|r3_lattice3(X14,X15,X18)|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l3_lattices(X14)))&(~r1_lattices(X14,X15,esk2_3(X14,X15,X18))|r3_lattice3(X14,X15,X18)|~m1_subset_1(X15,u1_struct_0(X14))|(v2_struct_0(X14)|~l3_lattices(X14)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
cnf(c_0_22, plain, (r3_lattice3(X1,k16_lattice3(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (X1=k16_lattice3(X2,a_2_4_lattice3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, plain, ![X20, X21, X22, X24]:(((m1_subset_1(esk3_3(X20,X21,X22),u1_struct_0(X21))|~r2_hidden(X20,a_2_4_lattice3(X21,X22))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21)))|~m1_subset_1(X22,u1_struct_0(X21))))&((X20=esk3_3(X20,X21,X22)|~r2_hidden(X20,a_2_4_lattice3(X21,X22))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21)))|~m1_subset_1(X22,u1_struct_0(X21))))&(r3_lattices(X21,X22,esk3_3(X20,X21,X22))|~r2_hidden(X20,a_2_4_lattice3(X21,X22))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21)))|~m1_subset_1(X22,u1_struct_0(X21))))))&(~m1_subset_1(X24,u1_struct_0(X21))|(X20!=X24|~r3_lattices(X21,X22,X24))|r2_hidden(X20,a_2_4_lattice3(X21,X22))|(v2_struct_0(X21)|(~v10_lattices(X21)|(~v4_lattice3(X21)|~l3_lattices(X21)))|~m1_subset_1(X22,u1_struct_0(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_4_lattice3])])])])])])).
cnf(c_0_25, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, plain, (r3_lattice3(X1,X2,a_2_4_lattice3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27, plain, (r2_hidden(X3,a_2_4_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattices(X2,X4,X1)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_28, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v13_lattices(X1)&l3_lattices(X1))))&k5_lattices(X1)=k15_lattice3(X1,k1_xboole_0)))), inference(assume_negation,[status(cth)],[t50_lattice3])).
fof(c_0_29, plain, ![X26, X27]:(~r2_hidden(X26,X27)|~r1_tarski(X27,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_30, plain, ![X28]:r1_tarski(k1_xboole_0,X28), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_31, plain, ![X29, X30, X31]:((~r1_lattices(X29,X30,X31)|k2_lattices(X29,X30,X31)=X30|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v8_lattices(X29)|(~v9_lattices(X29)|~l3_lattices(X29)))))&(k2_lattices(X29,X30,X31)!=X30|r1_lattices(X29,X30,X31)|~m1_subset_1(X31,u1_struct_0(X29))|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|(~v8_lattices(X29)|(~v9_lattices(X29)|~l3_lattices(X29)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_lattices])])])])])).
cnf(c_0_32, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r2_hidden(X3,a_2_4_lattice3(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33, plain, (r2_hidden(X1,a_2_4_lattice3(X2,X3))|v2_struct_0(X2)|~r3_lattices(X2,X3,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_27])).
fof(c_0_34, plain, ![X32]:((~v2_struct_0(X32)|(v2_struct_0(X32)|~v10_lattices(X32))|~l3_lattices(X32))&((v4_lattices(X32)|(v2_struct_0(X32)|~v10_lattices(X32))|~l3_lattices(X32))&((v5_lattices(X32)|(v2_struct_0(X32)|~v10_lattices(X32))|~l3_lattices(X32))&((v6_lattices(X32)|(v2_struct_0(X32)|~v10_lattices(X32))|~l3_lattices(X32))&((v7_lattices(X32)|(v2_struct_0(X32)|~v10_lattices(X32))|~l3_lattices(X32))&((v8_lattices(X32)|(v2_struct_0(X32)|~v10_lattices(X32))|~l3_lattices(X32))&(v9_lattices(X32)|(v2_struct_0(X32)|~v10_lattices(X32))|~l3_lattices(X32)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_35, plain, ![X33, X34, X35, X37]:((m1_subset_1(esk5_3(X33,X34,X35),u1_struct_0(X33))|r3_lattices(X33,k15_lattice3(X33,X34),k15_lattice3(X33,X35))|(v2_struct_0(X33)|(~v10_lattices(X33)|(~v4_lattice3(X33)|~l3_lattices(X33)))))&((r2_hidden(esk5_3(X33,X34,X35),X34)|r3_lattices(X33,k15_lattice3(X33,X34),k15_lattice3(X33,X35))|(v2_struct_0(X33)|(~v10_lattices(X33)|(~v4_lattice3(X33)|~l3_lattices(X33)))))&(~m1_subset_1(X37,u1_struct_0(X33))|(~r3_lattices(X33,esk5_3(X33,X34,X35),X37)|~r2_hidden(X37,X35))|r3_lattices(X33,k15_lattice3(X33,X34),k15_lattice3(X33,X35))|(v2_struct_0(X33)|(~v10_lattices(X33)|(~v4_lattice3(X33)|~l3_lattices(X33))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_lattice3])])])])])])).
fof(c_0_36, negated_conjecture, ((~v2_struct_0(esk4_0)&(v10_lattices(esk4_0)&(v4_lattice3(esk4_0)&l3_lattices(esk4_0))))&(v2_struct_0(esk4_0)|(~v10_lattices(esk4_0)|(~v13_lattices(esk4_0)|~l3_lattices(esk4_0)))|k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_28])])])])).
cnf(c_0_37, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_39, plain, ![X38, X39, X40, X43, X44, X45]:(((m1_subset_1(esk6_3(X38,X39,X40),u1_struct_0(X39))|~r2_hidden(X38,a_2_5_lattice3(X39,X40))|(v2_struct_0(X39)|(~v10_lattices(X39)|(~v4_lattice3(X39)|~l3_lattices(X39)))))&((X38=esk6_3(X38,X39,X40)|~r2_hidden(X38,a_2_5_lattice3(X39,X40))|(v2_struct_0(X39)|(~v10_lattices(X39)|(~v4_lattice3(X39)|~l3_lattices(X39)))))&((m1_subset_1(esk7_3(X38,X39,X40),u1_struct_0(X39))|~r2_hidden(X38,a_2_5_lattice3(X39,X40))|(v2_struct_0(X39)|(~v10_lattices(X39)|(~v4_lattice3(X39)|~l3_lattices(X39)))))&((r3_lattices(X39,esk6_3(X38,X39,X40),esk7_3(X38,X39,X40))|~r2_hidden(X38,a_2_5_lattice3(X39,X40))|(v2_struct_0(X39)|(~v10_lattices(X39)|(~v4_lattice3(X39)|~l3_lattices(X39)))))&(r2_hidden(esk7_3(X38,X39,X40),X40)|~r2_hidden(X38,a_2_5_lattice3(X39,X40))|(v2_struct_0(X39)|(~v10_lattices(X39)|(~v4_lattice3(X39)|~l3_lattices(X39)))))))))&(~m1_subset_1(X44,u1_struct_0(X39))|(X38!=X44|(~m1_subset_1(X45,u1_struct_0(X39))|(~r3_lattices(X39,X44,X45)|~r2_hidden(X45,X43))))|r2_hidden(X38,a_2_5_lattice3(X39,X43))|(v2_struct_0(X39)|(~v10_lattices(X39)|(~v4_lattice3(X39)|~l3_lattices(X39)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_5_lattice3])])])])])])])).
cnf(c_0_40, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, negated_conjecture, (l3_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, negated_conjecture, (v4_lattice3(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, negated_conjecture, (v10_lattices(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_50, plain, (r2_hidden(esk7_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (k2_lattices(X1,X2,X3)=X2|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk5_3(esk4_0,X1,X2),X1)|r3_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,X2))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
fof(c_0_53, plain, ![X46, X47]:(v2_struct_0(X46)|~l3_lattices(X46)|m1_subset_1(k15_lattice3(X46,X47),u1_struct_0(X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_54, plain, (v2_struct_0(X1)|~r2_hidden(X2,a_2_5_lattice3(X1,k1_xboole_0))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_55, plain, ![X48, X49]:((k15_lattice3(X48,X49)=k15_lattice3(X48,a_2_5_lattice3(X48,X49))|(v2_struct_0(X48)|(~v10_lattices(X48)|(~v4_lattice3(X48)|~l3_lattices(X48)))))&(k16_lattice3(X48,X49)=k16_lattice3(X48,a_2_6_lattice3(X48,X49))|(v2_struct_0(X48)|(~v10_lattices(X48)|(~v4_lattice3(X48)|~l3_lattices(X48)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_lattice3])])])])])).
cnf(c_0_56, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,X2))=k15_lattice3(esk4_0,X1)|r2_hidden(esk5_3(esk4_0,X1,X2),X1)|~m1_subset_1(k15_lattice3(esk4_0,X2),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,X1),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_57, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_58, negated_conjecture, (r3_lattices(esk4_0,k15_lattice3(esk4_0,a_2_5_lattice3(X1,k1_xboole_0)),k15_lattice3(esk4_0,X2))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_54, c_0_52])).
cnf(c_0_59, plain, (k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,X2))=k15_lattice3(esk4_0,X1)|r2_hidden(esk5_3(esk4_0,X1,X2),X1)|~m1_subset_1(k15_lattice3(esk4_0,X1),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_45])]), c_0_48])).
fof(c_0_61, plain, ![X50, X51, X52]:(((k2_lattices(X50,X51,X52)=X51|~m1_subset_1(X52,u1_struct_0(X50))|X51!=k5_lattices(X50)|~m1_subset_1(X51,u1_struct_0(X50))|~v13_lattices(X50)|(v2_struct_0(X50)|~l1_lattices(X50)))&(k2_lattices(X50,X52,X51)=X51|~m1_subset_1(X52,u1_struct_0(X50))|X51!=k5_lattices(X50)|~m1_subset_1(X51,u1_struct_0(X50))|~v13_lattices(X50)|(v2_struct_0(X50)|~l1_lattices(X50))))&((m1_subset_1(esk8_2(X50,X51),u1_struct_0(X50))|X51=k5_lattices(X50)|~m1_subset_1(X51,u1_struct_0(X50))|~v13_lattices(X50)|(v2_struct_0(X50)|~l1_lattices(X50)))&(k2_lattices(X50,X51,esk8_2(X50,X51))!=X51|k2_lattices(X50,esk8_2(X50,X51),X51)!=X51|X51=k5_lattices(X50)|~m1_subset_1(X51,u1_struct_0(X50))|~v13_lattices(X50)|(v2_struct_0(X50)|~l1_lattices(X50))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattices])])])])])])).
fof(c_0_62, plain, ![X54]:(v2_struct_0(X54)|~l1_lattices(X54)|m1_subset_1(k5_lattices(X54),u1_struct_0(X54))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_lattices])])])).
cnf(c_0_63, negated_conjecture, (r3_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),k15_lattice3(esk4_0,X1))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_64, plain, (X1=k15_lattice3(X2,a_2_3_lattice3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_65, plain, ![X55, X56, X57]:((~v6_lattices(X55)|(~m1_subset_1(X56,u1_struct_0(X55))|(~m1_subset_1(X57,u1_struct_0(X55))|k2_lattices(X55,X56,X57)=k2_lattices(X55,X57,X56)))|(v2_struct_0(X55)|~l1_lattices(X55)))&((m1_subset_1(esk9_1(X55),u1_struct_0(X55))|v6_lattices(X55)|(v2_struct_0(X55)|~l1_lattices(X55)))&((m1_subset_1(esk10_1(X55),u1_struct_0(X55))|v6_lattices(X55)|(v2_struct_0(X55)|~l1_lattices(X55)))&(k2_lattices(X55,esk9_1(X55),esk10_1(X55))!=k2_lattices(X55,esk10_1(X55),esk9_1(X55))|v6_lattices(X55)|(v2_struct_0(X55)|~l1_lattices(X55)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_lattices])])])])])])).
fof(c_0_66, plain, ![X60]:((l1_lattices(X60)|~l3_lattices(X60))&(l2_lattices(X60)|~l3_lattices(X60))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_67, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,X1),k15_lattice3(esk4_0,X2))=k15_lattice3(esk4_0,X1)|r2_hidden(esk5_3(esk4_0,X1,X2),X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_57]), c_0_45])]), c_0_48])).
cnf(c_0_68, plain, (k2_lattices(X1,X2,X3)=X3|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|X3!=k5_lattices(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v13_lattices(X1)|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_69, plain, (v2_struct_0(X1)|m1_subset_1(k5_lattices(X1),u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_70, negated_conjecture, (r3_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_71, negated_conjecture, (v2_struct_0(esk4_0)|~v10_lattices(esk4_0)|~v13_lattices(esk4_0)|~l3_lattices(esk4_0)|k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_72, plain, ![X61, X63, X64]:(((m1_subset_1(esk11_1(X61),u1_struct_0(X61))|~v13_lattices(X61)|(v2_struct_0(X61)|~l1_lattices(X61)))&((k2_lattices(X61,esk11_1(X61),X63)=esk11_1(X61)|~m1_subset_1(X63,u1_struct_0(X61))|~v13_lattices(X61)|(v2_struct_0(X61)|~l1_lattices(X61)))&(k2_lattices(X61,X63,esk11_1(X61))=esk11_1(X61)|~m1_subset_1(X63,u1_struct_0(X61))|~v13_lattices(X61)|(v2_struct_0(X61)|~l1_lattices(X61)))))&((m1_subset_1(esk12_2(X61,X64),u1_struct_0(X61))|~m1_subset_1(X64,u1_struct_0(X61))|v13_lattices(X61)|(v2_struct_0(X61)|~l1_lattices(X61)))&(k2_lattices(X61,X64,esk12_2(X61,X64))!=X64|k2_lattices(X61,esk12_2(X61,X64),X64)!=X64|~m1_subset_1(X64,u1_struct_0(X61))|v13_lattices(X61)|(v2_struct_0(X61)|~l1_lattices(X61))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_lattices])])])])])])).
cnf(c_0_73, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~v6_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_74, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_75, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_76, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,a_2_5_lattice3(X1,k1_xboole_0)),k15_lattice3(esk4_0,X2))=k15_lattice3(esk4_0,a_2_5_lattice3(X1,k1_xboole_0))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_54, c_0_67])).
cnf(c_0_77, plain, (k2_lattices(X1,X2,k5_lattices(X1))=k5_lattices(X1)|v2_struct_0(X1)|~l1_lattices(X1)|~v13_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_68]), c_0_69])).
cnf(c_0_78, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)=k15_lattice3(esk4_0,k1_xboole_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_70]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_79, negated_conjecture, (k5_lattices(esk4_0)!=k15_lattice3(esk4_0,k1_xboole_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_47]), c_0_45])]), c_0_48])).
cnf(c_0_80, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk12_2(X1,X2))!=X2|k2_lattices(X1,esk12_2(X1,X2),X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_81, plain, (k2_lattices(X1,X2,X3)=k2_lattices(X1,X3,X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])).
cnf(c_0_82, plain, (m1_subset_1(esk12_2(X1,X2),u1_struct_0(X1))|v13_lattices(X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_83, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),k15_lattice3(esk4_0,X1))=k15_lattice3(esk4_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_59]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_84, negated_conjecture, (~l1_lattices(esk4_0)|~v13_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))|~m1_subset_1(k5_lattices(esk4_0),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_48]), c_0_79])).
cnf(c_0_85, plain, (v13_lattices(X1)|v2_struct_0(X1)|k2_lattices(X1,X2,esk12_2(X1,X2))!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82]), c_0_75])).
cnf(c_0_86, negated_conjecture, (k2_lattices(esk4_0,k15_lattice3(esk4_0,k1_xboole_0),X1)=k15_lattice3(esk4_0,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_64]), c_0_45]), c_0_46]), c_0_47])]), c_0_48])).
cnf(c_0_87, negated_conjecture, (~l1_lattices(esk4_0)|~v13_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_69]), c_0_48])).
cnf(c_0_88, negated_conjecture, (v13_lattices(esk4_0)|~m1_subset_1(esk12_2(esk4_0,k15_lattice3(esk4_0,k1_xboole_0)),u1_struct_0(esk4_0))|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_45]), c_0_47])]), c_0_48])).
cnf(c_0_89, negated_conjecture, (~l1_lattices(esk4_0)|~v13_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_57]), c_0_45])]), c_0_48])).
cnf(c_0_90, negated_conjecture, (~l1_lattices(esk4_0)|~m1_subset_1(k15_lattice3(esk4_0,k1_xboole_0),u1_struct_0(esk4_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_82]), c_0_48]), c_0_89])).
cnf(c_0_91, negated_conjecture, (~l1_lattices(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_57]), c_0_45])]), c_0_48])).
cnf(c_0_92, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_75]), c_0_45])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 93
# Proof object clause steps            : 56
# Proof object formula steps           : 37
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 70
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 59
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 58
# Processed clauses                    : 320
# ...of these trivial                  : 1
# ...subsumed                          : 72
# ...remaining for further processing  : 247
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 0
# Generated clauses                    : 602
# ...of the previous two non-trivial   : 570
# Contextual simplify-reflections      : 24
# Paramodulations                      : 596
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 168
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 158
# Current number of unprocessed clauses: 365
# ...number of literals in the above   : 2686
# Current number of archived formulas  : 0
# Current number of archived clauses   : 73
# Clause-clause subsumption calls (NU) : 11860
# Rec. Clause-clause subsumption calls : 1016
# Non-unit clause-clause subsumptions  : 108
# Unit Clause-clause subsumption calls : 78
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 26771

# -------------------------------------------------
# User time                : 0.026 s
# System time              : 0.000 s
# Total time               : 0.026 s
# Maximum resident set size: 31508 pages
