# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', d4_xboole_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', d1_tarski)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', d10_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', t28_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', t2_xboole_1)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', t123_zfmisc_1)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', t134_zfmisc_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', l13_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', t113_zfmisc_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', t7_tarski)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', t2_boole)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', t6_xboole_0)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/94564aa5b30f45a7bf2a8b6471349bf5.E_input', d8_xboole_0)).
fof(c_0_17, plain, ![X7, X8, X9, X10, X11, X12, X13, X14]:((((r2_hidden(X10,X7)|~r2_hidden(X10,X9)|X9!=k3_xboole_0(X7,X8))&(r2_hidden(X10,X8)|~r2_hidden(X10,X9)|X9!=k3_xboole_0(X7,X8)))&(~r2_hidden(X11,X7)|~r2_hidden(X11,X8)|r2_hidden(X11,X9)|X9!=k3_xboole_0(X7,X8)))&((~r2_hidden(esk1_3(X12,X13,X14),X14)|(~r2_hidden(esk1_3(X12,X13,X14),X12)|~r2_hidden(esk1_3(X12,X13,X14),X13))|X14=k3_xboole_0(X12,X13))&((r2_hidden(esk1_3(X12,X13,X14),X12)|r2_hidden(esk1_3(X12,X13,X14),X14)|X14=k3_xboole_0(X12,X13))&(r2_hidden(esk1_3(X12,X13,X14),X13)|r2_hidden(esk1_3(X12,X13,X14),X14)|X14=k3_xboole_0(X12,X13))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_18, plain, ![X18, X19, X20, X21, X22, X23]:(((~r2_hidden(X20,X19)|X20=X18|X19!=k1_tarski(X18))&(X21!=X18|r2_hidden(X21,X19)|X19!=k1_tarski(X18)))&((~r2_hidden(esk2_2(X22,X23),X23)|esk2_2(X22,X23)!=X22|X23=k1_tarski(X22))&(r2_hidden(esk2_2(X22,X23),X23)|esk2_2(X22,X23)=X22|X23=k1_tarski(X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_19, plain, ![X39, X40, X41, X42, X43, X44, X45, X46]:((((r2_hidden(X42,X39)|~r2_hidden(X42,X41)|X41!=k4_xboole_0(X39,X40))&(~r2_hidden(X42,X40)|~r2_hidden(X42,X41)|X41!=k4_xboole_0(X39,X40)))&(~r2_hidden(X43,X39)|r2_hidden(X43,X40)|r2_hidden(X43,X41)|X41!=k4_xboole_0(X39,X40)))&((~r2_hidden(esk8_3(X44,X45,X46),X46)|(~r2_hidden(esk8_3(X44,X45,X46),X44)|r2_hidden(esk8_3(X44,X45,X46),X45))|X46=k4_xboole_0(X44,X45))&((r2_hidden(esk8_3(X44,X45,X46),X44)|r2_hidden(esk8_3(X44,X45,X46),X46)|X46=k4_xboole_0(X44,X45))&(~r2_hidden(esk8_3(X44,X45,X46),X45)|r2_hidden(esk8_3(X44,X45,X46),X46)|X46=k4_xboole_0(X44,X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, plain, ![X48, X49, X50]:((~v1_xboole_0(X48)|~r2_hidden(X49,X48))&(r2_hidden(esk9_1(X50),X50)|v1_xboole_0(X50))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_22, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X25, X26, X27, X28, X29]:((~r1_tarski(X25,X26)|(~r2_hidden(X27,X25)|r2_hidden(X27,X26)))&((r2_hidden(esk3_2(X28,X29),X28)|r1_tarski(X28,X29))&(~r2_hidden(esk3_2(X28,X29),X29)|r1_tarski(X28,X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_24, plain, ![X5, X6]:(((r1_tarski(X5,X6)|X5!=X6)&(r1_tarski(X6,X5)|X5!=X6))&(~r1_tarski(X5,X6)|~r1_tarski(X6,X5)|X5=X6)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (r2_hidden(esk9_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X52, X53]:k3_xboole_0(X52,X53)=k3_xboole_0(X53,X52), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_29, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, plain, ![X16, X17]:(~r1_tarski(X16,X17)|k3_xboole_0(X16,X17)=X16), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_32, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk9_1(k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_36, plain, ![X54]:r1_tarski(k1_xboole_0,X54), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_37, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, plain, (esk3_2(k1_tarski(X1),X2)=X1|r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_39, plain, ![X31, X32, X33, X34]:k2_zfmisc_1(k3_xboole_0(X31,X32),k3_xboole_0(X33,X34))=k3_xboole_0(k2_zfmisc_1(X31,X33),k2_zfmisc_1(X32,X34)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_40, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_32])).
fof(c_0_42, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
fof(c_0_43, plain, ![X55]:(~v1_xboole_0(X55)|X55=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_44, plain, (v1_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X3)))|~r2_hidden(esk9_1(k3_xboole_0(X1,k4_xboole_0(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_45, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk9_1(k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_46, plain, ![X56, X57]:((k2_zfmisc_1(X56,X57)!=k1_xboole_0|(X56=k1_xboole_0|X57=k1_xboole_0))&((X56!=k1_xboole_0|k2_zfmisc_1(X56,X57)=k1_xboole_0)&(X57!=k1_xboole_0|k2_zfmisc_1(X56,X57)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_47, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_48, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_50, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_40, c_0_41])).
fof(c_0_52, negated_conjecture, (k2_zfmisc_1(esk4_0,esk5_0)=k2_zfmisc_1(esk6_0,esk7_0)&(esk4_0!=k1_xboole_0&(esk5_0!=k1_xboole_0&(esk4_0!=esk6_0|esk5_0!=esk7_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_42])])])).
cnf(c_0_53, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, plain, (v1_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_55, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_57, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_58, plain, (v1_xboole_0(X1)|r1_tarski(k1_tarski(esk9_1(X1)),X1)), inference(spm,[status(thm)],[c_0_49, c_0_27])).
cnf(c_0_59, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))=k2_zfmisc_1(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_60, negated_conjecture, (k2_zfmisc_1(esk4_0,esk5_0)=k2_zfmisc_1(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (k1_tarski(esk9_1(k1_xboole_0))=k1_xboole_0|v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_65, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_66, plain, ![X58, X59, X61]:((r2_hidden(esk10_2(X58,X59),X59)|~r2_hidden(X58,X59))&(~r2_hidden(X61,X59)|~r2_hidden(X61,esk10_2(X58,X59))|~r2_hidden(X58,X59))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
cnf(c_0_67, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk6_0,esk7_0),k2_zfmisc_1(esk4_0,X1))=k2_zfmisc_1(esk4_0,k3_xboole_0(X1,esk5_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_35])).
cnf(c_0_68, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,k4_xboole_0(X4,X2)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_61]), c_0_62])).
cnf(c_0_69, plain, (r2_hidden(esk3_2(k4_xboole_0(X1,X2),X3),X1)|r1_tarski(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_63, c_0_30])).
cnf(c_0_70, plain, (X1=esk9_1(k1_xboole_0)|~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_64]), c_0_65])).
cnf(c_0_71, plain, (r2_hidden(esk10_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_73, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_74, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_75, negated_conjecture, (k2_zfmisc_1(esk4_0,k3_xboole_0(esk5_0,k4_xboole_0(X1,esk7_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_35])).
cnf(c_0_76, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_77, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_37, c_0_69])).
cnf(c_0_78, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,esk10_2(X3,X2))|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_79, plain, (esk10_2(X1,k1_xboole_0)=esk9_1(k1_xboole_0)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_80, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_72, c_0_48])).
cnf(c_0_81, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_82, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_73])).
cnf(c_0_83, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_84, negated_conjecture, (k3_xboole_0(esk5_0,k4_xboole_0(X1,esk7_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
cnf(c_0_85, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_77]), c_0_35])).
cnf(c_0_86, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X2,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])).
cnf(c_0_87, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk6_0,esk7_0),k2_zfmisc_1(X1,esk5_0))=k2_zfmisc_1(k3_xboole_0(esk4_0,X1),esk5_0)), inference(spm,[status(thm)],[c_0_81, c_0_60])).
cnf(c_0_88, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k4_xboole_0(X3,X1),X4))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_61]), c_0_82])).
cnf(c_0_89, plain, (r2_hidden(esk3_2(k3_xboole_0(X1,X2),X3),X2)|r1_tarski(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_26, c_0_30])).
cnf(c_0_90, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk5_0),k2_zfmisc_1(esk6_0,esk7_0))=k2_zfmisc_1(k3_xboole_0(X1,esk4_0),esk5_0)), inference(spm,[status(thm)],[c_0_81, c_0_60])).
cnf(c_0_91, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_35]), c_0_50])).
cnf(c_0_92, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_83])).
cnf(c_0_93, negated_conjecture, (k4_xboole_0(esk5_0,esk7_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_94, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_79]), c_0_86])).
cnf(c_0_95, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk4_0,k4_xboole_0(X1,esk6_0)),esk5_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_96, negated_conjecture, (esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_97, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_37, c_0_89])).
cnf(c_0_98, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk7_0),k2_zfmisc_1(esk6_0,esk5_0))=k2_zfmisc_1(k3_xboole_0(X1,esk4_0),esk5_0)), inference(rw,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_99, negated_conjecture, (r2_hidden(X1,esk7_0)|~r2_hidden(X1,esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94])).
cnf(c_0_100, negated_conjecture, (k3_xboole_0(esk4_0,k4_xboole_0(X1,esk6_0))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_95]), c_0_96])).
cnf(c_0_101, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,esk4_0),esk5_0),k2_zfmisc_1(esk6_0,esk5_0))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_102, negated_conjecture, (r2_hidden(esk3_2(esk5_0,X1),esk7_0)|r1_tarski(esk5_0,X1)), inference(spm,[status(thm)],[c_0_99, c_0_30])).
cnf(c_0_103, negated_conjecture, (k4_xboole_0(esk4_0,esk6_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_100, c_0_85])).
cnf(c_0_104, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk6_0,esk7_0),k2_zfmisc_1(esk6_0,esk5_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_51]), c_0_60])).
cnf(c_0_105, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk6_0,esk4_0),esk5_0)=k2_zfmisc_1(esk6_0,k3_xboole_0(esk7_0,esk5_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_59]), c_0_35])).
cnf(c_0_106, negated_conjecture, (r1_tarski(esk5_0,esk7_0)), inference(spm,[status(thm)],[c_0_37, c_0_102])).
cnf(c_0_107, negated_conjecture, (r2_hidden(X1,esk6_0)|~r2_hidden(X1,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_103]), c_0_94])).
cnf(c_0_108, negated_conjecture, (k2_zfmisc_1(esk6_0,k3_xboole_0(esk7_0,esk5_0))=k2_zfmisc_1(esk6_0,esk7_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_104]), c_0_87]), c_0_35]), c_0_105])).
cnf(c_0_109, negated_conjecture, (k3_xboole_0(esk7_0,esk5_0)=esk5_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_106]), c_0_35])).
cnf(c_0_110, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk6_0,esk7_0),k2_zfmisc_1(k4_xboole_0(X1,esk4_0),X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_88, c_0_60])).
fof(c_0_111, plain, ![X62]:k3_xboole_0(X62,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_112, negated_conjecture, (r2_hidden(esk3_2(esk4_0,X1),esk6_0)|r1_tarski(esk4_0,X1)), inference(spm,[status(thm)],[c_0_107, c_0_30])).
cnf(c_0_113, negated_conjecture, (k2_zfmisc_1(esk6_0,esk5_0)=k2_zfmisc_1(esk6_0,esk7_0)), inference(rw,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_114, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk6_0,k4_xboole_0(X1,esk4_0)),esk7_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_81, c_0_110])).
cnf(c_0_115, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_116, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_97, c_0_35])).
cnf(c_0_117, negated_conjecture, (r1_tarski(esk4_0,esk6_0)), inference(spm,[status(thm)],[c_0_37, c_0_112])).
cnf(c_0_118, negated_conjecture, (k2_zfmisc_1(esk6_0,k3_xboole_0(esk5_0,X1))=k2_zfmisc_1(esk6_0,k3_xboole_0(esk7_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_113]), c_0_59])).
cnf(c_0_119, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk6_0,esk4_0),esk7_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_114, c_0_85])).
cnf(c_0_120, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_115, c_0_35])).
cnf(c_0_121, plain, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_116, c_0_81])).
cnf(c_0_122, negated_conjecture, (k3_xboole_0(esk6_0,esk4_0)=esk4_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_117]), c_0_35])).
cnf(c_0_123, negated_conjecture, (k2_zfmisc_1(esk6_0,k3_xboole_0(esk7_0,k4_xboole_0(X1,esk5_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_61]), c_0_62])).
cnf(c_0_124, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk6_0,esk4_0),k3_xboole_0(esk7_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_119]), c_0_120])).
cnf(c_0_125, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,X1),k2_zfmisc_1(esk6_0,X1))), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_126, negated_conjecture, (k2_zfmisc_1(esk6_0,k4_xboole_0(esk7_0,esk5_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_123, c_0_85])).
cnf(c_0_127, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk6_0,esk4_0),esk5_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_124, c_0_109])).
cnf(c_0_128, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,k4_xboole_0(esk7_0,esk5_0)),k1_xboole_0)), inference(spm,[status(thm)],[c_0_125, c_0_126])).
fof(c_0_129, plain, ![X63, X64]:((r2_hidden(esk11_2(X63,X64),X64)|~r2_xboole_0(X63,X64))&(~r2_hidden(esk11_2(X63,X64),X63)|~r2_xboole_0(X63,X64))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
cnf(c_0_130, negated_conjecture, (k4_xboole_0(esk6_0,esk4_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_127]), c_0_96])).
cnf(c_0_131, negated_conjecture, (k2_zfmisc_1(esk4_0,k4_xboole_0(esk7_0,esk5_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_128]), c_0_115])).
cnf(c_0_132, plain, (~r2_hidden(esk11_2(X1,X2),X1)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_133, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_130]), c_0_94])).
cnf(c_0_134, negated_conjecture, (k4_xboole_0(esk7_0,esk5_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_131]), c_0_76])).
cnf(c_0_135, negated_conjecture, (~r2_xboole_0(esk4_0,X1)|~r2_hidden(esk11_2(esk4_0,X1),esk6_0)), inference(spm,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_136, plain, (r2_hidden(esk11_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
fof(c_0_137, plain, ![X66, X67]:(((r1_tarski(X66,X67)|~r2_xboole_0(X66,X67))&(X66!=X67|~r2_xboole_0(X66,X67)))&(~r1_tarski(X66,X67)|X66=X67|r2_xboole_0(X66,X67))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
cnf(c_0_138, negated_conjecture, (r2_hidden(X1,esk5_0)|~r2_hidden(X1,esk7_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_134]), c_0_94])).
cnf(c_0_139, negated_conjecture, (~r2_xboole_0(esk4_0,esk6_0)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_140, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_141, negated_conjecture, (~r2_xboole_0(esk5_0,X1)|~r2_hidden(esk11_2(esk5_0,X1),esk7_0)), inference(spm,[status(thm)],[c_0_132, c_0_138])).
cnf(c_0_142, negated_conjecture, (esk4_0!=esk6_0|esk5_0!=esk7_0), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_143, negated_conjecture, (esk4_0=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_117])])).
cnf(c_0_144, negated_conjecture, (~r2_xboole_0(esk5_0,esk7_0)), inference(spm,[status(thm)],[c_0_141, c_0_136])).
cnf(c_0_145, negated_conjecture, (esk5_0!=esk7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_142, c_0_143])])).
cnf(c_0_146, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_140]), c_0_106])]), c_0_145]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 147
# Proof object clause steps            : 112
# Proof object formula steps           : 35
# Proof object conjectures             : 52
# Proof object clause conjectures      : 49
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 71
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 44
# Processed clauses                    : 2311
# ...of these trivial                  : 134
# ...subsumed                          : 1524
# ...remaining for further processing  : 653
# Other redundant clauses eliminated   : 30
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 207
# Generated clauses                    : 26955
# ...of the previous two non-trivial   : 17074
# Contextual simplify-reflections      : 4
# Paramodulations                      : 26917
# Factorizations                       : 8
# Equation resolutions                 : 30
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 397
#    Positive orientable unit clauses  : 105
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 24
#    Non-unit-clauses                  : 265
# Current number of unprocessed clauses: 14367
# ...number of literals in the above   : 34949
# Current number of archived formulas  : 0
# Current number of archived clauses   : 243
# Clause-clause subsumption calls (NU) : 16128
# Rec. Clause-clause subsumption calls : 13748
# Non-unit clause-clause subsumptions  : 660
# Unit Clause-clause subsumption calls : 5887
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 294
# BW rewrite match successes           : 182
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 264304

# -------------------------------------------------
# User time                : 0.152 s
# System time              : 0.003 s
# Total time               : 0.155 s
# Maximum resident set size: 28628 pages
