# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/9472917a8d15449d8e0ad28dd3cbcef3.E_input', d5_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/9472917a8d15449d8e0ad28dd3cbcef3.E_input', t3_xboole_0)).
fof(t84_xboole_1, conjecture, ![X1, X2, X3]:~((~(r1_xboole_0(X1,X2))&(r1_xboole_0(X1,X3)&r1_xboole_0(X1,k4_xboole_0(X2,X3))))), file('data/MPTP/proofs/9472917a8d15449d8e0ad28dd3cbcef3.E_input', t84_xboole_1)).
fof(c_0_3, plain, ![X8, X9, X10, X11, X12, X13, X14, X15]:((((r2_hidden(X11,X8)|~r2_hidden(X11,X10)|X10!=k4_xboole_0(X8,X9))&(~r2_hidden(X11,X9)|~r2_hidden(X11,X10)|X10!=k4_xboole_0(X8,X9)))&(~r2_hidden(X12,X8)|r2_hidden(X12,X9)|r2_hidden(X12,X10)|X10!=k4_xboole_0(X8,X9)))&((~r2_hidden(esk4_3(X13,X14,X15),X15)|(~r2_hidden(esk4_3(X13,X14,X15),X13)|r2_hidden(esk4_3(X13,X14,X15),X14))|X15=k4_xboole_0(X13,X14))&((r2_hidden(esk4_3(X13,X14,X15),X13)|r2_hidden(esk4_3(X13,X14,X15),X15)|X15=k4_xboole_0(X13,X14))&(~r2_hidden(esk4_3(X13,X14,X15),X14)|r2_hidden(esk4_3(X13,X14,X15),X15)|X15=k4_xboole_0(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_4, plain, ![X17, X18, X20, X21, X22]:(((r2_hidden(esk5_2(X17,X18),X17)|r1_xboole_0(X17,X18))&(r2_hidden(esk5_2(X17,X18),X18)|r1_xboole_0(X17,X18)))&(~r2_hidden(X22,X20)|~r2_hidden(X22,X21)|~r1_xboole_0(X20,X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_5, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3]:~((~(r1_xboole_0(X1,X2))&(r1_xboole_0(X1,X3)&r1_xboole_0(X1,k4_xboole_0(X2,X3)))))), inference(assume_negation,[status(cth)],[t84_xboole_1])).
cnf(c_0_9, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_11, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk5_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_12, plain, (r2_hidden(esk5_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_13, negated_conjecture, (~r1_xboole_0(esk1_0,esk2_0)&(r1_xboole_0(esk1_0,esk3_0)&r1_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
cnf(c_0_14, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_15, plain, (k4_xboole_0(X1,X2)=X1|~r2_hidden(esk4_3(X1,X2,X1),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_6, c_0_9])).
cnf(c_0_16, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_9]), c_0_9])).
cnf(c_0_17, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_18, negated_conjecture, (r1_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, negated_conjecture, (r1_xboole_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22, negated_conjecture, (r1_xboole_0(k4_xboole_0(esk2_0,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_24, negated_conjecture, (r1_xboole_0(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_17, c_0_19])).
cnf(c_0_25, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk5_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_20, c_0_12])).
cnf(c_0_26, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk2_0,esk3_0))=esk1_0), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (k4_xboole_0(esk1_0,esk3_0)=esk1_0), inference(spm,[status(thm)],[c_0_21, c_0_24])).
cnf(c_0_29, negated_conjecture, (r1_xboole_0(X1,esk1_0)|~r2_hidden(esk5_2(X1,esk1_0),k4_xboole_0(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30, plain, (r2_hidden(esk5_2(X1,X2),k4_xboole_0(X1,X3))|r2_hidden(esk5_2(X1,X2),X3)|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_27, c_0_7])).
cnf(c_0_31, negated_conjecture, (r1_xboole_0(X1,esk1_0)|~r2_hidden(esk5_2(X1,esk1_0),esk3_0)), inference(spm,[status(thm)],[c_0_25, c_0_28])).
cnf(c_0_32, negated_conjecture, (r1_xboole_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
cnf(c_0_33, negated_conjecture, (~r1_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_34, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_32]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 28
# Proof object formula steps           : 7
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 3
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 299
# ...of these trivial                  : 42
# ...subsumed                          : 107
# ...remaining for further processing  : 150
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 4
# Generated clauses                    : 2274
# ...of the previous two non-trivial   : 1273
# Contextual simplify-reflections      : 4
# Paramodulations                      : 2246
# Factorizations                       : 10
# Equation resolutions                 : 18
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 145
#    Positive orientable unit clauses  : 62
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 80
# Current number of unprocessed clauses: 983
# ...number of literals in the above   : 2510
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 1047
# Rec. Clause-clause subsumption calls : 806
# Non-unit clause-clause subsumptions  : 83
# Unit Clause-clause subsumption calls : 174
# Rewrite failures with RHS unbound    : 26
# BW rewrite match attempts            : 140
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 21401

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.003 s
# Total time               : 0.020 s
# Maximum resident set size: 25940 pages
