# No SInE strategy applied
# Auto-Mode selected heuristic H_____047_B31_F1_PI_AE_R4_CS_SP_S4c
# and selection function SelectCQPrecWNTNp.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', redefinition_k6_domain_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t75_enumset1)).
fof(t50_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2)))))), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t50_orders_2)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', d6_enumset1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', dt_l1_orders_2)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t35_orders_2)).
fof(t49_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>~((r2_hidden(X2,X3)&r2_hidden(X2,k2_orders_2(X1,X3))))))), file('data/MPTP/proofs/949cf5cc62684e0eb9e756ca271430df.E_input', t49_orders_2)).
fof(c_0_14, plain, ![X11, X12]:(v1_xboole_0(X11)|~m1_subset_1(X12,X11)|k6_domain_1(X11,X12)=k1_tarski(X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_15, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_16, plain, ![X14, X15]:k1_enumset1(X14,X14,X15)=k2_tarski(X14,X15), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_17, plain, ![X16, X17, X18]:k2_enumset1(X16,X16,X17,X18)=k1_enumset1(X16,X17,X18), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_18, plain, ![X19, X20, X21, X22]:k3_enumset1(X19,X19,X20,X21,X22)=k2_enumset1(X19,X20,X21,X22), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_19, plain, ![X23, X24, X25, X26, X27]:k4_enumset1(X23,X23,X24,X25,X26,X27)=k3_enumset1(X23,X24,X25,X26,X27), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_20, plain, ![X28, X29, X30, X31, X32, X33]:k5_enumset1(X28,X28,X29,X30,X31,X32,X33)=k4_enumset1(X28,X29,X30,X31,X32,X33), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_21, plain, ![X34, X35, X36, X37, X38, X39, X40]:k6_enumset1(X34,X34,X35,X36,X37,X38,X39,X40)=k5_enumset1(X34,X35,X36,X37,X38,X39,X40), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_22, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~(r2_hidden(X2,k2_orders_2(X1,k6_domain_1(u1_struct_0(X1),X2))))))), inference(assume_negation,[status(cth)],[t50_orders_2])).
cnf(c_0_23, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_31, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
fof(c_0_32, plain, ![X43, X44, X45, X46, X47, X48, X49, X50, X51, X52, X53, X54, X55, X56, X57, X58, X59, X60, X61, X62]:(((~r2_hidden(X52,X51)|(X52=X43|(X52=X44|(X52=X45|(X52=X46|(X52=X47|(X52=X48|(X52=X49|X52=X50)))))))|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((X53!=X43|r2_hidden(X53,X51)|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((X53!=X44|r2_hidden(X53,X51)|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((X53!=X45|r2_hidden(X53,X51)|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((X53!=X46|r2_hidden(X53,X51)|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((X53!=X47|r2_hidden(X53,X51)|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((X53!=X48|r2_hidden(X53,X51)|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((X53!=X49|r2_hidden(X53,X51)|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&(X53!=X50|r2_hidden(X53,X51)|X51!=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))))))))))&(((esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)!=X54|~r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61))&((esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)!=X55|~r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61))&((esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)!=X56|~r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61))&((esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)!=X57|~r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61))&((esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)!=X58|~r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61))&((esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)!=X59|~r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61))&((esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)!=X60|~r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61))&(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)!=X61|~r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61)))))))))&(r2_hidden(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62),X62)|(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)=X54|(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)=X55|(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)=X56|(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)=X57|(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)=X58|(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)=X59|(esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)=X60|esk3_9(X54,X55,X56,X57,X58,X59,X60,X61,X62)=X61)))))))|X62=k6_enumset1(X54,X55,X56,X57,X58,X59,X60,X61)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
fof(c_0_33, plain, ![X64]:(v2_struct_0(X64)|~l1_struct_0(X64)|~v1_xboole_0(u1_struct_0(X64))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_34, plain, (k6_domain_1(X1,X2)=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24]), c_0_25]), c_0_26]), c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X4,X5,X6,X7,X8,X9,X10,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_38, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_40, plain, ![X65]:(~l1_orders_2(X65)|l1_struct_0(X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_41, plain, ![X66, X67]:((v6_orders_2(k6_domain_1(u1_struct_0(X66),X67),X66)|~m1_subset_1(X67,u1_struct_0(X66))|(v2_struct_0(X66)|(~v3_orders_2(X66)|~l1_orders_2(X66))))&(m1_subset_1(k6_domain_1(u1_struct_0(X66),X67),k1_zfmisc_1(u1_struct_0(X66)))|~m1_subset_1(X67,u1_struct_0(X66))|(v2_struct_0(X66)|(~v3_orders_2(X66)|~l1_orders_2(X66))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
fof(c_0_42, plain, ![X68, X69, X70]:(v2_struct_0(X68)|(~v3_orders_2(X68)|(~v4_orders_2(X68)|(~v5_orders_2(X68)|~l1_orders_2(X68))))|(~m1_subset_1(X69,u1_struct_0(X68))|(~m1_subset_1(X70,k1_zfmisc_1(u1_struct_0(X68)))|(~r2_hidden(X69,X70)|~r2_hidden(X69,k2_orders_2(X68,X70)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_orders_2])])])])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|X2!=k6_enumset1(X3,X4,X5,X6,X7,X8,X9,X1)), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_44, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_45, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_46, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_47, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_49, plain, (v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,X3)|~r2_hidden(X2,k2_orders_2(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, plain, (r2_hidden(X1,k6_enumset1(X2,X3,X4,X5,X6,X7,X8,X1))), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_52, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_35]), c_0_46]), c_0_48])]), c_0_39])).
cnf(c_0_54, plain, (v2_struct_0(X1)|~r2_hidden(X2,k2_orders_2(X1,k6_enumset1(X3,X4,X5,X6,X7,X8,X9,X2)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k6_enumset1(X3,X4,X5,X6,X7,X8,X9,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk2_0,k2_orders_2(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_57, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_53, c_0_52])).
cnf(c_0_59, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_46]), c_0_56]), c_0_57]), c_0_48]), c_0_58]), c_0_35])]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 31
# Proof object formula steps           : 29
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 31
# Processed clauses                    : 63
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 63
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 135
# ...of the previous two non-trivial   : 128
# Contextual simplify-reflections      : 0
# Paramodulations                      : 109
# Factorizations                       : 0
# Equation resolutions                 : 26
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 32
# Current number of unprocessed clauses: 90
# ...number of literals in the above   : 1026
# Current number of archived formulas  : 0
# Current number of archived clauses   : 12
# Clause-clause subsumption calls (NU) : 231
# Rec. Clause-clause subsumption calls : 121
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 30
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5447

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.005 s
# Total time               : 0.014 s
# Maximum resident set size: 28700 pages
