# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t51_mcart_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4)))))))))), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', t51_mcart_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', d1_tarski)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', t50_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', d3_zfmisc_1)).
fof(t29_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k3_mcart_1(X3,X4,X5))))))), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', t29_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', d3_mcart_1)).
fof(t48_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4)))))))), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', t48_mcart_1)).
fof(t20_mcart_1, axiom, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', t20_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', t7_mcart_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', t113_zfmisc_1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/94b85280690b4664845c2f137a05f26d.E_input', t152_zfmisc_1)).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4))))))))))), inference(assume_negation,[status(cth)],[t51_mcart_1])).
fof(c_0_12, plain, ![X10, X11, X12, X13, X14, X15]:(((~r2_hidden(X12,X11)|X12=X10|X11!=k1_tarski(X10))&(X13!=X10|r2_hidden(X13,X11)|X11!=k1_tarski(X10)))&((~r2_hidden(esk5_2(X14,X15),X15)|esk5_2(X14,X15)!=X14|X15=k1_tarski(X14))&(r2_hidden(esk5_2(X14,X15),X15)|esk5_2(X14,X15)=X14|X15=k1_tarski(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_13, plain, ![X17, X18, X19, X20]:((k5_mcart_1(X17,X18,X19,X20)=k1_mcart_1(k1_mcart_1(X20))|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0)&((k6_mcart_1(X17,X18,X19,X20)=k2_mcart_1(k1_mcart_1(X20))|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0)&(k7_mcart_1(X17,X18,X19,X20)=k2_mcart_1(X20)|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
fof(c_0_14, plain, ![X21, X22, X23]:k3_zfmisc_1(X21,X22,X23)=k2_zfmisc_1(k2_zfmisc_1(X21,X22),X23), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_15, negated_conjecture, (esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&(esk4_0=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|(esk4_0=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
fof(c_0_16, plain, ![X36, X38, X39, X40]:((r2_hidden(esk6_1(X36),X36)|X36=k1_xboole_0)&((~r2_hidden(X38,X36)|esk6_1(X36)!=k3_mcart_1(X38,X39,X40)|X36=k1_xboole_0)&(~r2_hidden(X39,X36)|esk6_1(X36)!=k3_mcart_1(X38,X39,X40)|X36=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_mcart_1])])])])])])])).
fof(c_0_17, plain, ![X28, X29, X30]:k3_mcart_1(X28,X29,X30)=k4_tarski(k4_tarski(X28,X29),X30), inference(variable_rename,[status(thm)],[d3_mcart_1])).
cnf(c_0_18, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_19, plain, ![X24, X25, X26, X27]:(X24=k1_xboole_0|(X25=k1_xboole_0|(X26=k1_xboole_0|(~m1_subset_1(X27,k3_zfmisc_1(X24,X25,X26))|X27=k3_mcart_1(k5_mcart_1(X24,X25,X26,X27),k6_mcart_1(X24,X25,X26,X27),k7_mcart_1(X24,X25,X26,X27)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_mcart_1])])])])).
cnf(c_0_20, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k3_mcart_1(X3,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (r2_hidden(esk6_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk4_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_22, c_0_21])).
cnf(c_0_30, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_31, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_32, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_33, plain, ![X31, X32, X33]:((X31!=k1_mcart_1(X31)|X31!=k4_tarski(X32,X33))&(X31!=k2_mcart_1(X31)|X31!=k4_tarski(X32,X33))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_mcart_1])])])])])).
fof(c_0_34, plain, ![X34, X35]:(k1_mcart_1(k4_tarski(X34,X35))=X34&k2_mcart_1(k4_tarski(X34,X35))=X35), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_35, plain, (X2=k1_xboole_0|esk6_1(X2)!=k4_tarski(k4_tarski(X3,X1),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_36, plain, (esk6_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_37, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4=k4_tarski(k4_tarski(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4)),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_21]), c_0_24])).
cnf(c_0_38, negated_conjecture, (k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=k2_mcart_1(esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_39, plain, (X1!=k2_mcart_1(X1)|X1!=k4_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (k1_tarski(k4_tarski(k4_tarski(X1,X2),X3))=k1_xboole_0|~r2_hidden(X2,k1_tarski(k4_tarski(k4_tarski(X1,X2),X3)))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_42, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)),k2_mcart_1(esk4_0))=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_29])]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_43, negated_conjecture, (esk4_0=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)|esk4_0=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_44, plain, (k4_tarski(X1,X2)!=X2), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_39]), c_0_40])).
cnf(c_0_45, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_46, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk6_1(X2)!=k3_mcart_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_47, negated_conjecture, (k1_tarski(esk4_0)=k1_xboole_0|~r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k1_tarski(esk4_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, negated_conjecture, (k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_43]), c_0_29])]), c_0_44]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_49, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_45])])).
fof(c_0_50, plain, ![X41, X42]:((k2_zfmisc_1(X41,X42)!=k1_xboole_0|(X41=k1_xboole_0|X42=k1_xboole_0))&((X41!=k1_xboole_0|k2_zfmisc_1(X41,X42)=k1_xboole_0)&(X42!=k1_xboole_0|k2_zfmisc_1(X41,X42)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_51, plain, (X2=k1_xboole_0|esk6_1(X2)!=k4_tarski(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_46, c_0_24])).
cnf(c_0_52, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk4_0|k1_tarski(esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])])).
fof(c_0_53, plain, ![X43, X44]:~r2_hidden(X43,k2_zfmisc_1(X43,X44)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_54, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_55, plain, (k1_tarski(k4_tarski(k4_tarski(X1,X2),X3))=k1_xboole_0|~r2_hidden(X1,k1_tarski(k4_tarski(k4_tarski(X1,X2),X3)))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_36])])).
cnf(c_0_56, negated_conjecture, (k4_tarski(k4_tarski(esk4_0,k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)),k2_mcart_1(esk4_0))=esk4_0|k1_tarski(esk4_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_52]), c_0_38]), c_0_29])]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_57, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_58, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_54])).
cnf(c_0_59, negated_conjecture, (k1_tarski(esk4_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_49])])).
cnf(c_0_60, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_61, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_59]), c_0_60]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 39
# Proof object formula steps           : 23
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 24
# Processed clauses                    : 97
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 84
# Other redundant clauses eliminated   : 15
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 6
# Generated clauses                    : 129
# ...of the previous two non-trivial   : 105
# Contextual simplify-reflections      : 0
# Paramodulations                      : 113
# Factorizations                       : 2
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 46
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 55
# ...number of literals in the above   : 207
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 165
# Rec. Clause-clause subsumption calls : 54
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 115
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3251

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28796 pages
