# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t70_enumset1)).
fof(t146_funct_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>r1_tarski(X1,k10_relat_1(X2,k9_relat_1(X2,X1))))), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t146_funct_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t100_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', idempotence_k3_xboole_0)).
fof(t139_funct_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(k7_relat_1(X3,X1),X2)=k3_xboole_0(X1,k10_relat_1(X3,X2))), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t139_funct_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t1_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', l32_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', d10_xboole_0)).
fof(t169_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t169_relat_1)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t87_relat_1)).
fof(t189_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k7_relat_1(X1,k1_relat_1(X2))=k7_relat_1(X1,k1_relat_1(k7_relat_1(X2,k1_relat_1(X1)))))), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t189_relat_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t71_relat_1)).
fof(dt_k6_relat_1, axiom, ![X1]:v1_relat_1(k6_relat_1(X1)), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', dt_k6_relat_1)).
fof(dt_k7_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2))), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', dt_k7_relat_1)).
fof(t97_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k1_relat_1(X2),X1)=>k7_relat_1(X2,X1)=X2)), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t97_relat_1)).
fof(t98_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k7_relat_1(X1,k1_relat_1(X1))=X1), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t98_relat_1)).
fof(t146_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t146_relat_1)).
fof(t162_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(r1_tarski(X2,X3)=>k9_relat_1(k7_relat_1(X1,X3),X2)=k9_relat_1(X1,X2))), file('data/MPTP/proofs/94b8735d216f44969f9c455e504f77ea.E_input', t162_relat_1)).
fof(c_0_19, plain, ![X4, X5]:k1_setfam_1(k2_tarski(X4,X5))=k3_xboole_0(X4,X5), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_20, plain, ![X6, X7]:k1_enumset1(X6,X6,X7)=k2_tarski(X6,X7), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>r1_tarski(X1,k10_relat_1(X2,k9_relat_1(X2,X1)))))), inference(assume_negation,[status(cth)],[t146_funct_1])).
fof(c_0_22, plain, ![X10, X11]:k4_xboole_0(X10,X11)=k5_xboole_0(X10,k3_xboole_0(X10,X11)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_23, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X12]:k3_xboole_0(X12,X12)=X12, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_26, plain, ![X13, X14, X15]:(~v1_relat_1(X15)|k10_relat_1(k7_relat_1(X15,X13),X14)=k3_xboole_0(X13,k10_relat_1(X15,X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t139_funct_1])])).
fof(c_0_27, plain, ![X16, X17, X18]:(~r1_tarski(X16,X17)|~r1_tarski(X17,X18)|r1_tarski(X16,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_28, negated_conjecture, (v1_relat_1(esk2_0)&(r1_tarski(esk1_0,k1_relat_1(esk2_0))&~r1_tarski(esk1_0,k10_relat_1(esk2_0,k9_relat_1(esk2_0,esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
fof(c_0_29, plain, ![X19, X20]:((k4_xboole_0(X19,X20)!=k1_xboole_0|r1_tarski(X19,X20))&(~r1_tarski(X19,X20)|k4_xboole_0(X19,X20)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_30, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k3_xboole_0(X2,k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_34, plain, ![X21, X22]:(((r1_tarski(X21,X22)|X21!=X22)&(r1_tarski(X22,X21)|X21!=X22))&(~r1_tarski(X21,X22)|~r1_tarski(X22,X21)|X21=X22)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_35, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, negated_conjecture, (r1_tarski(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, plain, (k1_setfam_1(k1_enumset1(X1,X1,X1))=X1), inference(rw,[status(thm)],[c_0_32, c_0_31])).
cnf(c_0_40, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k1_setfam_1(k1_enumset1(X2,X2,k10_relat_1(X1,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_33, c_0_31])).
fof(c_0_41, plain, ![X23]:(~v1_relat_1(X23)|k10_relat_1(X23,k2_relat_1(X23))=k1_relat_1(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t169_relat_1])])).
cnf(c_0_42, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, negated_conjecture, (r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_45, plain, ![X24, X25]:(~v1_relat_1(X25)|r1_tarski(k1_relat_1(k7_relat_1(X25,X24)),X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (k10_relat_1(k7_relat_1(X1,k10_relat_1(X1,X2)),X2)=k10_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, plain, (k10_relat_1(X1,k2_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, plain, (k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_42, c_0_38])).
cnf(c_0_50, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_51, negated_conjecture, (r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(X2,esk1_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_35, c_0_44])).
cnf(c_0_52, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_53, plain, ![X26, X27]:(~v1_relat_1(X26)|(~v1_relat_1(X27)|k7_relat_1(X26,k1_relat_1(X27))=k7_relat_1(X26,k1_relat_1(k7_relat_1(X27,k1_relat_1(X26)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t189_relat_1])])])).
fof(c_0_54, plain, ![X28]:(k1_relat_1(k6_relat_1(X28))=X28&k2_relat_1(k6_relat_1(X28))=X28), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_55, plain, ![X29]:v1_relat_1(k6_relat_1(X29)), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
cnf(c_0_56, plain, (r1_tarski(X1,k10_relat_1(X2,X3))|k5_xboole_0(X1,k10_relat_1(k7_relat_1(X2,X1),X3))!=k1_xboole_0|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_46, c_0_40])).
cnf(c_0_57, plain, (k10_relat_1(k7_relat_1(X1,k1_relat_1(X1)),k2_relat_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_58, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_39]), c_0_50])])).
cnf(c_0_59, negated_conjecture, (r1_tarski(X1,k1_relat_1(esk2_0))|~r1_tarski(X1,k1_relat_1(k7_relat_1(X2,esk1_0)))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, plain, (k7_relat_1(X1,k1_relat_1(X2))=k7_relat_1(X1,k1_relat_1(k7_relat_1(X2,k1_relat_1(X1))))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_61, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (r1_tarski(k1_relat_1(X1),k10_relat_1(X1,k2_relat_1(X1)))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])])).
cnf(c_0_64, negated_conjecture, (r1_tarski(k1_relat_1(k7_relat_1(X1,k1_relat_1(k7_relat_1(X2,esk1_0)))),k1_relat_1(esk2_0))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_59, c_0_52])).
cnf(c_0_65, plain, (k7_relat_1(k6_relat_1(X1),k1_relat_1(k7_relat_1(X2,X1)))=k7_relat_1(k6_relat_1(X1),k1_relat_1(X2))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])])).
cnf(c_0_66, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k7_relat_1(X3,X2)))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_35, c_0_52])).
fof(c_0_67, plain, ![X30, X31]:(~v1_relat_1(X30)|v1_relat_1(k7_relat_1(X30,X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])).
cnf(c_0_68, plain, (r1_tarski(X1,k10_relat_1(X2,k2_relat_1(X2)))|~r1_tarski(X1,k1_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_35, c_0_63])).
cnf(c_0_69, negated_conjecture, (r1_tarski(k1_relat_1(k7_relat_1(k6_relat_1(esk1_0),k1_relat_1(X1))),k1_relat_1(esk2_0))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_62])])).
cnf(c_0_70, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,k1_relat_1(k7_relat_1(X2,X3)))),X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_66, c_0_52])).
cnf(c_0_71, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_72, plain, ![X32, X33]:(~v1_relat_1(X33)|(~r1_tarski(k1_relat_1(X33),X32)|k7_relat_1(X33,X32)=X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t97_relat_1])])).
cnf(c_0_73, plain, (r1_tarski(X1,k10_relat_1(X2,k2_relat_1(X2)))|~r1_tarski(X3,k1_relat_1(X2))|~r1_tarski(X1,X3)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_35, c_0_68])).
cnf(c_0_74, negated_conjecture, (r1_tarski(k1_relat_1(k7_relat_1(k6_relat_1(esk1_0),X1)),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_61]), c_0_62])])).
cnf(c_0_75, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_76, plain, ![X34]:(~v1_relat_1(X34)|k7_relat_1(X34,k1_relat_1(X34))=X34), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t98_relat_1])])).
cnf(c_0_77, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_78, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k7_relat_1(X3,k1_relat_1(k7_relat_1(X4,X2)))))|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_35, c_0_70])).
cnf(c_0_79, plain, (k1_setfam_1(k1_enumset1(X1,X1,k1_relat_1(X2)))=k10_relat_1(k7_relat_1(k7_relat_1(X2,k1_relat_1(X2)),X1),k2_relat_1(X2))|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_57]), c_0_71])).
cnf(c_0_80, plain, (k7_relat_1(X1,X2)=X1|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_81, negated_conjecture, (r1_tarski(X1,k10_relat_1(esk2_0,k2_relat_1(esk2_0)))|~r1_tarski(X1,k1_relat_1(k7_relat_1(k6_relat_1(esk1_0),X2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])])).
cnf(c_0_82, plain, (k7_relat_1(X1,k1_relat_1(X1))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_83, plain, (k10_relat_1(k6_relat_1(X1),X1)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_77]), c_0_61]), c_0_62])])).
cnf(c_0_84, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k7_relat_1(k6_relat_1(X2),k1_relat_1(X3))))|~v1_relat_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_65]), c_0_62])])).
cnf(c_0_85, plain, (k5_xboole_0(X1,k10_relat_1(k7_relat_1(X2,X1),X3))=k1_xboole_0|~r1_tarski(X1,k10_relat_1(X2,X3))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_49, c_0_40])).
cnf(c_0_86, plain, (k1_setfam_1(k1_enumset1(X1,X1,k1_relat_1(X2)))=k10_relat_1(k7_relat_1(X2,k1_relat_1(X2)),k2_relat_1(X2))|~r1_tarski(k1_relat_1(k7_relat_1(X2,k1_relat_1(X2))),X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_71])).
cnf(c_0_87, negated_conjecture, (r1_tarski(k1_relat_1(k7_relat_1(k6_relat_1(esk1_0),X1)),k10_relat_1(esk2_0,k2_relat_1(esk2_0)))), inference(spm,[status(thm)],[c_0_81, c_0_50])).
cnf(c_0_88, plain, (k7_relat_1(k6_relat_1(X1),X1)=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_61]), c_0_62])])).
cnf(c_0_89, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k10_relat_1(k7_relat_1(k6_relat_1(X2),X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_83]), c_0_62])])).
cnf(c_0_90, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relat_1(k7_relat_1(k6_relat_1(X2),X3)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_61]), c_0_62])])).
cnf(c_0_91, plain, (v1_relat_1(k7_relat_1(k6_relat_1(X1),k1_relat_1(X2)))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_65]), c_0_62])])).
cnf(c_0_92, plain, (k1_setfam_1(k1_enumset1(X1,X1,k1_relat_1(X2)))=k10_relat_1(k7_relat_1(X2,X1),k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_40, c_0_48])).
cnf(c_0_93, plain, (k5_xboole_0(k1_relat_1(X1),k10_relat_1(X1,X2))=k1_xboole_0|~r1_tarski(k1_relat_1(X1),k10_relat_1(X1,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_85, c_0_82])).
cnf(c_0_94, negated_conjecture, (k10_relat_1(k7_relat_1(k6_relat_1(esk1_0),k10_relat_1(esk2_0,k2_relat_1(esk2_0))),esk1_0)=esk1_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_61]), c_0_61]), c_0_88]), c_0_77]), c_0_83]), c_0_62])]), c_0_89])).
cnf(c_0_95, plain, (r1_tarski(k1_relat_1(k7_relat_1(k6_relat_1(X1),X2)),X1)), inference(spm,[status(thm)],[c_0_90, c_0_50])).
cnf(c_0_96, plain, (v1_relat_1(k7_relat_1(k6_relat_1(X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_61]), c_0_62])])).
cnf(c_0_97, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_98, plain, (r1_tarski(X1,k1_relat_1(X2))|k5_xboole_0(X1,k10_relat_1(k7_relat_1(X2,X1),k2_relat_1(X2)))!=k1_xboole_0|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_46, c_0_92])).
cnf(c_0_99, negated_conjecture, (k5_xboole_0(k1_relat_1(k7_relat_1(k6_relat_1(esk1_0),k10_relat_1(esk2_0,k2_relat_1(esk2_0)))),esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95]), c_0_96])])).
cnf(c_0_100, plain, (k1_relat_1(k7_relat_1(X1,X2))=X2|~r1_tarski(X2,k1_relat_1(k7_relat_1(X1,X2)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_97, c_0_52])).
cnf(c_0_101, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,k1_relat_1(X2))),k1_relat_1(k7_relat_1(X2,k1_relat_1(X1))))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_52, c_0_60])).
cnf(c_0_102, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k10_relat_1(k7_relat_1(k6_relat_1(X2),X1),X2))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_77]), c_0_61]), c_0_62])])).
cnf(c_0_103, negated_conjecture, (k5_xboole_0(k1_relat_1(k7_relat_1(k6_relat_1(esk1_0),k1_relat_1(esk2_0))),esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_48]), c_0_75])])).
cnf(c_0_104, plain, (k1_relat_1(k7_relat_1(X1,k1_relat_1(X2)))=k1_relat_1(k7_relat_1(X2,k1_relat_1(X1)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_60]), c_0_101])).
cnf(c_0_105, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,X2)!=k1_xboole_0|~r1_tarski(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_80]), c_0_83]), c_0_61]), c_0_62])])).
cnf(c_0_106, negated_conjecture, (k5_xboole_0(k1_relat_1(k7_relat_1(esk2_0,esk1_0)),esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_61]), c_0_75]), c_0_62])])).
fof(c_0_107, plain, ![X35]:(~v1_relat_1(X35)|k9_relat_1(X35,k1_relat_1(X35))=k2_relat_1(X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t146_relat_1])])).
fof(c_0_108, plain, ![X36, X37, X38]:(~v1_relat_1(X36)|(~r1_tarski(X37,X38)|k9_relat_1(k7_relat_1(X36,X38),X37)=k9_relat_1(X36,X37))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t162_relat_1])])])).
cnf(c_0_109, negated_conjecture, (r1_tarski(k1_relat_1(k7_relat_1(esk2_0,esk1_0)),esk1_0)|~r1_tarski(esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0)))), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_110, plain, (r1_tarski(k1_relat_1(k7_relat_1(k6_relat_1(X1),k1_relat_1(X2))),k1_relat_1(k7_relat_1(X2,X1)))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_61]), c_0_62])])).
cnf(c_0_111, plain, (k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_112, plain, (k9_relat_1(k7_relat_1(X1,X3),X2)=k9_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_113, negated_conjecture, (k1_relat_1(k7_relat_1(esk2_0,esk1_0))=esk1_0|~r1_tarski(esk1_0,k1_relat_1(k7_relat_1(esk2_0,esk1_0)))), inference(spm,[status(thm)],[c_0_97, c_0_109])).
cnf(c_0_114, plain, (r1_tarski(X1,k1_relat_1(k7_relat_1(X2,X1)))|~r1_tarski(X1,k1_relat_1(X2))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_80]), c_0_61]), c_0_61]), c_0_62])])).
cnf(c_0_115, plain, (k9_relat_1(X1,k1_relat_1(k7_relat_1(X1,X2)))=k2_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_71]), c_0_52])).
cnf(c_0_116, negated_conjecture, (k1_relat_1(k7_relat_1(esk2_0,esk1_0))=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_36]), c_0_75])])).
cnf(c_0_117, plain, (r1_tarski(X1,k10_relat_1(X2,k2_relat_1(k7_relat_1(X2,X1))))|k5_xboole_0(X1,k1_relat_1(k7_relat_1(X2,X1)))!=k1_xboole_0|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_48]), c_0_71])).
cnf(c_0_118, negated_conjecture, (k2_relat_1(k7_relat_1(esk2_0,esk1_0))=k9_relat_1(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_75])])).
cnf(c_0_119, negated_conjecture, (~r1_tarski(esk1_0,k10_relat_1(esk2_0,k9_relat_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_120, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_116]), c_0_58]), c_0_75])]), c_0_119]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 121
# Proof object clause steps            : 82
# Proof object formula steps           : 39
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 51
# Proof object simplifying inferences  : 80
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 22
# Processed clauses                    : 4907
# ...of these trivial                  : 14
# ...subsumed                          : 4179
# ...remaining for further processing  : 714
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 38
# Backward-rewritten                   : 10
# Generated clauses                    : 44502
# ...of the previous two non-trivial   : 41568
# Contextual simplify-reflections      : 132
# Paramodulations                      : 44500
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 643
#    Positive orientable unit clauses  : 51
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 585
# Current number of unprocessed clauses: 36522
# ...number of literals in the above   : 152559
# Current number of archived formulas  : 0
# Current number of archived clauses   : 72
# Clause-clause subsumption calls (NU) : 68676
# Rec. Clause-clause subsumption calls : 38335
# Non-unit clause-clause subsumptions  : 3107
# Unit Clause-clause subsumption calls : 1112
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 342
# BW rewrite match successes           : 16
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1033799

# -------------------------------------------------
# User time                : 0.507 s
# System time              : 0.013 s
# Total time               : 0.520 s
# Maximum resident set size: 35048 pages
