# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', d10_xboole_0)).
fof(t4_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3))))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', t4_tsep_1)).
fof(t23_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(m1_pre_topc(X2,X3)=>((r1_tsep_1(X2,X4)&r1_tsep_1(X4,X2))|(~(r1_tsep_1(X3,X4))&~(r1_tsep_1(X4,X3))))))))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', t23_tmap_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', dt_m1_pre_topc)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', cc1_pre_topc)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', l3_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', t3_subset)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', t3_xboole_0)).
fof(d3_tsep_1, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_struct_0(X2)=>(r1_tsep_1(X1,X2)<=>r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', d3_tsep_1)).
fof(symmetry_r1_tsep_1, axiom, ![X1, X2]:((l1_struct_0(X1)&l1_struct_0(X2))=>(r1_tsep_1(X1,X2)=>r1_tsep_1(X2,X1))), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', symmetry_r1_tsep_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/94e1300aea754a039180f4b87cec758e.E_input', dt_l1_pre_topc)).
fof(c_0_11, plain, ![X28, X29]:(((r1_tarski(X28,X29)|X28!=X29)&(r1_tarski(X29,X28)|X28!=X29))&(~r1_tarski(X28,X29)|~r1_tarski(X29,X28)|X28=X29)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_12, plain, ![X25, X26, X27]:((~r1_tarski(u1_struct_0(X26),u1_struct_0(X27))|m1_pre_topc(X26,X27)|~m1_pre_topc(X27,X25)|~m1_pre_topc(X26,X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25)))&(~m1_pre_topc(X26,X27)|r1_tarski(u1_struct_0(X26),u1_struct_0(X27))|~m1_pre_topc(X27,X25)|~m1_pre_topc(X26,X25)|(~v2_pre_topc(X25)|~l1_pre_topc(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])).
cnf(c_0_13, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(m1_pre_topc(X2,X3)=>((r1_tsep_1(X2,X4)&r1_tsep_1(X4,X2))|(~(r1_tsep_1(X3,X4))&~(r1_tsep_1(X4,X3)))))))))), inference(assume_negation,[status(cth)],[t23_tmap_1])).
cnf(c_0_15, plain, (m1_pre_topc(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_13])).
fof(c_0_17, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&(m1_pre_topc(esk2_0,esk3_0)&((~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk4_0,esk2_0))&(r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(esk4_0,esk3_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_18, plain, ![X32, X33]:(~l1_pre_topc(X32)|(~m1_pre_topc(X33,X32)|l1_pre_topc(X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_19, plain, ![X30, X31]:(~v2_pre_topc(X30)|~l1_pre_topc(X30)|(~m1_pre_topc(X31,X30)|v2_pre_topc(X31))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
fof(c_0_20, plain, ![X9, X10, X11]:(~m1_subset_1(X10,k1_zfmisc_1(X9))|(~r2_hidden(X11,X10)|r2_hidden(X11,X9))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_21, plain, ![X13, X14]:((~m1_subset_1(X13,k1_zfmisc_1(X14))|r1_tarski(X13,X14))&(~r1_tarski(X13,X14)|m1_subset_1(X13,k1_zfmisc_1(X14)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_22, plain, (m1_pre_topc(X1,X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_29, plain, ![X19, X20, X22, X23, X24]:(((r2_hidden(esk5_2(X19,X20),X19)|r1_xboole_0(X19,X20))&(r2_hidden(esk5_2(X19,X20),X20)|r1_xboole_0(X19,X20)))&(~r2_hidden(X24,X22)|~r2_hidden(X24,X23)|~r1_xboole_0(X22,X23))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_30, plain, ![X15, X16]:((~r1_tsep_1(X15,X16)|r1_xboole_0(u1_struct_0(X15),u1_struct_0(X16))|~l1_struct_0(X16)|~l1_struct_0(X15))&(~r1_xboole_0(u1_struct_0(X15),u1_struct_0(X16))|r1_tsep_1(X15,X16)|~l1_struct_0(X16)|~l1_struct_0(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tsep_1])])])])).
cnf(c_0_31, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_33, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_34, negated_conjecture, (m1_pre_topc(esk4_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_35, negated_conjecture, (l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_23]), c_0_24])])).
cnf(c_0_36, negated_conjecture, (v2_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_37, negated_conjecture, (m1_pre_topc(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_28]), c_0_24]), c_0_25])])).
cnf(c_0_38, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_28]), c_0_24])])).
cnf(c_0_39, negated_conjecture, (v2_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_24]), c_0_25])])).
cnf(c_0_40, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~r1_tsep_1(X1,X2)|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r1_tarski(X3,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_43, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk4_0))|~m1_pre_topc(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36])])).
cnf(c_0_44, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk3_0))|~m1_pre_topc(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_37]), c_0_38]), c_0_39])])).
cnf(c_0_45, plain, (~l1_struct_0(X1)|~l1_struct_0(X2)|~r2_hidden(X3,u1_struct_0(X1))|~r2_hidden(X3,u1_struct_0(X2))|~r1_tsep_1(X2,X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_47, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk4_0))|~r2_hidden(X1,u1_struct_0(X2))|~m1_pre_topc(X2,esk4_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, plain, (r2_hidden(esk5_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_49, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk3_0))|~r2_hidden(X1,u1_struct_0(X2))|~m1_pre_topc(X2,esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_44])).
fof(c_0_50, plain, ![X17, X18]:(~l1_struct_0(X17)|~l1_struct_0(X18)|(~r1_tsep_1(X17,X18)|r1_tsep_1(X18,X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_tsep_1])])).
cnf(c_0_51, plain, (r1_xboole_0(u1_struct_0(X1),X2)|~l1_struct_0(X1)|~l1_struct_0(X3)|~r2_hidden(esk5_2(u1_struct_0(X1),X2),u1_struct_0(X3))|~r1_tsep_1(X3,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_52, negated_conjecture, (r1_xboole_0(X1,u1_struct_0(X2))|r2_hidden(esk5_2(X1,u1_struct_0(X2)),u1_struct_0(esk4_0))|~m1_pre_topc(X2,esk4_0)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53, negated_conjecture, (r1_xboole_0(X1,u1_struct_0(X2))|r2_hidden(esk5_2(X1,u1_struct_0(X2)),u1_struct_0(esk3_0))|~m1_pre_topc(X2,esk3_0)), inference(spm,[status(thm)],[c_0_49, c_0_48])).
cnf(c_0_54, plain, (r1_tsep_1(X2,X1)|~l1_struct_0(X1)|~l1_struct_0(X2)|~r1_tsep_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_55, negated_conjecture, (r1_tsep_1(esk3_0,esk4_0)|r1_tsep_1(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_56, plain, ![X12]:(~l1_pre_topc(X12)|l1_struct_0(X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_57, plain, (r1_tsep_1(X1,X2)|~r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~l1_struct_0(X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_58, negated_conjecture, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~l1_struct_0(esk4_0)|~l1_struct_0(X1)|~r1_tsep_1(esk4_0,X1)|~m1_pre_topc(X2,esk4_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59, negated_conjecture, (r1_xboole_0(u1_struct_0(X1),u1_struct_0(X2))|~l1_struct_0(esk3_0)|~l1_struct_0(X1)|~r1_tsep_1(esk3_0,X1)|~m1_pre_topc(X2,esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_53])).
cnf(c_0_60, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)|~l1_struct_0(esk4_0)|~l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_61, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_62, negated_conjecture, (r1_tsep_1(X1,X2)|~l1_struct_0(esk4_0)|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_tsep_1(esk4_0,X1)|~m1_pre_topc(X2,esk4_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_63, negated_conjecture, (r1_tsep_1(X1,X2)|~l1_struct_0(esk3_0)|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_tsep_1(esk3_0,X1)|~m1_pre_topc(X2,esk3_0)), inference(spm,[status(thm)],[c_0_57, c_0_59])).
cnf(c_0_64, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)|~l1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_38])])).
cnf(c_0_65, negated_conjecture, (r1_tsep_1(X1,X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_tsep_1(esk4_0,X1)|~m1_pre_topc(X2,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_61]), c_0_35])])).
cnf(c_0_66, negated_conjecture, (r1_tsep_1(X1,X2)|~l1_struct_0(X2)|~l1_struct_0(X1)|~r1_tsep_1(esk3_0,X1)|~m1_pre_topc(X2,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_61]), c_0_38])])).
cnf(c_0_67, negated_conjecture, (r1_tsep_1(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_61]), c_0_35])])).
cnf(c_0_68, negated_conjecture, (r1_tsep_1(X1,X2)|~l1_struct_0(X1)|~r1_tsep_1(esk4_0,X1)|~m1_pre_topc(X2,esk4_0)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_65, c_0_61])).
cnf(c_0_69, negated_conjecture, (r1_tsep_1(X1,X2)|~l1_struct_0(X1)|~r1_tsep_1(esk3_0,X1)|~m1_pre_topc(X2,esk3_0)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_66, c_0_61])).
cnf(c_0_70, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_71, negated_conjecture, (r1_tsep_1(esk3_0,esk4_0)|~l1_struct_0(esk3_0)|~l1_struct_0(esk4_0)), inference(spm,[status(thm)],[c_0_54, c_0_67])).
cnf(c_0_72, negated_conjecture, (r1_tsep_1(X1,X2)|~r1_tsep_1(esk4_0,X1)|~m1_pre_topc(X2,esk4_0)|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_68, c_0_61])).
cnf(c_0_73, negated_conjecture, (r1_tsep_1(X1,X2)|~r1_tsep_1(esk3_0,X1)|~m1_pre_topc(X2,esk3_0)|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_69, c_0_61])).
cnf(c_0_74, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_70]), c_0_24])])).
cnf(c_0_75, negated_conjecture, (m1_pre_topc(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_76, negated_conjecture, (r1_tsep_1(esk3_0,esk4_0)|~l1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_61]), c_0_38])])).
cnf(c_0_77, negated_conjecture, (r1_tsep_1(X1,esk4_0)|~r1_tsep_1(esk4_0,X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_35]), c_0_34])])).
cnf(c_0_78, negated_conjecture, (~r1_tsep_1(esk2_0,esk4_0)|~r1_tsep_1(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_79, negated_conjecture, (r1_tsep_1(X1,esk2_0)|~r1_tsep_1(esk3_0,X1)|~l1_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])])).
cnf(c_0_80, negated_conjecture, (r1_tsep_1(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_61]), c_0_35])])).
cnf(c_0_81, negated_conjecture, (~r1_tsep_1(esk4_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_74]), c_0_78])).
cnf(c_0_82, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_35]), c_0_80])]), c_0_81]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 60
# Proof object formula steps           : 23
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 37
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 844
# ...of these trivial                  : 3
# ...subsumed                          : 378
# ...remaining for further processing  : 463
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 138
# Backward-rewritten                   : 3
# Generated clauses                    : 1583
# ...of the previous two non-trivial   : 1429
# Contextual simplify-reflections      : 10
# Paramodulations                      : 1581
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 292
#    Positive orientable unit clauses  : 18
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 263
# Current number of unprocessed clauses: 519
# ...number of literals in the above   : 3266
# Current number of archived formulas  : 0
# Current number of archived clauses   : 169
# Clause-clause subsumption calls (NU) : 20678
# Rec. Clause-clause subsumption calls : 7231
# Non-unit clause-clause subsumptions  : 330
# Unit Clause-clause subsumption calls : 465
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 30788

# -------------------------------------------------
# User time                : 0.033 s
# System time              : 0.003 s
# Total time               : 0.036 s
# Maximum resident set size: 28852 pages
