# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/94fb2e0ad6194373b33cc87d48c2415f.E_input', t3_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/94fb2e0ad6194373b33cc87d48c2415f.E_input', d4_xboole_0)).
fof(t117_zfmisc_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/94fb2e0ad6194373b33cc87d48c2415f.E_input', t117_zfmisc_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/94fb2e0ad6194373b33cc87d48c2415f.E_input', l54_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/94fb2e0ad6194373b33cc87d48c2415f.E_input', d3_tarski)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/94fb2e0ad6194373b33cc87d48c2415f.E_input', t65_xboole_1)).
fof(c_0_6, plain, ![X5, X6, X8, X9, X10]:(((r2_hidden(esk1_2(X5,X6),X5)|r1_xboole_0(X5,X6))&(r2_hidden(esk1_2(X5,X6),X6)|r1_xboole_0(X5,X6)))&(~r2_hidden(X10,X8)|~r2_hidden(X10,X9)|~r1_xboole_0(X8,X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_7, plain, ![X18, X19, X20, X21, X22, X23, X24, X25]:((((r2_hidden(X21,X18)|~r2_hidden(X21,X20)|X20!=k3_xboole_0(X18,X19))&(r2_hidden(X21,X19)|~r2_hidden(X21,X20)|X20!=k3_xboole_0(X18,X19)))&(~r2_hidden(X22,X18)|~r2_hidden(X22,X19)|r2_hidden(X22,X20)|X20!=k3_xboole_0(X18,X19)))&((~r2_hidden(esk3_3(X23,X24,X25),X25)|(~r2_hidden(esk3_3(X23,X24,X25),X23)|~r2_hidden(esk3_3(X23,X24,X25),X24))|X25=k3_xboole_0(X23,X24))&((r2_hidden(esk3_3(X23,X24,X25),X23)|r2_hidden(esk3_3(X23,X24,X25),X25)|X25=k3_xboole_0(X23,X24))&(r2_hidden(esk3_3(X23,X24,X25),X24)|r2_hidden(esk3_3(X23,X24,X25),X25)|X25=k3_xboole_0(X23,X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_8, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3)))))), inference(assume_negation,[status(cth)],[t117_zfmisc_1])).
fof(c_0_11, plain, ![X30, X31, X32, X33]:(((r2_hidden(X30,X32)|~r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(X32,X33)))&(r2_hidden(X31,X33)|~r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(X32,X33))))&(~r2_hidden(X30,X32)|~r2_hidden(X31,X33)|r2_hidden(k4_tarski(X30,X31),k2_zfmisc_1(X32,X33)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_12, plain, ![X12, X13, X14, X15, X16]:((~r1_tarski(X12,X13)|(~r2_hidden(X14,X12)|r2_hidden(X14,X13)))&((r2_hidden(esk2_2(X15,X16),X15)|r1_tarski(X15,X16))&(~r2_hidden(esk2_2(X15,X16),X16)|r1_tarski(X15,X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_13, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk1_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_15, plain, (r2_hidden(esk1_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_16, plain, ![X11]:r1_xboole_0(X11,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_17, negated_conjecture, (esk4_0!=k1_xboole_0&((r1_tarski(k2_zfmisc_1(esk5_0,esk4_0),k2_zfmisc_1(esk6_0,esk4_0))|r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0)))&~r1_tarski(esk5_0,esk6_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (X1=k3_xboole_0(X2,X3)|r2_hidden(esk3_3(X2,X3,X1),X3)|~r2_hidden(esk3_3(X2,X3,X1),X4)|~r1_xboole_0(X4,X1)), inference(spm,[status(thm)],[c_0_8, c_0_13])).
cnf(c_0_21, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk5_0,esk4_0),k2_zfmisc_1(esk6_0,esk4_0))|r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|r2_hidden(k4_tarski(X3,esk2_2(X1,X2)),k2_zfmisc_1(X4,X1))|~r2_hidden(X3,X4)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_26, plain, (X1=k3_xboole_0(X2,X3)|r2_hidden(esk3_3(X2,X3,X1),X3)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_20, c_0_13])).
cnf(c_0_27, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))|r2_hidden(X1,k2_zfmisc_1(esk6_0,esk4_0))|~r2_hidden(X1,k2_zfmisc_1(esk5_0,esk4_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, plain, (r1_tarski(X1,X2)|r1_tarski(X3,X4)|r2_hidden(k4_tarski(esk2_2(X1,X2),esk2_2(X3,X4)),k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_25, c_0_19])).
cnf(c_0_30, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk3_3(X1,X2,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_33, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))|r1_tarski(esk4_0,X1)|r1_tarski(esk5_0,X2)|r2_hidden(k4_tarski(esk2_2(esk5_0,X2),esk2_2(esk4_0,X1)),k2_zfmisc_1(esk6_0,esk4_0))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (r2_hidden(X1,k3_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r2_hidden(esk3_3(X1,X2,k1_xboole_0),X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_8, c_0_31])).
cnf(c_0_36, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_37, plain, (X3=k3_xboole_0(X1,X2)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~r2_hidden(esk3_3(X1,X2,X3),X1)|~r2_hidden(esk3_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_38, plain, (k3_xboole_0(X1,X2)=X2|r2_hidden(esk3_3(X1,X2,X2),X2)), inference(ef,[status(thm)],[c_0_13])).
cnf(c_0_39, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_40, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))|r1_tarski(esk5_0,X1)|r1_tarski(esk4_0,X2)|r2_hidden(esk2_2(esk5_0,X1),esk6_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41, negated_conjecture, (~r1_tarski(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_42, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk3_3(X1,X2,k1_xboole_0),k3_xboole_0(X3,X2))|~r2_hidden(esk3_3(X1,X2,k1_xboole_0),X3)), inference(spm,[status(thm)],[c_0_34, c_0_31])).
cnf(c_0_43, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk3_3(X1,X2,k1_xboole_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_27])])).
cnf(c_0_44, plain, (k3_xboole_0(X1,X2)=X2|~r2_hidden(esk3_3(X1,X2,X2),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_38])).
cnf(c_0_45, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))|r1_tarski(esk4_0,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_46, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk3_3(X1,X2,k1_xboole_0),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_44, c_0_38])).
cnf(c_0_48, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))|r2_hidden(X1,X2)|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_23, c_0_45])).
cnf(c_0_49, plain, (X1=k1_xboole_0|r2_hidden(esk3_3(X1,X1,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_50, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_51, plain, (k3_xboole_0(X1,X2)=X2|~r2_hidden(esk3_3(X1,X2,X2),X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_8, c_0_38])).
cnf(c_0_52, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))|r2_hidden(esk3_3(esk4_0,esk4_0,k1_xboole_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_53, plain, (X1=k3_xboole_0(X2,X3)|r2_hidden(esk3_3(X2,X3,X1),X1)|~r2_hidden(esk3_3(X2,X3,X1),X4)|~r1_xboole_0(X4,X3)), inference(spm,[status(thm)],[c_0_8, c_0_13])).
cnf(c_0_54, plain, (k3_xboole_0(X1,X2)=X2|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_51, c_0_38])).
cnf(c_0_55, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))|~r1_xboole_0(X1,esk4_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_52]), c_0_47]), c_0_50])).
cnf(c_0_56, plain, (X1=k3_xboole_0(X2,X3)|r2_hidden(esk3_3(X2,X3,X1),X1)|~r1_xboole_0(X3,X3)), inference(spm,[status(thm)],[c_0_53, c_0_13])).
cnf(c_0_57, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_54, c_0_27])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk4_0,esk5_0),k2_zfmisc_1(esk4_0,esk6_0))), inference(spm,[status(thm)],[c_0_55, c_0_27])).
cnf(c_0_59, plain, (X1=k1_xboole_0|r2_hidden(esk3_3(X2,k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_27]), c_0_57])).
cnf(c_0_60, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk4_0,esk6_0))|~r2_hidden(X1,k2_zfmisc_1(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_23, c_0_58])).
cnf(c_0_61, plain, (X1=k1_xboole_0|r1_tarski(X2,X3)|r2_hidden(k4_tarski(esk3_3(X4,k1_xboole_0,X1),esk2_2(X2,X3)),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_25, c_0_59])).
cnf(c_0_62, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_63, negated_conjecture, (r1_tarski(esk5_0,X1)|r2_hidden(k4_tarski(esk3_3(X2,k1_xboole_0,esk4_0),esk2_2(esk5_0,X1)),k2_zfmisc_1(esk4_0,esk6_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_50])).
cnf(c_0_64, negated_conjecture, (r1_tarski(esk5_0,X1)|r2_hidden(esk2_2(esk5_0,X1),esk6_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_65, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_64]), c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 66
# Proof object clause steps            : 53
# Proof object formula steps           : 13
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 19406
# ...of these trivial                  : 273
# ...subsumed                          : 15354
# ...remaining for further processing  : 3779
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 170
# Backward-rewritten                   : 1385
# Generated clauses                    : 447519
# ...of the previous two non-trivial   : 441168
# Contextual simplify-reflections      : 7
# Paramodulations                      : 447367
# Factorizations                       : 124
# Equation resolutions                 : 28
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2224
#    Positive orientable unit clauses  : 70
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 2152
# Current number of unprocessed clauses: 421199
# ...number of literals in the above   : 2332951
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1555
# Clause-clause subsumption calls (NU) : 1970000
# Rec. Clause-clause subsumption calls : 1043679
# Non-unit clause-clause subsumptions  : 15525
# Unit Clause-clause subsumption calls : 3889
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 441
# BW rewrite match successes           : 20
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 11389327

# -------------------------------------------------
# User time                : 5.998 s
# System time              : 0.136 s
# Total time               : 6.134 s
# Maximum resident set size: 320236 pages
