# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d17_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>![X3]:(X3=k4_orders_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>m2_orders_2(X4,X1,X2))))), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', d17_orders_2)).
fof(t87_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>~(k3_tarski(k4_orders_2(X1,X2))=k1_xboole_0))), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', t87_orders_2)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', d4_tarski)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', t2_xboole_1)).
fof(existence_m2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_orders_1(X2,k1_orders_1(u1_struct_0(X1))))=>?[X3]:m2_orders_2(X3,X1,X2)), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', existence_m2_orders_2)).
fof(d16_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>![X3]:((v6_orders_2(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>(m2_orders_2(X3,X1,X2)<=>(~(X3=k1_xboole_0)&(r2_wellord1(u1_orders_2(X1),X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>k1_funct_1(X2,k1_orders_2(X1,k3_orders_2(X1,X3,X4)))=X4)))))))), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', d16_orders_2)).
fof(dt_m2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_orders_1(X2,k1_orders_1(u1_struct_0(X1))))=>![X3]:(m2_orders_2(X3,X1,X2)=>(v6_orders_2(X3,X1)&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', dt_m2_orders_2)).
fof(t6_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6, X7]:((r2_hidden(X3,X4)&(r2_hidden(X4,X5)&(r2_hidden(X5,X6)&(r2_hidden(X6,X7)&r2_hidden(X7,X2)))))=>r1_xboole_0(X3,X1)))))), file('data/MPTP/proofs/9569accbb24941e0aad451c3e76afa73.E_input', t6_mcart_1)).
fof(c_0_9, plain, ![X32, X33, X34, X35, X36, X37]:(((~r2_hidden(X35,X34)|m2_orders_2(X35,X32,X33)|X34!=k4_orders_2(X32,X33)|~m1_orders_1(X33,k1_orders_1(u1_struct_0(X32)))|(v2_struct_0(X32)|(~v3_orders_2(X32)|(~v4_orders_2(X32)|(~v5_orders_2(X32)|~l1_orders_2(X32))))))&(~m2_orders_2(X36,X32,X33)|r2_hidden(X36,X34)|X34!=k4_orders_2(X32,X33)|~m1_orders_1(X33,k1_orders_1(u1_struct_0(X32)))|(v2_struct_0(X32)|(~v3_orders_2(X32)|(~v4_orders_2(X32)|(~v5_orders_2(X32)|~l1_orders_2(X32)))))))&((~r2_hidden(esk7_3(X32,X33,X37),X37)|~m2_orders_2(esk7_3(X32,X33,X37),X32,X33)|X37=k4_orders_2(X32,X33)|~m1_orders_1(X33,k1_orders_1(u1_struct_0(X32)))|(v2_struct_0(X32)|(~v3_orders_2(X32)|(~v4_orders_2(X32)|(~v5_orders_2(X32)|~l1_orders_2(X32))))))&(r2_hidden(esk7_3(X32,X33,X37),X37)|m2_orders_2(esk7_3(X32,X33,X37),X32,X33)|X37=k4_orders_2(X32,X33)|~m1_orders_1(X33,k1_orders_1(u1_struct_0(X32)))|(v2_struct_0(X32)|(~v3_orders_2(X32)|(~v4_orders_2(X32)|(~v5_orders_2(X32)|~l1_orders_2(X32)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_orders_2])])])])])])])).
fof(c_0_10, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))=>~(k3_tarski(k4_orders_2(X1,X2))=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t87_orders_2])).
fof(c_0_11, plain, ![X10, X11, X12, X14, X15, X16, X17, X19]:((((r2_hidden(X12,esk3_3(X10,X11,X12))|~r2_hidden(X12,X11)|X11!=k3_tarski(X10))&(r2_hidden(esk3_3(X10,X11,X12),X10)|~r2_hidden(X12,X11)|X11!=k3_tarski(X10)))&(~r2_hidden(X14,X15)|~r2_hidden(X15,X10)|r2_hidden(X14,X11)|X11!=k3_tarski(X10)))&((~r2_hidden(esk4_2(X16,X17),X17)|(~r2_hidden(esk4_2(X16,X17),X19)|~r2_hidden(X19,X16))|X17=k3_tarski(X16))&((r2_hidden(esk4_2(X16,X17),esk5_2(X16,X17))|r2_hidden(esk4_2(X16,X17),X17)|X17=k3_tarski(X16))&(r2_hidden(esk5_2(X16,X17),X16)|r2_hidden(esk4_2(X16,X17),X17)|X17=k3_tarski(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
cnf(c_0_12, plain, (r2_hidden(X1,X4)|v2_struct_0(X2)|~m2_orders_2(X1,X2,X3)|X4!=k4_orders_2(X2,X3)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_orders_1(esk2_0,k1_orders_1(u1_struct_0(esk1_0)))&k3_tarski(k4_orders_2(esk1_0,esk2_0))=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])).
fof(c_0_14, plain, ![X42, X43]:(~r2_hidden(X42,X43)|~r1_tarski(X43,X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_15, plain, ![X21]:r1_tarski(k1_xboole_0,X21), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_16, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|X4!=k3_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (r2_hidden(X1,k4_orders_2(X2,X3))|v2_struct_0(X2)|~m2_orders_2(X1,X2,X3)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (m1_orders_1(esk2_0,k1_orders_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_24, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_26, plain, ![X39, X40]:(v2_struct_0(X39)|(~v3_orders_2(X39)|(~v4_orders_2(X39)|(~v5_orders_2(X39)|~l1_orders_2(X39))))|~m1_orders_1(X40,k1_orders_1(u1_struct_0(X39)))|m2_orders_2(esk8_2(X39,X40),X39,X40)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[existence_m2_orders_2])])])])).
cnf(c_0_27, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_28, negated_conjecture, (r2_hidden(X1,k4_orders_2(esk1_0,esk2_0))|~m2_orders_2(X1,esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23])).
cnf(c_0_29, negated_conjecture, (k3_tarski(k4_orders_2(esk1_0,esk2_0))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, plain, (v2_struct_0(X1)|m2_orders_2(esk8_2(X1,X2),X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_32, plain, ![X44, X45, X46, X47]:(((X46!=k1_xboole_0|~m2_orders_2(X46,X44,X45)|(~v6_orders_2(X46,X44)|~m1_subset_1(X46,k1_zfmisc_1(u1_struct_0(X44))))|~m1_orders_1(X45,k1_orders_1(u1_struct_0(X44)))|(v2_struct_0(X44)|(~v3_orders_2(X44)|(~v4_orders_2(X44)|(~v5_orders_2(X44)|~l1_orders_2(X44))))))&((r2_wellord1(u1_orders_2(X44),X46)|~m2_orders_2(X46,X44,X45)|(~v6_orders_2(X46,X44)|~m1_subset_1(X46,k1_zfmisc_1(u1_struct_0(X44))))|~m1_orders_1(X45,k1_orders_1(u1_struct_0(X44)))|(v2_struct_0(X44)|(~v3_orders_2(X44)|(~v4_orders_2(X44)|(~v5_orders_2(X44)|~l1_orders_2(X44))))))&(~m1_subset_1(X47,u1_struct_0(X44))|(~r2_hidden(X47,X46)|k1_funct_1(X45,k1_orders_2(X44,k3_orders_2(X44,X46,X47)))=X47)|~m2_orders_2(X46,X44,X45)|(~v6_orders_2(X46,X44)|~m1_subset_1(X46,k1_zfmisc_1(u1_struct_0(X44))))|~m1_orders_1(X45,k1_orders_1(u1_struct_0(X44)))|(v2_struct_0(X44)|(~v3_orders_2(X44)|(~v4_orders_2(X44)|(~v5_orders_2(X44)|~l1_orders_2(X44))))))))&((m1_subset_1(esk9_3(X44,X45,X46),u1_struct_0(X44))|~r2_wellord1(u1_orders_2(X44),X46)|X46=k1_xboole_0|m2_orders_2(X46,X44,X45)|(~v6_orders_2(X46,X44)|~m1_subset_1(X46,k1_zfmisc_1(u1_struct_0(X44))))|~m1_orders_1(X45,k1_orders_1(u1_struct_0(X44)))|(v2_struct_0(X44)|(~v3_orders_2(X44)|(~v4_orders_2(X44)|(~v5_orders_2(X44)|~l1_orders_2(X44))))))&((r2_hidden(esk9_3(X44,X45,X46),X46)|~r2_wellord1(u1_orders_2(X44),X46)|X46=k1_xboole_0|m2_orders_2(X46,X44,X45)|(~v6_orders_2(X46,X44)|~m1_subset_1(X46,k1_zfmisc_1(u1_struct_0(X44))))|~m1_orders_1(X45,k1_orders_1(u1_struct_0(X44)))|(v2_struct_0(X44)|(~v3_orders_2(X44)|(~v4_orders_2(X44)|(~v5_orders_2(X44)|~l1_orders_2(X44))))))&(k1_funct_1(X45,k1_orders_2(X44,k3_orders_2(X44,X46,esk9_3(X44,X45,X46))))!=esk9_3(X44,X45,X46)|~r2_wellord1(u1_orders_2(X44),X46)|X46=k1_xboole_0|m2_orders_2(X46,X44,X45)|(~v6_orders_2(X46,X44)|~m1_subset_1(X46,k1_zfmisc_1(u1_struct_0(X44))))|~m1_orders_1(X45,k1_orders_1(u1_struct_0(X44)))|(v2_struct_0(X44)|(~v3_orders_2(X44)|(~v4_orders_2(X44)|(~v5_orders_2(X44)|~l1_orders_2(X44))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_orders_2])])])])])])).
fof(c_0_33, plain, ![X29, X30, X31]:((v6_orders_2(X31,X29)|~m2_orders_2(X31,X29,X30)|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_orders_1(X30,k1_orders_1(u1_struct_0(X29)))))&(m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X29)))|~m2_orders_2(X31,X29,X30)|(v2_struct_0(X29)|(~v3_orders_2(X29)|(~v4_orders_2(X29)|(~v5_orders_2(X29)|~l1_orders_2(X29))))|~m1_orders_1(X30,k1_orders_1(u1_struct_0(X29)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_orders_2])])])])])).
cnf(c_0_34, negated_conjecture, (~m2_orders_2(X1,esk1_0,esk2_0)|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_35, negated_conjecture, (m2_orders_2(esk8_2(esk1_0,esk2_0),esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_18]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23])).
fof(c_0_36, plain, ![X22, X24, X25, X26, X27, X28]:((r2_hidden(esk6_1(X22),X22)|X22=k1_xboole_0)&(~r2_hidden(X24,X25)|(~r2_hidden(X25,X26)|(~r2_hidden(X26,X27)|(~r2_hidden(X27,X28)|~r2_hidden(X28,esk6_1(X22)))))|r1_xboole_0(X24,X22)|X22=k1_xboole_0)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_mcart_1])])])])])])])).
cnf(c_0_37, plain, (v2_struct_0(X2)|X1!=k1_xboole_0|~m2_orders_2(X1,X2,X3)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (v6_orders_2(X1,X2)|v2_struct_0(X2)|~m2_orders_2(X1,X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, negated_conjecture, (~r2_hidden(X1,esk8_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (r2_hidden(esk6_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41, plain, (v2_struct_0(X1)|X2!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m2_orders_2(X2,X1,X3)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_42, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|v2_struct_0(X2)|~m2_orders_2(X1,X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_orders_1(X3,k1_orders_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, negated_conjecture, (esk8_2(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44, plain, (v2_struct_0(X1)|~m2_orders_2(k1_xboole_0,X1,X2)|~m1_orders_1(X2,k1_orders_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_41]), c_0_42])).
cnf(c_0_45, negated_conjecture, (m2_orders_2(k1_xboole_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_35, c_0_43])).
cnf(c_0_46, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_18]), c_0_45]), c_0_19]), c_0_20]), c_0_21]), c_0_22])]), c_0_23]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 28
# Proof object formula steps           : 19
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 435
# ...of these trivial                  : 0
# ...subsumed                          : 306
# ...remaining for further processing  : 129
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 4
# Generated clauses                    : 777
# ...of the previous two non-trivial   : 584
# Contextual simplify-reflections      : 8
# Paramodulations                      : 754
# Factorizations                       : 0
# Equation resolutions                 : 23
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 94
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 82
# Current number of unprocessed clauses: 206
# ...number of literals in the above   : 924
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 1777
# Rec. Clause-clause subsumption calls : 680
# Non-unit clause-clause subsumptions  : 157
# Unit Clause-clause subsumption calls : 58
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 13026

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.003 s
# Total time               : 0.018 s
# Maximum resident set size: 28616 pages
