# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t22_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))=>k1_funct_1(k5_relat_1(X3,X2),X1)=k1_funct_1(X2,k1_funct_1(X3,X1))))), file('data/MPTP/proofs/95aa6a1b96254973b340fd5d0dd8e156.E_input', t22_funct_1)).
fof(d8_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))&(r2_hidden(X3,k1_relat_1(X1))&k1_funct_1(X1,X2)=k1_funct_1(X1,X3)))=>X2=X3))), file('data/MPTP/proofs/95aa6a1b96254973b340fd5d0dd8e156.E_input', d8_funct_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/95aa6a1b96254973b340fd5d0dd8e156.E_input', dt_k5_relat_1)).
fof(fc2_funct_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&(v1_relat_1(X2)&v1_funct_1(X2)))=>(v1_relat_1(k5_relat_1(X1,X2))&v1_funct_1(k5_relat_1(X1,X2)))), file('data/MPTP/proofs/95aa6a1b96254973b340fd5d0dd8e156.E_input', fc2_funct_1)).
fof(t21_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))<=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X1),k1_relat_1(X2)))))), file('data/MPTP/proofs/95aa6a1b96254973b340fd5d0dd8e156.E_input', t21_funct_1)).
fof(t46_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X1)&v2_funct_1(X2))=>v2_funct_1(k5_relat_1(X1,X2))))), file('data/MPTP/proofs/95aa6a1b96254973b340fd5d0dd8e156.E_input', t46_funct_1)).
fof(c_0_6, plain, ![X4, X5, X6]:(~v1_relat_1(X5)|~v1_funct_1(X5)|(~v1_relat_1(X6)|~v1_funct_1(X6)|(~r2_hidden(X4,k1_relat_1(k5_relat_1(X6,X5)))|k1_funct_1(k5_relat_1(X6,X5),X4)=k1_funct_1(X5,k1_funct_1(X6,X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_funct_1])])])).
fof(c_0_7, plain, ![X7, X8, X9]:((~v2_funct_1(X7)|(~r2_hidden(X8,k1_relat_1(X7))|(~r2_hidden(X9,k1_relat_1(X7))|k1_funct_1(X7,X8)!=k1_funct_1(X7,X9))|X8=X9)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(((r2_hidden(esk1_1(X7),k1_relat_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&((r2_hidden(esk2_1(X7),k1_relat_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(k1_funct_1(X7,esk1_1(X7))=k1_funct_1(X7,esk2_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))))&(esk1_1(X7)!=esk2_1(X7)|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_funct_1])])])])])).
fof(c_0_8, plain, ![X12, X13]:(~v1_relat_1(X12)|~v1_relat_1(X13)|v1_relat_1(k5_relat_1(X12,X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
fof(c_0_9, plain, ![X14, X15]:((v1_relat_1(k5_relat_1(X14,X15))|(~v1_relat_1(X14)|~v1_funct_1(X14)|(~v1_relat_1(X15)|~v1_funct_1(X15))))&(v1_funct_1(k5_relat_1(X14,X15))|(~v1_relat_1(X14)|~v1_funct_1(X14)|(~v1_relat_1(X15)|~v1_funct_1(X15))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_funct_1])])])).
cnf(c_0_10, plain, (k1_funct_1(k5_relat_1(X2,X1),X3)=k1_funct_1(X1,k1_funct_1(X2,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X3,k1_relat_1(k5_relat_1(X2,X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (r2_hidden(esk2_1(X1),k1_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (v1_funct_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (r2_hidden(esk1_1(X1),k1_relat_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, plain, (k1_funct_1(X1,esk1_1(X1))=k1_funct_1(X1,esk2_1(X1))|v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, plain, (k1_funct_1(k5_relat_1(X1,X2),esk2_1(k5_relat_1(X1,X2)))=k1_funct_1(X2,k1_funct_1(X1,esk2_1(k5_relat_1(X1,X2))))|v2_funct_1(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])).
fof(c_0_17, plain, ![X16, X17, X18]:(((r2_hidden(X16,k1_relat_1(X18))|~r2_hidden(X16,k1_relat_1(k5_relat_1(X18,X17)))|(~v1_relat_1(X18)|~v1_funct_1(X18))|(~v1_relat_1(X17)|~v1_funct_1(X17)))&(r2_hidden(k1_funct_1(X18,X16),k1_relat_1(X17))|~r2_hidden(X16,k1_relat_1(k5_relat_1(X18,X17)))|(~v1_relat_1(X18)|~v1_funct_1(X18))|(~v1_relat_1(X17)|~v1_funct_1(X17))))&(~r2_hidden(X16,k1_relat_1(X18))|~r2_hidden(k1_funct_1(X18,X16),k1_relat_1(X17))|r2_hidden(X16,k1_relat_1(k5_relat_1(X18,X17)))|(~v1_relat_1(X18)|~v1_funct_1(X18))|(~v1_relat_1(X17)|~v1_funct_1(X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_funct_1])])])])).
cnf(c_0_18, plain, (k1_funct_1(k5_relat_1(X1,X2),esk1_1(k5_relat_1(X1,X2)))=k1_funct_1(X2,k1_funct_1(X1,esk1_1(k5_relat_1(X1,X2))))|v2_funct_1(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_14]), c_0_12]), c_0_13])).
cnf(c_0_19, plain, (k1_funct_1(k5_relat_1(X1,X2),esk1_1(k5_relat_1(X1,X2)))=k1_funct_1(X2,k1_funct_1(X1,esk2_1(k5_relat_1(X1,X2))))|v2_funct_1(k5_relat_1(X1,X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_12]), c_0_13])).
cnf(c_0_20, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relat_1(X3))|~r2_hidden(X2,k1_relat_1(k5_relat_1(X1,X3)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (X2=X3|~v2_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))|~r2_hidden(X3,k1_relat_1(X1))|k1_funct_1(X1,X2)!=k1_funct_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, plain, (k1_funct_1(X1,k1_funct_1(X2,esk2_1(k5_relat_1(X2,X1))))=k1_funct_1(X1,k1_funct_1(X2,esk1_1(k5_relat_1(X2,X1))))|v2_funct_1(k5_relat_1(X2,X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, plain, (v2_funct_1(k5_relat_1(X1,X2))|r2_hidden(k1_funct_1(X1,esk2_1(k5_relat_1(X1,X2))),k1_relat_1(X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_11]), c_0_12]), c_0_13])).
cnf(c_0_24, plain, (k1_funct_1(X1,esk2_1(k5_relat_1(X1,X2)))=X3|v2_funct_1(k5_relat_1(X1,X2))|k1_funct_1(X2,k1_funct_1(X1,esk1_1(k5_relat_1(X1,X2))))!=k1_funct_1(X2,X3)|~v2_funct_1(X2)|~r2_hidden(X3,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])).
cnf(c_0_25, plain, (v2_funct_1(k5_relat_1(X1,X2))|r2_hidden(k1_funct_1(X1,esk1_1(k5_relat_1(X1,X2))),k1_relat_1(X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_14]), c_0_12]), c_0_13])).
cnf(c_0_26, plain, (k1_funct_1(X1,esk2_1(k5_relat_1(X1,X2)))=k1_funct_1(X3,esk1_1(k5_relat_1(X3,X2)))|v2_funct_1(k5_relat_1(X3,X2))|v2_funct_1(k5_relat_1(X1,X2))|k1_funct_1(X2,k1_funct_1(X1,esk1_1(k5_relat_1(X1,X2))))!=k1_funct_1(X2,k1_funct_1(X3,esk1_1(k5_relat_1(X3,X2))))|~v2_funct_1(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_27, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,k1_relat_1(k5_relat_1(X2,X3)))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (k1_funct_1(X1,esk2_1(k5_relat_1(X1,X2)))=k1_funct_1(X1,esk1_1(k5_relat_1(X1,X2)))|v2_funct_1(k5_relat_1(X1,X2))|~v2_funct_1(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_29, plain, (v2_funct_1(k5_relat_1(X1,X2))|r2_hidden(esk2_1(k5_relat_1(X1,X2)),k1_relat_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_11]), c_0_12]), c_0_13])).
cnf(c_0_30, plain, (esk2_1(k5_relat_1(X1,X2))=X3|v2_funct_1(k5_relat_1(X1,X2))|k1_funct_1(X1,esk1_1(k5_relat_1(X1,X2)))!=k1_funct_1(X1,X3)|~v2_funct_1(X1)|~v2_funct_1(X2)|~r2_hidden(X3,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_28]), c_0_29])).
cnf(c_0_31, plain, (v2_funct_1(k5_relat_1(X1,X2))|r2_hidden(esk1_1(k5_relat_1(X1,X2)),k1_relat_1(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_14]), c_0_12]), c_0_13])).
fof(c_0_32, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X1)&v2_funct_1(X2))=>v2_funct_1(k5_relat_1(X1,X2)))))), inference(assume_negation,[status(cth)],[t46_funct_1])).
cnf(c_0_33, plain, (esk2_1(k5_relat_1(X1,X2))=esk1_1(k5_relat_1(X1,X3))|v2_funct_1(k5_relat_1(X1,X3))|v2_funct_1(k5_relat_1(X1,X2))|k1_funct_1(X1,esk1_1(k5_relat_1(X1,X2)))!=k1_funct_1(X1,esk1_1(k5_relat_1(X1,X3)))|~v2_funct_1(X1)|~v2_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_34, negated_conjecture, ((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((v1_relat_1(esk4_0)&v1_funct_1(esk4_0))&((v2_funct_1(esk3_0)&v2_funct_1(esk4_0))&~v2_funct_1(k5_relat_1(esk3_0,esk4_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_32])])])).
cnf(c_0_35, plain, (v2_funct_1(X1)|esk1_1(X1)!=esk2_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_36, plain, (esk2_1(k5_relat_1(X1,X2))=esk1_1(k5_relat_1(X1,X2))|v2_funct_1(k5_relat_1(X1,X2))|~v2_funct_1(X1)|~v2_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_33])).
cnf(c_0_37, negated_conjecture, (~v2_funct_1(k5_relat_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_38, plain, (v2_funct_1(k5_relat_1(X1,X2))|~v2_funct_1(X1)|~v2_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_12]), c_0_13])).
cnf(c_0_39, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, negated_conjecture, (v2_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_45, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42]), c_0_43]), c_0_44])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 33
# Proof object formula steps           : 13
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 354
# ...of these trivial                  : 0
# ...subsumed                          : 126
# ...remaining for further processing  : 228
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 0
# Generated clauses                    : 1250
# ...of the previous two non-trivial   : 1226
# Contextual simplify-reflections      : 270
# Paramodulations                      : 1247
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 203
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 196
# Current number of unprocessed clauses: 908
# ...number of literals in the above   : 13381
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 61762
# Rec. Clause-clause subsumption calls : 988
# Non-unit clause-clause subsumptions  : 403
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 64432

# -------------------------------------------------
# User time                : 0.041 s
# System time              : 0.000 s
# Total time               : 0.041 s
# Maximum resident set size: 28752 pages
