# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_tops_2, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>v2_tops_2(k7_setfam_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/95d8cb6f1c4d4025a650b9813a340c67.E_input', t17_tops_2)).
fof(d8_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(X3=k7_setfam_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(r2_hidden(X4,X3)<=>r2_hidden(k3_subset_1(X1,X4),X2)))))), file('data/MPTP/proofs/95d8cb6f1c4d4025a650b9813a340c67.E_input', d8_setfam_1)).
fof(t49_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r2_hidden(k3_subset_1(X1,X3),k7_setfam_1(X1,X2))<=>r2_hidden(X3,X2)))), file('data/MPTP/proofs/95d8cb6f1c4d4025a650b9813a340c67.E_input', t49_setfam_1)).
fof(dt_k7_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>m1_subset_1(k7_setfam_1(X1,X2),k1_zfmisc_1(k1_zfmisc_1(X1)))), file('data/MPTP/proofs/95d8cb6f1c4d4025a650b9813a340c67.E_input', dt_k7_setfam_1)).
fof(t16_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v2_tops_2(X2,X1)<=>v1_tops_2(k7_setfam_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/95d8cb6f1c4d4025a650b9813a340c67.E_input', t16_tops_2)).
fof(c_0_5, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>v2_tops_2(k7_setfam_1(u1_struct_0(X1),X2),X1))))), inference(assume_negation,[status(cth)],[t17_tops_2])).
fof(c_0_6, plain, ![X14, X15, X16, X17]:(((~r2_hidden(X17,X16)|r2_hidden(k3_subset_1(X14,X17),X15)|~m1_subset_1(X17,k1_zfmisc_1(X14))|X16!=k7_setfam_1(X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k1_zfmisc_1(X14)))|~m1_subset_1(X15,k1_zfmisc_1(k1_zfmisc_1(X14))))&(~r2_hidden(k3_subset_1(X14,X17),X15)|r2_hidden(X17,X16)|~m1_subset_1(X17,k1_zfmisc_1(X14))|X16!=k7_setfam_1(X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k1_zfmisc_1(X14)))|~m1_subset_1(X15,k1_zfmisc_1(k1_zfmisc_1(X14)))))&((m1_subset_1(esk3_3(X14,X15,X16),k1_zfmisc_1(X14))|X16=k7_setfam_1(X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k1_zfmisc_1(X14)))|~m1_subset_1(X15,k1_zfmisc_1(k1_zfmisc_1(X14))))&((~r2_hidden(esk3_3(X14,X15,X16),X16)|~r2_hidden(k3_subset_1(X14,esk3_3(X14,X15,X16)),X15)|X16=k7_setfam_1(X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k1_zfmisc_1(X14)))|~m1_subset_1(X15,k1_zfmisc_1(k1_zfmisc_1(X14))))&(r2_hidden(esk3_3(X14,X15,X16),X16)|r2_hidden(k3_subset_1(X14,esk3_3(X14,X15,X16)),X15)|X16=k7_setfam_1(X14,X15)|~m1_subset_1(X16,k1_zfmisc_1(k1_zfmisc_1(X14)))|~m1_subset_1(X15,k1_zfmisc_1(k1_zfmisc_1(X14))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_setfam_1])])])])])).
fof(c_0_7, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))&((~v1_tops_2(esk2_0,esk1_0)|~v2_tops_2(k7_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0))&(v1_tops_2(esk2_0,esk1_0)|v2_tops_2(k7_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
fof(c_0_8, plain, ![X7, X8, X9]:((~r2_hidden(k3_subset_1(X7,X9),k7_setfam_1(X7,X8))|r2_hidden(X9,X8)|~m1_subset_1(X9,k1_zfmisc_1(X7))|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(X7))))&(~r2_hidden(X9,X8)|r2_hidden(k3_subset_1(X7,X9),k7_setfam_1(X7,X8))|~m1_subset_1(X9,k1_zfmisc_1(X7))|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_setfam_1])])])])).
cnf(c_0_9, plain, (r2_hidden(esk3_3(X1,X2,X3),X3)|r2_hidden(k3_subset_1(X1,esk3_3(X1,X2,X3)),X2)|X3=k7_setfam_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, ![X12, X13]:(~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(X12)))|m1_subset_1(k7_setfam_1(X12,X13),k1_zfmisc_1(k1_zfmisc_1(X12)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_setfam_1])])).
cnf(c_0_12, plain, (r2_hidden(X2,X3)|~r2_hidden(k3_subset_1(X1,X2),k7_setfam_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (k7_setfam_1(u1_struct_0(esk1_0),X1)=esk2_0|r2_hidden(k3_subset_1(u1_struct_0(esk1_0),esk3_3(u1_struct_0(esk1_0),X1,esk2_0)),X1)|r2_hidden(esk3_3(u1_struct_0(esk1_0),X1,esk2_0),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_14, plain, (m1_subset_1(k7_setfam_1(X2,X1),k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (X3=k7_setfam_1(X1,X2)|~r2_hidden(esk3_3(X1,X2,X3),X3)|~r2_hidden(k3_subset_1(X1,esk3_3(X1,X2,X3)),X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, plain, (r2_hidden(k3_subset_1(X3,X1),k7_setfam_1(X3,X2))|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, negated_conjecture, (k7_setfam_1(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),X1))=esk2_0|r2_hidden(esk3_3(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),X1),esk2_0),esk2_0)|r2_hidden(esk3_3(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),X1),esk2_0),X1)|~m1_subset_1(esk3_3(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),X1),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
cnf(c_0_18, plain, (m1_subset_1(esk3_3(X1,X2,X3),k1_zfmisc_1(X1))|X3=k7_setfam_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19, plain, (X1=k7_setfam_1(X2,k7_setfam_1(X2,X3))|~r2_hidden(esk3_3(X2,k7_setfam_1(X2,X3),X1),X1)|~r2_hidden(esk3_3(X2,k7_setfam_1(X2,X3),X1),X3)|~m1_subset_1(esk3_3(X2,k7_setfam_1(X2,X3),X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_14])).
cnf(c_0_20, negated_conjecture, (k7_setfam_1(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),X1))=esk2_0|r2_hidden(esk3_3(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),X1),esk2_0),esk2_0)|r2_hidden(esk3_3(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),X1),esk2_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_10])]), c_0_14])).
fof(c_0_21, plain, ![X10, X11]:((~v2_tops_2(X11,X10)|v1_tops_2(k7_setfam_1(u1_struct_0(X10),X11),X10)|~m1_subset_1(X11,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~l1_pre_topc(X10))&(~v1_tops_2(k7_setfam_1(u1_struct_0(X10),X11),X10)|v2_tops_2(X11,X10)|~m1_subset_1(X11,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X10))))|~l1_pre_topc(X10))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_tops_2])])])])).
cnf(c_0_22, plain, (X1=k7_setfam_1(X2,k7_setfam_1(X2,X3))|~r2_hidden(esk3_3(X2,k7_setfam_1(X2,X3),X1),X1)|~r2_hidden(esk3_3(X2,k7_setfam_1(X2,X3),X1),X3)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_18]), c_0_14])).
cnf(c_0_23, negated_conjecture, (k7_setfam_1(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),esk2_0))=esk2_0|r2_hidden(esk3_3(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),esk2_0),esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_20, c_0_10])).
cnf(c_0_24, plain, (v1_tops_2(k7_setfam_1(u1_struct_0(X2),X1),X2)|~v2_tops_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, negated_conjecture, (k7_setfam_1(u1_struct_0(esk1_0),k7_setfam_1(u1_struct_0(esk1_0),esk2_0))=esk2_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_10])]), c_0_23])).
cnf(c_0_26, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_27, negated_conjecture, (v1_tops_2(esk2_0,esk1_0)|v2_tops_2(k7_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_28, negated_conjecture, (v1_tops_2(esk2_0,esk1_0)|~m1_subset_1(k7_setfam_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_29, negated_conjecture, (~v1_tops_2(esk2_0,esk1_0)|~v2_tops_2(k7_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30, negated_conjecture, (v1_tops_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_14]), c_0_10])])).
cnf(c_0_31, plain, (v2_tops_2(X2,X1)|~v1_tops_2(k7_setfam_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, negated_conjecture, (~v2_tops_2(k7_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30])])).
cnf(c_0_33, negated_conjecture, (~m1_subset_1(k7_setfam_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_25]), c_0_26])]), c_0_30])]), c_0_32])).
cnf(c_0_34, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_14]), c_0_10])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 35
# Proof object clause steps            : 24
# Proof object formula steps           : 11
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 14
# Processed clauses                    : 57
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 52
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 4
# Generated clauses                    : 78
# ...of the previous two non-trivial   : 61
# Contextual simplify-reflections      : 12
# Paramodulations                      : 76
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 24
# Current number of unprocessed clauses: 32
# ...number of literals in the above   : 195
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 434
# Rec. Clause-clause subsumption calls : 51
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5004

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28660 pages
