# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t53_relset_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k8_relset_1(X1,X3,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X5,X6),X4)&r2_hidden(X6,X2))))))))), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', t53_relset_1)).
fof(redefinition_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k8_relset_1(X1,X2,X3,X4)=k10_relat_1(X3,X4)), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', redefinition_k8_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', t3_subset)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', d19_relat_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', fc6_relat_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', t4_subset)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', t166_relat_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', cc2_relset_1)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/95e5234c7b0b4255841ecd6cdaa27927.E_input', d14_relat_1)).
fof(c_0_10, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))=>![X5]:(m1_subset_1(X5,X1)=>(r2_hidden(X5,k8_relset_1(X1,X3,X4,X2))<=>?[X6]:(m1_subset_1(X6,X3)&(r2_hidden(k4_tarski(X5,X6),X4)&r2_hidden(X6,X2)))))))))), inference(assume_negation,[status(cth)],[t53_relset_1])).
fof(c_0_11, plain, ![X14, X15, X16, X17]:(~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))|k8_relset_1(X14,X15,X16,X17)=k10_relat_1(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_relset_1])])).
fof(c_0_12, negated_conjecture, ![X12]:(~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&(~v1_xboole_0(esk3_0)&(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk3_0)))&(m1_subset_1(esk5_0,esk1_0)&((~r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))|(~m1_subset_1(X12,esk3_0)|(~r2_hidden(k4_tarski(esk5_0,X12),esk4_0)|~r2_hidden(X12,esk2_0))))&((m1_subset_1(esk6_0,esk3_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0)))&((r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0)))&(r2_hidden(esk6_0,esk2_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])])])).
fof(c_0_13, plain, ![X47, X48]:((~m1_subset_1(X47,k1_zfmisc_1(X48))|r1_tarski(X47,X48))&(~r1_tarski(X47,X48)|m1_subset_1(X47,k1_zfmisc_1(X48)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_14, plain, ![X26, X27]:((~v5_relat_1(X27,X26)|r1_tarski(k2_relat_1(X27),X26)|~v1_relat_1(X27))&(~r1_tarski(k2_relat_1(X27),X26)|v5_relat_1(X27,X26)|~v1_relat_1(X27))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
cnf(c_0_15, plain, (k8_relset_1(X2,X3,X1,X4)=k10_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, ![X28, X29]:(~v1_relat_1(X28)|(~m1_subset_1(X29,k1_zfmisc_1(X28))|v1_relat_1(X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_18, plain, ![X33, X34]:v1_relat_1(k2_zfmisc_1(X33,X34)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_19, plain, ![X23, X24, X25]:(~r2_hidden(X23,X24)|~m1_subset_1(X24,k1_zfmisc_1(X25))|m1_subset_1(X23,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_20, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (~r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))|~m1_subset_1(X1,esk3_0)|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (k8_relset_1(esk1_0,esk3_0,esk4_0,X1)=k10_relat_1(esk4_0,X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_24, plain, ![X18, X19, X20, X22]:(((r2_hidden(esk7_3(X18,X19,X20),k2_relat_1(X20))|~r2_hidden(X18,k10_relat_1(X20,X19))|~v1_relat_1(X20))&((r2_hidden(k4_tarski(X18,esk7_3(X18,X19,X20)),X20)|~r2_hidden(X18,k10_relat_1(X20,X19))|~v1_relat_1(X20))&(r2_hidden(esk7_3(X18,X19,X20),X19)|~r2_hidden(X18,k10_relat_1(X20,X19))|~v1_relat_1(X20))))&(~r2_hidden(X22,k2_relat_1(X20))|(~r2_hidden(k4_tarski(X18,X22),X20)|~r2_hidden(X22,X19))|r2_hidden(X18,k10_relat_1(X20,X19))|~v1_relat_1(X20))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
cnf(c_0_25, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (m1_subset_1(k2_relat_1(X1),k1_zfmisc_1(X2))|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_29, plain, ![X30, X31, X32]:((v4_relat_1(X32,X30)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X30,X31))))&(v5_relat_1(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X30,X31))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_30, negated_conjecture, (~r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|~r2_hidden(k4_tarski(esk5_0,X1),esk4_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk3_0)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, plain, (r2_hidden(k4_tarski(X1,esk7_3(X1,X2,X3)),X3)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, negated_conjecture, (v1_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_16]), c_0_26])])).
cnf(c_0_33, plain, (m1_subset_1(X1,X2)|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k2_relat_1(X3))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, plain, (r2_hidden(esk7_3(X1,X2,X3),k2_relat_1(X3))|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, plain, ![X35, X36, X37, X38, X40, X41, X42, X43, X45]:((((r2_hidden(k4_tarski(X38,esk8_4(X35,X36,X37,X38)),X35)|~r2_hidden(X38,X37)|X37!=k10_relat_1(X35,X36)|~v1_relat_1(X35))&(r2_hidden(esk8_4(X35,X36,X37,X38),X36)|~r2_hidden(X38,X37)|X37!=k10_relat_1(X35,X36)|~v1_relat_1(X35)))&(~r2_hidden(k4_tarski(X40,X41),X35)|~r2_hidden(X41,X36)|r2_hidden(X40,X37)|X37!=k10_relat_1(X35,X36)|~v1_relat_1(X35)))&((~r2_hidden(esk9_3(X35,X42,X43),X43)|(~r2_hidden(k4_tarski(esk9_3(X35,X42,X43),X45),X35)|~r2_hidden(X45,X42))|X43=k10_relat_1(X35,X42)|~v1_relat_1(X35))&((r2_hidden(k4_tarski(esk9_3(X35,X42,X43),esk10_3(X35,X42,X43)),X35)|r2_hidden(esk9_3(X35,X42,X43),X43)|X43=k10_relat_1(X35,X42)|~v1_relat_1(X35))&(r2_hidden(esk10_3(X35,X42,X43),X42)|r2_hidden(esk9_3(X35,X42,X43),X43)|X43=k10_relat_1(X35,X42)|~v1_relat_1(X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
cnf(c_0_37, negated_conjecture, (~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~m1_subset_1(esk7_3(esk5_0,X1,esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
cnf(c_0_38, plain, (m1_subset_1(esk7_3(X1,X2,X3),X4)|~v5_relat_1(X3,X4)|~v1_relat_1(X3)|~r2_hidden(X1,k10_relat_1(X3,X2))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, negated_conjecture, (v5_relat_1(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_35, c_0_16])).
cnf(c_0_40, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk6_0,esk2_0)|r2_hidden(esk5_0,k8_relset_1(esk1_0,esk3_0,esk4_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_43, negated_conjecture, (~r2_hidden(esk7_3(esk5_0,X1,esk4_0),esk2_0)|~r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|~r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_32])])).
cnf(c_0_44, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_45, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_46, negated_conjecture, (r2_hidden(k4_tarski(esk5_0,esk6_0),esk4_0)|r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))), inference(rw,[status(thm)],[c_0_41, c_0_23])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|r2_hidden(esk6_0,esk2_0)), inference(rw,[status(thm)],[c_0_42, c_0_23])).
cnf(c_0_48, negated_conjecture, (~r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_32])])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk5_0,k10_relat_1(esk4_0,esk2_0))|r2_hidden(esk5_0,k10_relat_1(esk4_0,X1))|~r2_hidden(esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_32])])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk6_0,esk2_0)), inference(sr,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 31
# Proof object formula steps           : 21
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 98
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 94
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 7
# Generated clauses                    : 77
# ...of the previous two non-trivial   : 72
# Contextual simplify-reflections      : 0
# Paramodulations                      : 71
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 52
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 40
# Current number of unprocessed clauses: 31
# ...number of literals in the above   : 124
# Current number of archived formulas  : 0
# Current number of archived clauses   : 39
# Clause-clause subsumption calls (NU) : 1299
# Rec. Clause-clause subsumption calls : 808
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 20
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3827

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28792 pages
