# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l28_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(X4=k11_lattice3(X1,X2,X3)<=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))))))), file('data/MPTP/proofs/962405eb210e411b965fe65d9381ebe0.E_input', l28_lattice3)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/962405eb210e411b965fe65d9381ebe0.E_input', cc2_lattice3)).
fof(d11_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))))))), file('data/MPTP/proofs/962405eb210e411b965fe65d9381ebe0.E_input', d11_lattice3)).
fof(t15_lattice3, conjecture, ![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2)))), file('data/MPTP/proofs/962405eb210e411b965fe65d9381ebe0.E_input', t15_lattice3)).
fof(c_0_4, plain, ![X6, X7, X8, X9, X10]:(((r1_orders_2(X6,X9,X7)|X9!=k11_lattice3(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&((r1_orders_2(X6,X9,X8)|X9!=k11_lattice3(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(~m1_subset_1(X10,u1_struct_0(X6))|(~r1_orders_2(X6,X10,X7)|~r1_orders_2(X6,X10,X8)|r1_orders_2(X6,X10,X9))|X9!=k11_lattice3(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))))&((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r1_orders_2(X6,X9,X8)|~r1_orders_2(X6,X9,X7)|X9=k11_lattice3(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(((r1_orders_2(X6,esk1_4(X6,X7,X8,X9),X7)|~r1_orders_2(X6,X9,X8)|~r1_orders_2(X6,X9,X7)|X9=k11_lattice3(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))&(r1_orders_2(X6,esk1_4(X6,X7,X8,X9),X8)|~r1_orders_2(X6,X9,X8)|~r1_orders_2(X6,X9,X7)|X9=k11_lattice3(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6))))))&(~r1_orders_2(X6,esk1_4(X6,X7,X8,X9),X9)|~r1_orders_2(X6,X9,X8)|~r1_orders_2(X6,X9,X7)|X9=k11_lattice3(X6,X7,X8)|~m1_subset_1(X9,u1_struct_0(X6))|~m1_subset_1(X8,u1_struct_0(X6))|~m1_subset_1(X7,u1_struct_0(X6))|(v2_struct_0(X6)|(~v5_orders_2(X6)|(~v2_lattice3(X6)|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l28_lattice3])])])])])])).
fof(c_0_5, plain, ![X12]:(~l1_orders_2(X12)|(~v2_lattice3(X12)|~v2_struct_0(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
cnf(c_0_6, plain, (X4=k11_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,esk1_4(X1,X2,X3,X4),X4)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, ![X13, X14, X15, X17, X20]:(((m1_subset_1(esk2_3(X13,X14,X15),u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|~v2_lattice3(X13)|~l1_orders_2(X13))&((r1_orders_2(X13,esk2_3(X13,X14,X15),X14)|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|~v2_lattice3(X13)|~l1_orders_2(X13))&((r1_orders_2(X13,esk2_3(X13,X14,X15),X15)|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|~v2_lattice3(X13)|~l1_orders_2(X13))&(~m1_subset_1(X17,u1_struct_0(X13))|(~r1_orders_2(X13,X17,X14)|~r1_orders_2(X13,X17,X15)|r1_orders_2(X13,X17,esk2_3(X13,X14,X15)))|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|~v2_lattice3(X13)|~l1_orders_2(X13)))))&((m1_subset_1(esk3_1(X13),u1_struct_0(X13))|v2_lattice3(X13)|~l1_orders_2(X13))&((m1_subset_1(esk4_1(X13),u1_struct_0(X13))|v2_lattice3(X13)|~l1_orders_2(X13))&((m1_subset_1(esk5_2(X13,X20),u1_struct_0(X13))|~r1_orders_2(X13,X20,esk4_1(X13))|~r1_orders_2(X13,X20,esk3_1(X13))|~m1_subset_1(X20,u1_struct_0(X13))|v2_lattice3(X13)|~l1_orders_2(X13))&(((r1_orders_2(X13,esk5_2(X13,X20),esk3_1(X13))|~r1_orders_2(X13,X20,esk4_1(X13))|~r1_orders_2(X13,X20,esk3_1(X13))|~m1_subset_1(X20,u1_struct_0(X13))|v2_lattice3(X13)|~l1_orders_2(X13))&(r1_orders_2(X13,esk5_2(X13,X20),esk4_1(X13))|~r1_orders_2(X13,X20,esk4_1(X13))|~r1_orders_2(X13,X20,esk3_1(X13))|~m1_subset_1(X20,u1_struct_0(X13))|v2_lattice3(X13)|~l1_orders_2(X13)))&(~r1_orders_2(X13,esk5_2(X13,X20),X20)|~r1_orders_2(X13,X20,esk4_1(X13))|~r1_orders_2(X13,X20,esk3_1(X13))|~m1_subset_1(X20,u1_struct_0(X13))|v2_lattice3(X13)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_lattice3])])])])])).
cnf(c_0_9, plain, (X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk1_4(X2,X3,X4,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10, plain, (r1_orders_2(X2,X1,esk2_3(X2,X3,X4))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X1,X3)|~r1_orders_2(X2,X1,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (r1_orders_2(X1,esk1_4(X1,X2,X3,X4),X3)|X4=k11_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13, plain, (esk2_3(X1,X2,X3)=k11_lattice3(X1,X4,X5)|~r1_orders_2(X1,esk1_4(X1,X4,X5,esk2_3(X1,X2,X3)),X3)|~r1_orders_2(X1,esk1_4(X1,X4,X5,esk2_3(X1,X2,X3)),X2)|~r1_orders_2(X1,esk2_3(X1,X2,X3),X5)|~r1_orders_2(X1,esk2_3(X1,X2,X3),X4)|~m1_subset_1(esk1_4(X1,X4,X5,esk2_3(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])).
cnf(c_0_14, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk1_4(X2,X3,X4,X1),X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_12, c_0_7])).
cnf(c_0_15, plain, (r1_orders_2(X1,esk2_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, plain, (r1_orders_2(X1,esk1_4(X1,X2,X3,X4),X2)|X4=k11_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17, plain, (esk2_3(X1,X2,X3)=k11_lattice3(X1,X4,X3)|~r1_orders_2(X1,esk1_4(X1,X4,X3,esk2_3(X1,X2,X3)),X2)|~r1_orders_2(X1,esk2_3(X1,X2,X3),X4)|~m1_subset_1(esk1_4(X1,X4,X3,esk2_3(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_11]), c_0_15])).
cnf(c_0_18, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk1_4(X2,X3,X4,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_16, c_0_7])).
cnf(c_0_19, plain, (r1_orders_2(X1,esk2_3(X1,X2,X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20, plain, (m1_subset_1(esk1_4(X1,X2,X3,X4),u1_struct_0(X1))|X4=k11_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_21, plain, (r1_orders_2(X2,X1,X5)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X1,X3)|~r1_orders_2(X2,X1,X4)|X5!=k11_lattice3(X2,X3,X4)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22, plain, (esk2_3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~m1_subset_1(esk1_4(X1,X2,X3,esk2_3(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_11]), c_0_15]), c_0_19])).
cnf(c_0_23, plain, (X1=k11_lattice3(X2,X3,X4)|m1_subset_1(esk1_4(X2,X3,X4,X1),u1_struct_0(X2))|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)), inference(csr,[status(thm)],[c_0_20, c_0_7])).
cnf(c_0_24, plain, (r1_orders_2(X1,X2,k11_lattice3(X1,X3,X4))|~r1_orders_2(X1,X2,X4)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(k11_lattice3(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_21, c_0_7])])).
cnf(c_0_25, plain, (esk2_3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_11]), c_0_19]), c_0_15])).
fof(c_0_26, negated_conjecture, ~(![X1]:((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2))))), inference(assume_negation,[status(cth)],[t15_lattice3])).
cnf(c_0_27, plain, (k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X4,X5)|~r1_orders_2(X1,esk1_4(X1,X4,X5,k11_lattice3(X1,X2,X3)),X3)|~r1_orders_2(X1,esk1_4(X1,X4,X5,k11_lattice3(X1,X2,X3)),X2)|~r1_orders_2(X1,k11_lattice3(X1,X2,X3),X5)|~r1_orders_2(X1,k11_lattice3(X1,X2,X3),X4)|~m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_24]), c_0_23])).
cnf(c_0_28, plain, (m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_11, c_0_25])).
cnf(c_0_29, plain, (r1_orders_2(X1,k11_lattice3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_15, c_0_25])).
fof(c_0_30, negated_conjecture, ((v5_orders_2(esk6_0)&(v2_lattice3(esk6_0)&l1_orders_2(esk6_0)))&(m1_subset_1(esk7_0,u1_struct_0(esk6_0))&(m1_subset_1(esk8_0,u1_struct_0(esk6_0))&k11_lattice3(esk6_0,esk7_0,esk8_0)!=k11_lattice3(esk6_0,esk8_0,esk7_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
cnf(c_0_31, plain, (k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X4)|~r1_orders_2(X1,esk1_4(X1,X3,X4,k11_lattice3(X1,X2,X3)),X2)|~r1_orders_2(X1,k11_lattice3(X1,X2,X3),X4)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_18]), c_0_28]), c_0_29])).
cnf(c_0_32, plain, (r1_orders_2(X1,k11_lattice3(X1,X2,X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_25])).
cnf(c_0_33, negated_conjecture, (k11_lattice3(esk6_0,esk7_0,esk8_0)!=k11_lattice3(esk6_0,esk8_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_34, plain, (k11_lattice3(X1,X2,X3)=k11_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_14]), c_0_28]), c_0_29]), c_0_32])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk8_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, negated_conjecture, (l1_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, negated_conjecture, (v2_lattice3(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, negated_conjecture, (v5_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 41
# Proof object clause steps            : 32
# Proof object formula steps           : 9
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 27
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 228
# ...of these trivial                  : 0
# ...subsumed                          : 129
# ...remaining for further processing  : 99
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 0
# Generated clauses                    : 677
# ...of the previous two non-trivial   : 519
# Contextual simplify-reflections      : 40
# Paramodulations                      : 674
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 52
# Current number of unprocessed clauses: 321
# ...number of literals in the above   : 3771
# Current number of archived formulas  : 0
# Current number of archived clauses   : 38
# Clause-clause subsumption calls (NU) : 4907
# Rec. Clause-clause subsumption calls : 245
# Non-unit clause-clause subsumptions  : 183
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 31201

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.000 s
# Total time               : 0.024 s
# Maximum resident set size: 28660 pages
