# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', d8_lattice3)).
fof(d4_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>(v1_yellow_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&r1_lattice3(X1,u1_struct_0(X1),X2)))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', d4_yellow_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', t1_subset)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', t5_subset)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', t4_subset_1)).
fof(t42_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', t42_yellow_0)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', t30_yellow_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', t2_subset)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', d9_lattice3)).
fof(t16_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3))))))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', t16_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', dt_l1_orders_2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/9675ff23b885403ebbe65838e4e752f1.E_input', fc1_xboole_0)).
fof(c_0_13, plain, ![X5, X6, X7, X8]:((~r1_lattice3(X5,X6,X7)|(~m1_subset_1(X8,u1_struct_0(X5))|(~r2_hidden(X8,X6)|r1_orders_2(X5,X7,X8)))|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|r1_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&((r2_hidden(esk1_3(X5,X6,X7),X6)|r1_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))&(~r1_orders_2(X5,X7,esk1_3(X5,X6,X7))|r1_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|~l1_orders_2(X5))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
fof(c_0_14, plain, ![X18, X20]:(((m1_subset_1(esk5_1(X18),u1_struct_0(X18))|~v1_yellow_0(X18)|~l1_orders_2(X18))&(r1_lattice3(X18,u1_struct_0(X18),esk5_1(X18))|~v1_yellow_0(X18)|~l1_orders_2(X18)))&(~m1_subset_1(X20,u1_struct_0(X18))|~r1_lattice3(X18,u1_struct_0(X18),X20)|v1_yellow_0(X18)|~l1_orders_2(X18))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_yellow_0])])])])])).
fof(c_0_15, plain, ![X21, X22]:(~r2_hidden(X21,X22)|m1_subset_1(X21,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_16, plain, ![X23, X24, X25]:(~r2_hidden(X23,X24)|(~m1_subset_1(X24,k1_zfmisc_1(X25))|~v1_xboole_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_17, plain, ![X26]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X26)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t42_yellow_0])).
fof(c_0_19, plain, ![X30, X31, X32, X33, X34]:(((r2_lattice3(X30,X32,X31)|(X31!=k1_yellow_0(X30,X32)|~r1_yellow_0(X30,X32))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30)))&(~m1_subset_1(X33,u1_struct_0(X30))|(~r2_lattice3(X30,X32,X33)|r1_orders_2(X30,X31,X33))|(X31!=k1_yellow_0(X30,X32)|~r1_yellow_0(X30,X32))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30))))&(((X31=k1_yellow_0(X30,X34)|(m1_subset_1(esk6_3(X30,X31,X34),u1_struct_0(X30))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30)))&(r1_yellow_0(X30,X34)|(m1_subset_1(esk6_3(X30,X31,X34),u1_struct_0(X30))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30))))&(((X31=k1_yellow_0(X30,X34)|(r2_lattice3(X30,X34,esk6_3(X30,X31,X34))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30)))&(r1_yellow_0(X30,X34)|(r2_lattice3(X30,X34,esk6_3(X30,X31,X34))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30))))&((X31=k1_yellow_0(X30,X34)|(~r1_orders_2(X30,X31,esk6_3(X30,X31,X34))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30)))&(r1_yellow_0(X30,X34)|(~r1_orders_2(X30,X31,esk6_3(X30,X31,X34))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_20, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (r1_lattice3(X1,u1_struct_0(X1),esk5_1(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (m1_subset_1(esk5_1(X1),u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_24, plain, ![X27, X28]:(~m1_subset_1(X27,X28)|(v1_xboole_0(X28)|r2_hidden(X27,X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_27, plain, ![X36, X37, X38, X39]:((~r2_lattice3(X36,X37,X38)|(~m1_subset_1(X39,u1_struct_0(X36))|(~r2_hidden(X39,X37)|r1_orders_2(X36,X39,X38)))|~m1_subset_1(X38,u1_struct_0(X36))|~l1_orders_2(X36))&((m1_subset_1(esk7_3(X36,X37,X38),u1_struct_0(X36))|r2_lattice3(X36,X37,X38)|~m1_subset_1(X38,u1_struct_0(X36))|~l1_orders_2(X36))&((r2_hidden(esk7_3(X36,X37,X38),X37)|r2_lattice3(X36,X37,X38)|~m1_subset_1(X38,u1_struct_0(X36))|~l1_orders_2(X36))&(~r1_orders_2(X36,esk7_3(X36,X37,X38),X38)|r2_lattice3(X36,X37,X38)|~m1_subset_1(X38,u1_struct_0(X36))|~l1_orders_2(X36))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
fof(c_0_28, plain, ![X10, X11, X13, X14, X15]:(((m1_subset_1(esk2_2(X10,X11),u1_struct_0(X10))|~r2_yellow_0(X10,X11)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&((r1_lattice3(X10,X11,esk2_2(X10,X11))|~r2_yellow_0(X10,X11)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(~m1_subset_1(X13,u1_struct_0(X10))|(~r1_lattice3(X10,X11,X13)|r1_orders_2(X10,X13,esk2_2(X10,X11)))|~r2_yellow_0(X10,X11)|(~v5_orders_2(X10)|~l1_orders_2(X10)))))&((m1_subset_1(esk3_3(X10,X14,X15),u1_struct_0(X10))|~r1_lattice3(X10,X14,X15)|~m1_subset_1(X15,u1_struct_0(X10))|r2_yellow_0(X10,X14)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&((r1_lattice3(X10,X14,esk3_3(X10,X14,X15))|~r1_lattice3(X10,X14,X15)|~m1_subset_1(X15,u1_struct_0(X10))|r2_yellow_0(X10,X14)|(~v5_orders_2(X10)|~l1_orders_2(X10)))&(~r1_orders_2(X10,esk3_3(X10,X14,X15),X15)|~r1_lattice3(X10,X14,X15)|~m1_subset_1(X15,u1_struct_0(X10))|r2_yellow_0(X10,X14)|(~v5_orders_2(X10)|~l1_orders_2(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_yellow_0])])])])])])).
fof(c_0_29, negated_conjecture, ((~v2_struct_0(esk4_0)&(v5_orders_2(esk4_0)&(v1_yellow_0(esk4_0)&l1_orders_2(esk4_0))))&(~r1_yellow_0(esk4_0,k1_xboole_0)|~r2_yellow_0(esk4_0,u1_struct_0(esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
fof(c_0_30, plain, ![X29]:(v2_struct_0(X29)|~l1_struct_0(X29)|~v1_xboole_0(u1_struct_0(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_31, plain, (r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk6_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (r1_orders_2(X1,esk5_1(X1),X2)|~v1_yellow_0(X1)|~r2_hidden(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_33, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk6_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_35, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_36, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (r1_lattice3(X1,X2,esk3_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_41, plain, ![X41]:(~l1_orders_2(X41)|l1_struct_0(X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_42, plain, (r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,esk5_1(X1))|~v1_yellow_0(X1)|~v5_orders_2(X1)|~r2_hidden(esk6_3(X1,esk5_1(X1),X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_23])).
cnf(c_0_43, plain, (v1_xboole_0(u1_struct_0(X1))|r1_yellow_0(X1,X2)|r2_hidden(esk6_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_44, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~v1_xboole_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_46, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk3_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_47, plain, (r2_yellow_0(X1,X2)|r1_orders_2(X1,esk3_3(X1,X2,X3),X4)|~v5_orders_2(X1)|~r2_hidden(X4,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_37]), c_0_38])).
cnf(c_0_48, negated_conjecture, (~l1_struct_0(esk4_0)|~v1_xboole_0(u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50, negated_conjecture, (l1_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_51, plain, (v1_xboole_0(u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,esk5_1(X1))|~v1_yellow_0(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_23])).
cnf(c_0_52, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53, plain, (r2_yellow_0(X1,X2)|~v5_orders_2(X1)|~r2_hidden(X3,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_55, plain, (v1_xboole_0(u1_struct_0(X1))|r1_yellow_0(X1,k1_xboole_0)|~v1_yellow_0(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_23])).
cnf(c_0_56, negated_conjecture, (v1_yellow_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_57, negated_conjecture, (v5_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_58, plain, (r2_yellow_0(X1,u1_struct_0(X1))|~v1_yellow_0(X1)|~v5_orders_2(X1)|~r2_hidden(esk5_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_21]), c_0_23])).
cnf(c_0_59, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk5_1(X1),u1_struct_0(X1))|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_33, c_0_23])).
cnf(c_0_60, negated_conjecture, (~r1_yellow_0(esk4_0,k1_xboole_0)|~r2_yellow_0(esk4_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_61, negated_conjecture, (r1_yellow_0(esk4_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_57]), c_0_50])])).
cnf(c_0_62, plain, (v1_xboole_0(u1_struct_0(X1))|r2_yellow_0(X1,u1_struct_0(X1))|~v1_yellow_0(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63, negated_conjecture, (~r2_yellow_0(esk4_0,u1_struct_0(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])])).
cnf(c_0_64, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_56]), c_0_57]), c_0_50])]), c_0_54]), c_0_63]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 39
# Proof object formula steps           : 26
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 155
# ...of these trivial                  : 0
# ...subsumed                          : 21
# ...remaining for further processing  : 134
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 1
# Generated clauses                    : 159
# ...of the previous two non-trivial   : 151
# Contextual simplify-reflections      : 21
# Paramodulations                      : 157
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 90
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 81
# Current number of unprocessed clauses: 67
# ...number of literals in the above   : 566
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 2174
# Rec. Clause-clause subsumption calls : 226
# Non-unit clause-clause subsumptions  : 46
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7691

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.005 s
# Total time               : 0.016 s
# Maximum resident set size: 28812 pages
