# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d39_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>(v24_waybel_0(X1)<=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r3_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/96c289c58ecb40f08539e40af42986a4.E_input', d39_waybel_0)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/96c289c58ecb40f08539e40af42986a4.E_input', redefinition_r3_orders_2)).
fof(t75_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>(v24_waybel_0(X1)<=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))))=>r1_yellow_0(X1,X2)))), file('data/MPTP/proofs/96c289c58ecb40f08539e40af42986a4.E_input', t75_waybel_0)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/96c289c58ecb40f08539e40af42986a4.E_input', d7_yellow_0)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/96c289c58ecb40f08539e40af42986a4.E_input', t30_yellow_0)).
fof(c_0_5, plain, ![X6, X7, X9, X11]:(((m1_subset_1(esk1_2(X6,X7),u1_struct_0(X6))|(v1_xboole_0(X7)|(~v1_waybel_0(X7,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))))|~v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6))))&((r2_lattice3(X6,X7,esk1_2(X6,X7))|(v1_xboole_0(X7)|(~v1_waybel_0(X7,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))))|~v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6))))&(~m1_subset_1(X9,u1_struct_0(X6))|(~r2_lattice3(X6,X7,X9)|r3_orders_2(X6,esk1_2(X6,X7),X9))|(v1_xboole_0(X7)|(~v1_waybel_0(X7,X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))))|~v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6))))))&(((~v1_xboole_0(esk2_1(X6))|v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6))))&((v1_waybel_0(esk2_1(X6),X6)|v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6))))&(m1_subset_1(esk2_1(X6),k1_zfmisc_1(u1_struct_0(X6)))|v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6))))))&((m1_subset_1(esk3_2(X6,X11),u1_struct_0(X6))|~r2_lattice3(X6,esk2_1(X6),X11)|~m1_subset_1(X11,u1_struct_0(X6))|v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6))))&((r2_lattice3(X6,esk2_1(X6),esk3_2(X6,X11))|~r2_lattice3(X6,esk2_1(X6),X11)|~m1_subset_1(X11,u1_struct_0(X6))|v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6))))&(~r3_orders_2(X6,X11,esk3_2(X6,X11))|~r2_lattice3(X6,esk2_1(X6),X11)|~m1_subset_1(X11,u1_struct_0(X6))|v24_waybel_0(X6)|(v2_struct_0(X6)|(~v3_orders_2(X6)|~l1_orders_2(X6)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d39_waybel_0])])])])])])).
fof(c_0_6, plain, ![X13, X14, X15]:((~r3_orders_2(X13,X14,X15)|r1_orders_2(X13,X14,X15)|(v2_struct_0(X13)|(~v3_orders_2(X13)|~l1_orders_2(X13))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))))&(~r1_orders_2(X13,X14,X15)|r3_orders_2(X13,X14,X15)|(v2_struct_0(X13)|(~v3_orders_2(X13)|~l1_orders_2(X13))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
fof(c_0_7, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>(v24_waybel_0(X1)<=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))))=>r1_yellow_0(X1,X2))))), inference(assume_negation,[status(cth)],[t75_waybel_0])).
cnf(c_0_8, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~r3_orders_2(X1,X2,esk3_2(X1,X2))|~r2_lattice3(X1,esk2_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (m1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|v24_waybel_0(X1)|v2_struct_0(X1)|~r2_lattice3(X1,esk2_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_11, plain, ![X19, X20, X22, X23, X25, X26, X29]:(((m1_subset_1(esk6_2(X19,X20),u1_struct_0(X19))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&((r2_lattice3(X19,X20,esk6_2(X19,X20))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&((~m1_subset_1(X22,u1_struct_0(X19))|(~r2_lattice3(X19,X20,X22)|r1_orders_2(X19,esk6_2(X19,X20),X22))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&((m1_subset_1(esk7_3(X19,X20,X23),u1_struct_0(X19))|~r2_lattice3(X19,X20,X23)|X23=esk6_2(X19,X20)|~m1_subset_1(X23,u1_struct_0(X19))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&((r2_lattice3(X19,X20,esk7_3(X19,X20,X23))|~r2_lattice3(X19,X20,X23)|X23=esk6_2(X19,X20)|~m1_subset_1(X23,u1_struct_0(X19))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19))&(~r1_orders_2(X19,X23,esk7_3(X19,X20,X23))|~r2_lattice3(X19,X20,X23)|X23=esk6_2(X19,X20)|~m1_subset_1(X23,u1_struct_0(X19))|~r1_yellow_0(X19,X20)|~l1_orders_2(X19)))))))&(((m1_subset_1(esk9_3(X19,X25,X26),u1_struct_0(X19))|m1_subset_1(esk8_3(X19,X25,X26),u1_struct_0(X19))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(((r2_lattice3(X19,X25,esk9_3(X19,X25,X26))|m1_subset_1(esk8_3(X19,X25,X26),u1_struct_0(X19))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(~m1_subset_1(X29,u1_struct_0(X19))|(~r2_lattice3(X19,X25,X29)|r1_orders_2(X19,esk9_3(X19,X25,X26),X29))|m1_subset_1(esk8_3(X19,X25,X26),u1_struct_0(X19))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19)))&(esk9_3(X19,X25,X26)!=X26|m1_subset_1(esk8_3(X19,X25,X26),u1_struct_0(X19))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))))&(((m1_subset_1(esk9_3(X19,X25,X26),u1_struct_0(X19))|r2_lattice3(X19,X25,esk8_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(((r2_lattice3(X19,X25,esk9_3(X19,X25,X26))|r2_lattice3(X19,X25,esk8_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(~m1_subset_1(X29,u1_struct_0(X19))|(~r2_lattice3(X19,X25,X29)|r1_orders_2(X19,esk9_3(X19,X25,X26),X29))|r2_lattice3(X19,X25,esk8_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19)))&(esk9_3(X19,X25,X26)!=X26|r2_lattice3(X19,X25,esk8_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))))&((m1_subset_1(esk9_3(X19,X25,X26),u1_struct_0(X19))|~r1_orders_2(X19,X26,esk8_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(((r2_lattice3(X19,X25,esk9_3(X19,X25,X26))|~r1_orders_2(X19,X26,esk8_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))&(~m1_subset_1(X29,u1_struct_0(X19))|(~r2_lattice3(X19,X25,X29)|r1_orders_2(X19,esk9_3(X19,X25,X26),X29))|~r1_orders_2(X19,X26,esk8_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19)))&(esk9_3(X19,X25,X26)!=X26|~r1_orders_2(X19,X26,esk8_3(X19,X25,X26))|~r2_lattice3(X19,X25,X26)|~m1_subset_1(X26,u1_struct_0(X19))|r1_yellow_0(X19,X25)|~l1_orders_2(X19))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
fof(c_0_12, negated_conjecture, ![X18]:((~v2_struct_0(esk4_0)&(v3_orders_2(esk4_0)&(v5_orders_2(esk4_0)&l1_orders_2(esk4_0))))&((((~v1_xboole_0(esk5_0)|~v24_waybel_0(esk4_0))&((v1_waybel_0(esk5_0,esk4_0)|~v24_waybel_0(esk4_0))&(m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk4_0)))|~v24_waybel_0(esk4_0))))&(~r1_yellow_0(esk4_0,esk5_0)|~v24_waybel_0(esk4_0)))&(v24_waybel_0(esk4_0)|(v1_xboole_0(X18)|(~v1_waybel_0(X18,esk4_0)|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(esk4_0))))|r1_yellow_0(esk4_0,X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])])])).
cnf(c_0_13, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~r1_orders_2(X1,X2,esk3_2(X1,X2))|~r2_lattice3(X1,esk2_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10])).
cnf(c_0_14, plain, (r1_orders_2(X2,esk6_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (m1_subset_1(esk6_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (v24_waybel_0(esk4_0)|v1_xboole_0(X1)|r1_yellow_0(esk4_0,X1)|~v1_waybel_0(X1,esk4_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (m1_subset_1(esk2_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|v24_waybel_0(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18, negated_conjecture, (l1_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (v3_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_21, plain, ![X30, X31, X32, X33, X34]:(((r2_lattice3(X30,X32,X31)|(X31!=k1_yellow_0(X30,X32)|~r1_yellow_0(X30,X32))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30)))&(~m1_subset_1(X33,u1_struct_0(X30))|(~r2_lattice3(X30,X32,X33)|r1_orders_2(X30,X31,X33))|(X31!=k1_yellow_0(X30,X32)|~r1_yellow_0(X30,X32))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30))))&(((X31=k1_yellow_0(X30,X34)|(m1_subset_1(esk10_3(X30,X31,X34),u1_struct_0(X30))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30)))&(r1_yellow_0(X30,X34)|(m1_subset_1(esk10_3(X30,X31,X34),u1_struct_0(X30))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30))))&(((X31=k1_yellow_0(X30,X34)|(r2_lattice3(X30,X34,esk10_3(X30,X31,X34))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30)))&(r1_yellow_0(X30,X34)|(r2_lattice3(X30,X34,esk10_3(X30,X31,X34))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30))))&((X31=k1_yellow_0(X30,X34)|(~r1_orders_2(X30,X31,esk10_3(X30,X31,X34))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30)))&(r1_yellow_0(X30,X34)|(~r1_orders_2(X30,X31,esk10_3(X30,X31,X34))|~r2_lattice3(X30,X34,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(~v5_orders_2(X30)|~l1_orders_2(X30))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_22, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, plain, (r3_orders_2(X2,esk1_2(X2,X3),X1)|v1_xboole_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~v1_waybel_0(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v24_waybel_0(X2)|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_24, plain, (m1_subset_1(esk1_2(X1,X2),u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,esk3_2(X1,esk6_2(X1,X2)))|~r2_lattice3(X1,esk2_1(X1),esk6_2(X1,X2))|~m1_subset_1(esk3_2(X1,esk6_2(X1,X2)),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])).
cnf(c_0_26, plain, (r2_lattice3(X1,esk2_1(X1),esk3_2(X1,X2))|v24_waybel_0(X1)|v2_struct_0(X1)|~r2_lattice3(X1,esk2_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_27, plain, (r2_lattice3(X1,X2,esk6_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_28, negated_conjecture, (r1_yellow_0(esk4_0,esk2_1(esk4_0))|v1_xboole_0(esk2_1(esk4_0))|v24_waybel_0(esk4_0)|~v1_waybel_0(esk2_1(esk4_0),esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_29, plain, (v1_waybel_0(esk2_1(X1),X1)|v24_waybel_0(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_30, plain, (r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk10_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (r1_orders_2(X1,esk1_2(X1,X2),X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v1_waybel_0(X2,X1)|~v24_waybel_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24])).
cnf(c_0_32, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~r1_yellow_0(X1,esk2_1(X1))|~m1_subset_1(esk3_2(X1,esk6_2(X1,esk2_1(X1))),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_15]), c_0_27])).
cnf(c_0_33, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~v1_xboole_0(esk2_1(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_34, negated_conjecture, (r1_yellow_0(esk4_0,esk2_1(esk4_0))|v1_xboole_0(esk2_1(esk4_0))|v24_waybel_0(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_35, plain, (r1_yellow_0(X1,X2)|v1_xboole_0(X3)|v2_struct_0(X1)|~v5_orders_2(X1)|~r2_lattice3(X1,X3,esk10_3(X1,esk1_2(X1,X3),X2))|~r2_lattice3(X1,X2,esk1_2(X1,X3))|~m1_subset_1(esk10_3(X1,esk1_2(X1,X3),X2),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_0(X3,X1)|~v24_waybel_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_24])).
cnf(c_0_36, plain, (r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk10_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_37, plain, (r2_lattice3(X1,X2,esk1_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_38, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~r1_yellow_0(X1,esk2_1(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_10]), c_0_15]), c_0_27])).
cnf(c_0_39, negated_conjecture, (r1_yellow_0(esk4_0,esk2_1(esk4_0))|v24_waybel_0(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_40, plain, (r1_yellow_0(X1,X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v5_orders_2(X1)|~m1_subset_1(esk10_3(X1,esk1_2(X1,X2),X2),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_0(X2,X1)|~v24_waybel_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_24]), c_0_37])).
cnf(c_0_41, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk10_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk4_0)))|~v24_waybel_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_43, negated_conjecture, (v24_waybel_0(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_18]), c_0_19])]), c_0_20])).
cnf(c_0_44, negated_conjecture, (v1_waybel_0(esk5_0,esk4_0)|~v24_waybel_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_45, negated_conjecture, (~r1_yellow_0(esk4_0,esk5_0)|~v24_waybel_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_46, negated_conjecture, (~v1_xboole_0(esk5_0)|~v24_waybel_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_47, plain, (r1_yellow_0(X1,X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_waybel_0(X2,X1)|~v24_waybel_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_24]), c_0_37])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43])])).
cnf(c_0_49, negated_conjecture, (v5_orders_2(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_50, negated_conjecture, (v1_waybel_0(esk5_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_43])])).
cnf(c_0_51, negated_conjecture, (~r1_yellow_0(esk4_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_43])])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_43])])).
cnf(c_0_53, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_43]), c_0_18]), c_0_19])]), c_0_51]), c_0_52]), c_0_20]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 43
# Proof object formula steps           : 11
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 46
# Processed clauses                    : 137
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 131
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 6
# Generated clauses                    : 96
# ...of the previous two non-trivial   : 71
# Contextual simplify-reflections      : 34
# Paramodulations                      : 94
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 70
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 61
# Current number of unprocessed clauses: 26
# ...number of literals in the above   : 306
# Current number of archived formulas  : 0
# Current number of archived clauses   : 59
# Clause-clause subsumption calls (NU) : 2753
# Rec. Clause-clause subsumption calls : 54
# Non-unit clause-clause subsumptions  : 47
# Unit Clause-clause subsumption calls : 82
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8918

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.003 s
# Total time               : 0.015 s
# Maximum resident set size: 28672 pages
