# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_tmap_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>(r1_subset_1(X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))&(k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X4)=X6))))))))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', t6_tmap_1)).
fof(redefinition_r1_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))=>(r1_subset_1(X1,X2)<=>r1_xboole_0(X1,X2))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', redefinition_r1_subset_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', d7_xboole_0)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', redefinition_k9_subset_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', t3_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', d10_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', t3_subset)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', cc5_funct_2)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', cc1_relset_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', t5_subset)).
fof(d4_partfun1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>(v1_partfun1(X2,X1)<=>k1_relat_1(X2)=X1)), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', d4_partfun1)).
fof(redefinition_k2_partfun1, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k2_partfun1(X1,X2,X3,X4)=k7_relat_1(X3,X4)), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', redefinition_k2_partfun1)).
fof(t192_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k7_relat_1(X2,X1)=k7_relat_1(X2,k3_xboole_0(k1_relat_1(X2),X1))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', t192_relat_1)).
fof(t207_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_xboole_0(X1,X2)=>k7_relat_1(k7_relat_1(X3,X1),X2)=k1_xboole_0)), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', t207_relat_1)).
fof(t209_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>X2=k7_relat_1(X2,X1)), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', t209_relat_1)).
fof(dt_k1_tmap_1, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))&((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))&((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2))))))))))=>(v1_funct_1(k1_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k1_tmap_1(X1,X2,X3,X4,X5,X6),k4_subset_1(X1,X3,X4),X2)&m1_subset_1(k1_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', dt_k1_tmap_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', cc1_subset_1)).
fof(d1_tmap_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))=>![X7]:((v1_funct_1(X7)&(v1_funct_2(X7,k4_subset_1(X1,X3,X4),X2)&m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))))=>(X7=k1_tmap_1(X1,X2,X3,X4,X5,X6)<=>(k2_partfun1(k4_subset_1(X1,X3,X4),X2,X7,X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,X7,X4)=X6)))))))))), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', d1_tmap_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/979a4ebad2cd4add95eaba2058d92261.E_input', fc1_xboole_0)).
fof(c_0_20, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>(r1_subset_1(X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))&(k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X4)=X6)))))))))), inference(assume_negation,[status(cth)],[t6_tmap_1])).
fof(c_0_21, plain, ![X44, X45]:((~r1_subset_1(X44,X45)|r1_xboole_0(X44,X45)|(v1_xboole_0(X44)|v1_xboole_0(X45)))&(~r1_xboole_0(X44,X45)|r1_subset_1(X44,X45)|(v1_xboole_0(X44)|v1_xboole_0(X45)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r1_subset_1])])])])).
fof(c_0_22, negated_conjecture, (~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&((~v1_xboole_0(esk3_0)&m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0)))&((~v1_xboole_0(esk4_0)&m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0)))&(r1_subset_1(esk3_0,esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,esk3_0,esk2_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk2_0)))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,esk4_0,esk2_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0)))))&(k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk1_0,esk3_0,esk4_0))!=k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk1_0,esk3_0,esk4_0))|(k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
fof(c_0_23, plain, ![X38, X39]:((~r1_xboole_0(X38,X39)|k3_xboole_0(X38,X39)=k1_xboole_0)&(k3_xboole_0(X38,X39)!=k1_xboole_0|r1_xboole_0(X38,X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_24, plain, (r1_xboole_0(X1,X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, negated_conjecture, (r1_subset_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, ![X32, X33, X34]:(~m1_subset_1(X34,k1_zfmisc_1(X32))|k9_subset_1(X32,X33,X34)=k3_xboole_0(X33,X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
fof(c_0_29, plain, ![X26, X27, X29, X30, X31]:(((r2_hidden(esk7_2(X26,X27),X26)|r1_xboole_0(X26,X27))&(r2_hidden(esk7_2(X26,X27),X27)|r1_xboole_0(X26,X27)))&(~r2_hidden(X31,X29)|~r2_hidden(X31,X30)|~r1_xboole_0(X29,X30))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_30, plain, ![X42, X43]:(((r1_tarski(X42,X43)|X42!=X43)&(r1_tarski(X43,X42)|X42!=X43))&(~r1_tarski(X42,X43)|~r1_tarski(X43,X42)|X42=X43)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_31, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, negated_conjecture, (r1_xboole_0(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_27])).
cnf(c_0_33, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_35, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_37, plain, ![X62, X63]:((~m1_subset_1(X62,k1_zfmisc_1(X63))|r1_tarski(X62,X63))&(~r1_tarski(X62,X63)|m1_subset_1(X62,k1_zfmisc_1(X63)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_38, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_39, plain, ![X18, X19, X20]:((v1_funct_1(X20)|(~v1_funct_1(X20)|~v1_funct_2(X20,X18,X19))|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19)))|v1_xboole_0(X19))&(v1_partfun1(X20,X18)|(~v1_funct_1(X20)|~v1_funct_2(X20,X18,X19))|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19)))|v1_xboole_0(X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
fof(c_0_40, plain, ![X59, X60, X61]:((v4_relat_1(X61,X59)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))&(v5_relat_1(X61,X60)|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_41, plain, ![X56, X57, X58]:(~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(X56,X57)))|v1_relat_1(X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_42, negated_conjecture, (k3_xboole_0(esk3_0,esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_43, negated_conjecture, (k3_xboole_0(X1,esk4_0)=k9_subset_1(esk1_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_44, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk7_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45, plain, (r2_hidden(esk7_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_46, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_38])).
fof(c_0_48, plain, ![X35, X36, X37]:(~r2_hidden(X35,X36)|(~m1_subset_1(X36,k1_zfmisc_1(X37))|~v1_xboole_0(X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_49, plain, ![X48, X49]:((~v1_partfun1(X49,X48)|k1_relat_1(X49)=X48|(~v1_relat_1(X49)|~v4_relat_1(X49,X48)))&(k1_relat_1(X49)!=X48|v1_partfun1(X49,X48)|(~v1_relat_1(X49)|~v4_relat_1(X49,X48)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_partfun1])])])).
cnf(c_0_50, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, negated_conjecture, (v1_funct_2(esk6_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_52, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_54, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_55, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_56, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_57, plain, ![X14, X15, X16, X17]:(~v1_funct_1(X16)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X15)))|k2_partfun1(X14,X15,X16,X17)=k7_relat_1(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_partfun1])])).
cnf(c_0_58, negated_conjecture, (v1_funct_2(esk5_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_59, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_61, negated_conjecture, (k1_xboole_0=k9_subset_1(esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_62, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_64, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_65, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_66, plain, ![X24, X25]:(~v1_relat_1(X25)|k7_relat_1(X25,X24)=k7_relat_1(X25,k3_xboole_0(k1_relat_1(X25),X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t192_relat_1])])).
cnf(c_0_67, plain, (k1_relat_1(X1)=X2|~v1_partfun1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_68, negated_conjecture, (v1_partfun1(esk6_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_69, negated_conjecture, (v4_relat_1(esk6_0,esk4_0)), inference(spm,[status(thm)],[c_0_55, c_0_53])).
cnf(c_0_70, negated_conjecture, (v1_relat_1(esk6_0)), inference(spm,[status(thm)],[c_0_56, c_0_53])).
cnf(c_0_71, plain, (k2_partfun1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_72, negated_conjecture, (v1_partfun1(esk5_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_58]), c_0_59]), c_0_60])]), c_0_54])).
cnf(c_0_73, negated_conjecture, (v4_relat_1(esk5_0,esk3_0)), inference(spm,[status(thm)],[c_0_55, c_0_60])).
cnf(c_0_74, negated_conjecture, (v1_relat_1(esk5_0)), inference(spm,[status(thm)],[c_0_56, c_0_60])).
fof(c_0_75, plain, ![X21, X22, X23]:(~v1_relat_1(X23)|(~r1_xboole_0(X21,X22)|k7_relat_1(k7_relat_1(X23,X21),X22)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t207_relat_1])])).
cnf(c_0_76, plain, (k3_xboole_0(X1,X2)=k9_subset_1(esk1_0,esk3_0,esk4_0)|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_31, c_0_61])).
cnf(c_0_77, negated_conjecture, (r1_xboole_0(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_62, c_0_32])).
cnf(c_0_78, negated_conjecture, (k3_xboole_0(X1,esk3_0)=k9_subset_1(esk1_0,X1,esk3_0)), inference(spm,[status(thm)],[c_0_33, c_0_63])).
cnf(c_0_79, plain, (k3_xboole_0(X1,X2)=k9_subset_1(X2,X1,X2)), inference(spm,[status(thm)],[c_0_33, c_0_64])).
cnf(c_0_80, plain, (r1_xboole_0(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_65, c_0_45])).
cnf(c_0_81, plain, (k7_relat_1(X1,X2)=k7_relat_1(X1,k3_xboole_0(k1_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_82, negated_conjecture, (k1_relat_1(esk6_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70])])).
cnf(c_0_83, negated_conjecture, (k7_relat_1(esk6_0,X1)=k2_partfun1(esk4_0,esk2_0,esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_53]), c_0_52])])).
cnf(c_0_84, negated_conjecture, (k1_relat_1(esk5_0)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_72]), c_0_73]), c_0_74])])).
cnf(c_0_85, plain, (k7_relat_1(k7_relat_1(X1,X2),X3)=k1_xboole_0|~v1_relat_1(X1)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_86, negated_conjecture, (k9_subset_1(esk1_0,esk4_0,esk3_0)=k9_subset_1(esk1_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])).
cnf(c_0_87, negated_conjecture, (k9_subset_1(esk1_0,X1,esk3_0)=k9_subset_1(esk3_0,X1,esk3_0)), inference(rw,[status(thm)],[c_0_78, c_0_79])).
fof(c_0_88, plain, ![X46, X47]:(~v1_relat_1(X47)|~v4_relat_1(X47,X46)|X47=k7_relat_1(X47,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t209_relat_1])])).
cnf(c_0_89, plain, (r1_xboole_0(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_80, c_0_64])).
cnf(c_0_90, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,k3_xboole_0(esk4_0,X1))=k2_partfun1(esk4_0,esk2_0,esk6_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_70])]), c_0_83]), c_0_83])).
cnf(c_0_91, negated_conjecture, (k7_relat_1(esk5_0,k3_xboole_0(esk3_0,X1))=k7_relat_1(esk5_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_84]), c_0_74])])).
cnf(c_0_92, negated_conjecture, (k7_relat_1(esk5_0,X1)=k2_partfun1(esk3_0,esk2_0,esk5_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_60]), c_0_59])])).
fof(c_0_93, plain, ![X50, X51, X52, X53, X54, X55]:((v1_funct_1(k1_tmap_1(X50,X51,X52,X53,X54,X55))|(v1_xboole_0(X50)|(v1_xboole_0(X51)|(v1_xboole_0(X52)|~m1_subset_1(X52,k1_zfmisc_1(X50))|(v1_xboole_0(X53)|~m1_subset_1(X53,k1_zfmisc_1(X50))|(~v1_funct_1(X54)|(~v1_funct_2(X54,X52,X51)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X51))))|(~v1_funct_1(X55)|(~v1_funct_2(X55,X53,X51)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X51)))))))))))&((v1_funct_2(k1_tmap_1(X50,X51,X52,X53,X54,X55),k4_subset_1(X50,X52,X53),X51)|(v1_xboole_0(X50)|(v1_xboole_0(X51)|(v1_xboole_0(X52)|~m1_subset_1(X52,k1_zfmisc_1(X50))|(v1_xboole_0(X53)|~m1_subset_1(X53,k1_zfmisc_1(X50))|(~v1_funct_1(X54)|(~v1_funct_2(X54,X52,X51)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X51))))|(~v1_funct_1(X55)|(~v1_funct_2(X55,X53,X51)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X51)))))))))))&(m1_subset_1(k1_tmap_1(X50,X51,X52,X53,X54,X55),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X50,X52,X53),X51)))|(v1_xboole_0(X50)|(v1_xboole_0(X51)|(v1_xboole_0(X52)|~m1_subset_1(X52,k1_zfmisc_1(X50))|(v1_xboole_0(X53)|~m1_subset_1(X53,k1_zfmisc_1(X50))|(~v1_funct_1(X54)|(~v1_funct_2(X54,X52,X51)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X51))))|(~v1_funct_1(X55)|(~v1_funct_2(X55,X53,X51)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X51))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tmap_1])])])])).
fof(c_0_94, plain, ![X40, X41]:(~v1_xboole_0(X40)|(~m1_subset_1(X41,k1_zfmisc_1(X40))|v1_xboole_0(X41))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_95, plain, (k7_relat_1(k7_relat_1(X1,X2),X3)=k9_subset_1(esk1_0,esk3_0,esk4_0)|~r1_xboole_0(X2,X3)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_85, c_0_61])).
cnf(c_0_96, negated_conjecture, (k9_subset_1(esk1_0,esk3_0,esk4_0)=k9_subset_1(esk3_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_97, plain, (X1=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_98, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_31, c_0_89])).
cnf(c_0_99, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk1_0,esk3_0,esk4_0))!=k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk1_0,esk3_0,esk4_0))|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_100, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk1_0,esk3_0,esk4_0))=k2_partfun1(esk4_0,esk2_0,esk6_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_78]), c_0_86])).
cnf(c_0_101, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k3_xboole_0(esk3_0,X1))=k2_partfun1(esk3_0,esk2_0,esk5_0,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_92]), c_0_92])).
cnf(c_0_102, negated_conjecture, (k9_subset_1(esk1_0,X1,esk4_0)=k9_subset_1(esk4_0,X1,esk4_0)), inference(rw,[status(thm)],[c_0_43, c_0_79])).
fof(c_0_103, plain, ![X64, X65, X66, X67, X68, X69, X70]:(((k2_partfun1(k4_subset_1(X64,X66,X67),X65,X70,X66)=X68|X70!=k1_tmap_1(X64,X65,X66,X67,X68,X69)|(~v1_funct_1(X70)|(~v1_funct_2(X70,k4_subset_1(X64,X66,X67),X65)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X64,X66,X67),X65)))))|k2_partfun1(X66,X65,X68,k9_subset_1(X64,X66,X67))!=k2_partfun1(X67,X65,X69,k9_subset_1(X64,X66,X67))|(~v1_funct_1(X69)|(~v1_funct_2(X69,X67,X65)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X65)))))|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X65)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X65)))))|(v1_xboole_0(X67)|~m1_subset_1(X67,k1_zfmisc_1(X64)))|(v1_xboole_0(X66)|~m1_subset_1(X66,k1_zfmisc_1(X64)))|v1_xboole_0(X65)|v1_xboole_0(X64))&(k2_partfun1(k4_subset_1(X64,X66,X67),X65,X70,X67)=X69|X70!=k1_tmap_1(X64,X65,X66,X67,X68,X69)|(~v1_funct_1(X70)|(~v1_funct_2(X70,k4_subset_1(X64,X66,X67),X65)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X64,X66,X67),X65)))))|k2_partfun1(X66,X65,X68,k9_subset_1(X64,X66,X67))!=k2_partfun1(X67,X65,X69,k9_subset_1(X64,X66,X67))|(~v1_funct_1(X69)|(~v1_funct_2(X69,X67,X65)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X65)))))|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X65)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X65)))))|(v1_xboole_0(X67)|~m1_subset_1(X67,k1_zfmisc_1(X64)))|(v1_xboole_0(X66)|~m1_subset_1(X66,k1_zfmisc_1(X64)))|v1_xboole_0(X65)|v1_xboole_0(X64)))&(k2_partfun1(k4_subset_1(X64,X66,X67),X65,X70,X66)!=X68|k2_partfun1(k4_subset_1(X64,X66,X67),X65,X70,X67)!=X69|X70=k1_tmap_1(X64,X65,X66,X67,X68,X69)|(~v1_funct_1(X70)|(~v1_funct_2(X70,k4_subset_1(X64,X66,X67),X65)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X64,X66,X67),X65)))))|k2_partfun1(X66,X65,X68,k9_subset_1(X64,X66,X67))!=k2_partfun1(X67,X65,X69,k9_subset_1(X64,X66,X67))|(~v1_funct_1(X69)|(~v1_funct_2(X69,X67,X65)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X65)))))|(~v1_funct_1(X68)|(~v1_funct_2(X68,X66,X65)|~m1_subset_1(X68,k1_zfmisc_1(k2_zfmisc_1(X66,X65)))))|(v1_xboole_0(X67)|~m1_subset_1(X67,k1_zfmisc_1(X64)))|(v1_xboole_0(X66)|~m1_subset_1(X66,k1_zfmisc_1(X64)))|v1_xboole_0(X65)|v1_xboole_0(X64))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tmap_1])])])])])).
cnf(c_0_104, plain, (m1_subset_1(k1_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_105, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_106, plain, (v1_funct_1(k1_tmap_1(X1,X2,X3,X4,X5,X6))|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_107, plain, (v1_funct_2(k1_tmap_1(X1,X2,X3,X4,X5,X6),k4_subset_1(X1,X3,X4),X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_108, plain, (k7_relat_1(k7_relat_1(X1,X2),X3)=k9_subset_1(esk3_0,esk4_0,esk3_0)|~r1_xboole_0(X2,X3)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_109, negated_conjecture, (k7_relat_1(esk6_0,esk4_0)=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_69]), c_0_70])])).
cnf(c_0_110, plain, (k3_xboole_0(X1,X2)=k9_subset_1(esk1_0,esk3_0,esk4_0)|~v1_xboole_0(X2)), inference(rw,[status(thm)],[c_0_98, c_0_61])).
cnf(c_0_111, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk4_0,esk2_0,esk6_0,esk3_0)!=k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk1_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_112, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(X1,esk3_0,X1))=k2_partfun1(esk3_0,esk2_0,esk5_0,X1)), inference(rw,[status(thm)],[c_0_101, c_0_79])).
cnf(c_0_113, negated_conjecture, (k9_subset_1(esk4_0,esk3_0,esk4_0)=k9_subset_1(esk3_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_96, c_0_102])).
cnf(c_0_114, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,X5,X3)=X6|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X4)|v1_xboole_0(X1)|X5!=k1_tmap_1(X1,X4,X2,X3,X7,X6)|~v1_funct_1(X5)|~v1_funct_2(X5,k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X2,X3),X4)))|k2_partfun1(X2,X4,X7,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_1(X6)|~v1_funct_2(X6,X3,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X7)|~v1_funct_2(X7,X2,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_115, plain, (m1_subset_1(k1_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_funct_2(X6,X4,X2)|~v1_funct_2(X5,X3,X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_116, plain, (v1_funct_1(k1_tmap_1(X1,X2,X3,X4,X5,X6))|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_funct_2(X6,X4,X2)|~v1_funct_2(X5,X3,X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[c_0_106, c_0_105])).
cnf(c_0_117, plain, (v1_funct_2(k1_tmap_1(X1,X2,X3,X4,X5,X6),k4_subset_1(X1,X3,X4),X2)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~v1_funct_2(X6,X4,X2)|~v1_funct_2(X5,X3,X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[c_0_107, c_0_105])).
cnf(c_0_118, negated_conjecture, (k7_relat_1(k7_relat_1(X1,esk4_0),esk3_0)=k9_subset_1(esk3_0,esk4_0,esk3_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_108, c_0_77])).
cnf(c_0_119, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,esk4_0)=esk6_0), inference(rw,[status(thm)],[c_0_109, c_0_83])).
cnf(c_0_120, plain, (k3_xboole_0(X1,X2)=k9_subset_1(esk3_0,esk4_0,esk3_0)|~v1_xboole_0(X2)), inference(rw,[status(thm)],[c_0_110, c_0_96])).
cnf(c_0_121, negated_conjecture, (k7_relat_1(esk5_0,esk3_0)=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_73]), c_0_74])])).
cnf(c_0_122, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk4_0,esk2_0,esk6_0,esk3_0)!=k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk3_0,esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_111, c_0_96])).
cnf(c_0_123, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk3_0,esk4_0,esk3_0))=k2_partfun1(esk3_0,esk2_0,esk5_0,esk4_0)), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_124, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,k1_tmap_1(X1,X4,X2,X3,X5,X6),X3)=X6|v1_xboole_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|k2_partfun1(X2,X4,X5,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_2(X5,X2,X4)|~v1_funct_2(X6,X3,X4)|~v1_funct_1(X5)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_114, c_0_105])]), c_0_115]), c_0_116]), c_0_117])).
cnf(c_0_125, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,esk3_0)=k2_partfun1(esk4_0,esk2_0,esk6_0,X1)|~v1_xboole_0(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_110]), c_0_100])).
cnf(c_0_126, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,esk3_0)=k9_subset_1(esk3_0,esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_83]), c_0_119]), c_0_83]), c_0_70])])).
cnf(c_0_127, plain, (k9_subset_1(esk3_0,esk4_0,esk3_0)=k9_subset_1(X1,X2,X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_79, c_0_120])).
cnf(c_0_128, negated_conjecture, (k7_relat_1(k7_relat_1(X1,esk3_0),esk4_0)=k9_subset_1(esk3_0,esk4_0,esk3_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_108, c_0_32])).
cnf(c_0_129, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,esk3_0)=esk5_0), inference(rw,[status(thm)],[c_0_121, c_0_92])).
cnf(c_0_130, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_131, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,X5,X2)=X6|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X4)|v1_xboole_0(X1)|X5!=k1_tmap_1(X1,X4,X2,X3,X6,X7)|~v1_funct_1(X5)|~v1_funct_2(X5,k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X2,X3),X4)))|k2_partfun1(X2,X4,X6,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X7,k9_subset_1(X1,X2,X3))|~v1_funct_1(X7)|~v1_funct_2(X7,X3,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X6)|~v1_funct_2(X6,X2,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_132, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk4_0,esk2_0,esk6_0,esk3_0)!=k2_partfun1(esk3_0,esk2_0,esk5_0,esk4_0)), inference(rw,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_133, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,X1,esk4_0),X2,k1_tmap_1(esk1_0,X2,X1,esk4_0,X3,X4),esk4_0)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k2_partfun1(X1,X2,X3,k9_subset_1(esk4_0,X1,esk4_0))!=k2_partfun1(esk4_0,X2,X4,k9_subset_1(esk4_0,X1,esk4_0))|~v1_funct_2(X4,esk4_0,X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk4_0,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_102]), c_0_34])]), c_0_26])).
cnf(c_0_134, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,X1)=k9_subset_1(esk3_0,esk4_0,esk3_0)|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_135, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,esk4_0)=k2_partfun1(esk3_0,esk2_0,esk5_0,X1)|~v1_xboole_0(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_127]), c_0_123])).
cnf(c_0_136, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,esk4_0)=k9_subset_1(esk3_0,esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_92]), c_0_129]), c_0_92]), c_0_74])])).
cnf(c_0_137, plain, (v1_xboole_0(k9_subset_1(esk1_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[c_0_130, c_0_61])).
cnf(c_0_138, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,k1_tmap_1(X1,X4,X2,X3,X5,X6),X2)=X5|v1_xboole_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|k2_partfun1(X2,X4,X5,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_2(X6,X3,X4)|~v1_funct_2(X5,X2,X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_131, c_0_105])]), c_0_115]), c_0_116]), c_0_117])).
cnf(c_0_139, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk3_0,esk2_0,esk5_0,esk4_0)!=k9_subset_1(esk3_0,esk4_0,esk3_0)), inference(rw,[status(thm)],[c_0_132, c_0_126])).
cnf(c_0_140, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,X1,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0),esk4_0)=esk6_0|v1_xboole_0(X1)|k2_partfun1(X1,esk2_0,X2,k9_subset_1(esk4_0,X1,esk4_0))!=k9_subset_1(esk3_0,esk4_0,esk3_0)|~v1_funct_2(X2,X1,esk2_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~v1_xboole_0(k9_subset_1(esk4_0,X1,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_141, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,X1)=k9_subset_1(esk3_0,esk4_0,esk3_0)|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_142, plain, (v1_xboole_0(k9_subset_1(esk3_0,esk4_0,esk3_0))), inference(rw,[status(thm)],[c_0_137, c_0_96])).
cnf(c_0_143, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,X1,esk4_0),X2,k1_tmap_1(esk1_0,X2,X1,esk4_0,X3,X4),X1)=X3|v1_xboole_0(X2)|v1_xboole_0(X1)|k2_partfun1(X1,X2,X3,k9_subset_1(esk4_0,X1,esk4_0))!=k2_partfun1(esk4_0,X2,X4,k9_subset_1(esk4_0,X1,esk4_0))|~v1_funct_2(X4,esk4_0,X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk4_0,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_102]), c_0_34])]), c_0_26])).
cnf(c_0_144, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_139, c_0_136])])).
cnf(c_0_145, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)=esk6_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_58]), c_0_59]), c_0_60]), c_0_63]), c_0_113]), c_0_142])]), c_0_27])).
cnf(c_0_146, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,X1,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0),X1)=X2|v1_xboole_0(X1)|k2_partfun1(X1,esk2_0,X2,k9_subset_1(esk4_0,X1,esk4_0))!=k9_subset_1(esk3_0,esk4_0,esk3_0)|~v1_funct_2(X2,X1,esk2_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~v1_xboole_0(k9_subset_1(esk4_0,X1,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_134]), c_0_51]), c_0_52]), c_0_53])]), c_0_54])).
cnf(c_0_147, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_144, c_0_145])])).
cnf(c_0_148, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_141]), c_0_58]), c_0_59]), c_0_60]), c_0_63]), c_0_113]), c_0_142])]), c_0_147]), c_0_27]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 149
# Proof object clause steps            : 109
# Proof object formula steps           : 40
# Proof object conjectures             : 67
# Proof object clause conjectures      : 64
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 37
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 42
# Proof object simplifying inferences  : 116
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 47
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 46
# Processed clauses                    : 3152
# ...of these trivial                  : 30
# ...subsumed                          : 2125
# ...remaining for further processing  : 997
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 27
# Backward-rewritten                   : 121
# Generated clauses                    : 7902
# ...of the previous two non-trivial   : 6834
# Contextual simplify-reflections      : 24
# Paramodulations                      : 7870
# Factorizations                       : 0
# Equation resolutions                 : 33
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 798
#    Positive orientable unit clauses  : 137
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 656
# Current number of unprocessed clauses: 3611
# ...number of literals in the above   : 23647
# Current number of archived formulas  : 0
# Current number of archived clauses   : 193
# Clause-clause subsumption calls (NU) : 136592
# Rec. Clause-clause subsumption calls : 48351
# Non-unit clause-clause subsumptions  : 1134
# Unit Clause-clause subsumption calls : 2972
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 369
# BW rewrite match successes           : 59
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 246151

# -------------------------------------------------
# User time                : 0.166 s
# System time              : 0.000 s
# Total time               : 0.166 s
# Maximum resident set size: 28804 pages
