# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B02_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t131_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k8_relat_1(X1,X3),k8_relat_1(X2,X3)))), file('data/MPTP/proofs/97c07bbaf54547e6b2415e0a4076431e.E_input', t131_relat_1)).
fof(d10_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k6_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>(r2_hidden(X3,X1)&X3=X4)))), file('data/MPTP/proofs/97c07bbaf54547e6b2415e0a4076431e.E_input', d10_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/97c07bbaf54547e6b2415e0a4076431e.E_input', d3_tarski)).
fof(t48_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k5_relat_1(X3,X1),k5_relat_1(X3,X2)))))), file('data/MPTP/proofs/97c07bbaf54547e6b2415e0a4076431e.E_input', t48_relat_1)).
fof(t123_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k8_relat_1(X1,X2)=k5_relat_1(X2,k6_relat_1(X1))), file('data/MPTP/proofs/97c07bbaf54547e6b2415e0a4076431e.E_input', t123_relat_1)).
fof(dt_k6_relat_1, axiom, ![X1]:v1_relat_1(k6_relat_1(X1)), file('data/MPTP/proofs/97c07bbaf54547e6b2415e0a4076431e.E_input', dt_k6_relat_1)).
fof(t73_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(![X3]:(r2_hidden(X3,X1)=>r2_hidden(k4_tarski(X3,X3),X2))=>r1_tarski(k6_relat_1(X1),X2))), file('data/MPTP/proofs/97c07bbaf54547e6b2415e0a4076431e.E_input', t73_relat_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k8_relat_1(X1,X3),k8_relat_1(X2,X3))))), inference(assume_negation,[status(cth)],[t131_relat_1])).
fof(c_0_8, plain, ![X16, X17, X18, X19, X20, X21]:((((r2_hidden(X18,X16)|~r2_hidden(k4_tarski(X18,X19),X17)|X17!=k6_relat_1(X16)|~v1_relat_1(X17))&(X18=X19|~r2_hidden(k4_tarski(X18,X19),X17)|X17!=k6_relat_1(X16)|~v1_relat_1(X17)))&(~r2_hidden(X20,X16)|X20!=X21|r2_hidden(k4_tarski(X20,X21),X17)|X17!=k6_relat_1(X16)|~v1_relat_1(X17)))&((~r2_hidden(k4_tarski(esk5_2(X16,X17),esk6_2(X16,X17)),X17)|(~r2_hidden(esk5_2(X16,X17),X16)|esk5_2(X16,X17)!=esk6_2(X16,X17))|X17=k6_relat_1(X16)|~v1_relat_1(X17))&((r2_hidden(esk5_2(X16,X17),X16)|r2_hidden(k4_tarski(esk5_2(X16,X17),esk6_2(X16,X17)),X17)|X17=k6_relat_1(X16)|~v1_relat_1(X17))&(esk5_2(X16,X17)=esk6_2(X16,X17)|r2_hidden(k4_tarski(esk5_2(X16,X17),esk6_2(X16,X17)),X17)|X17=k6_relat_1(X16)|~v1_relat_1(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_relat_1])])])])])])).
fof(c_0_9, plain, ![X25, X26, X27, X28, X29]:((~r1_tarski(X25,X26)|(~r2_hidden(X27,X25)|r2_hidden(X27,X26)))&((r2_hidden(esk7_2(X28,X29),X28)|r1_tarski(X28,X29))&(~r2_hidden(esk7_2(X28,X29),X29)|r1_tarski(X28,X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_10, negated_conjecture, (v1_relat_1(esk3_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k8_relat_1(esk1_0,esk3_0),k8_relat_1(esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_11, plain, ![X13, X14, X15]:(~v1_relat_1(X13)|(~v1_relat_1(X14)|(~v1_relat_1(X15)|(~r1_tarski(X13,X14)|r1_tarski(k5_relat_1(X15,X13),k5_relat_1(X15,X14)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t48_relat_1])])])).
fof(c_0_12, plain, ![X8, X9]:(~v1_relat_1(X9)|k8_relat_1(X8,X9)=k5_relat_1(X9,k6_relat_1(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t123_relat_1])])).
fof(c_0_13, plain, ![X24]:v1_relat_1(k6_relat_1(X24)), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)|X1!=X3|X4!=k6_relat_1(X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (r1_tarski(k5_relat_1(X3,X1),k5_relat_1(X3,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (k8_relat_1(X2,X1)=k5_relat_1(X1,k6_relat_1(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_20, plain, ![X10, X11]:((r2_hidden(esk4_2(X10,X11),X10)|r1_tarski(k6_relat_1(X10),X11)|~v1_relat_1(X11))&(~r2_hidden(k4_tarski(esk4_2(X10,X11),esk4_2(X10,X11)),X11)|r1_tarski(k6_relat_1(X10),X11)|~v1_relat_1(X11))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_relat_1])])])])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X1,X1),X2)|X2!=k6_relat_1(X3)|~r2_hidden(X1,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, plain, (r1_tarski(k5_relat_1(X1,X2),k8_relat_1(X3,X1))|~r1_tarski(X2,k6_relat_1(X3))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
cnf(c_0_24, plain, (r1_tarski(k6_relat_1(X1),X2)|~r2_hidden(k4_tarski(esk4_2(X1,X2),esk4_2(X1,X2)),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (r2_hidden(k4_tarski(X1,X1),X2)|X2!=k6_relat_1(esk2_0)|~r2_hidden(X1,esk1_0)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26, negated_conjecture, (~r1_tarski(k8_relat_1(esk1_0,esk3_0),k8_relat_1(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_27, plain, (r1_tarski(k8_relat_1(X1,X2),k8_relat_1(X3,X2))|~r1_tarski(k6_relat_1(X1),k6_relat_1(X3))|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_18]), c_0_19])])).
cnf(c_0_28, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29, negated_conjecture, (r1_tarski(k6_relat_1(X1),X2)|X2!=k6_relat_1(esk2_0)|~r2_hidden(esk4_2(X1,X2),esk1_0)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(k6_relat_1(X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, negated_conjecture, (~r1_tarski(k6_relat_1(esk1_0),k6_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_32, negated_conjecture, (r1_tarski(k6_relat_1(esk1_0),X1)|X1!=k6_relat_1(esk2_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_19])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 34
# Proof object clause steps            : 19
# Proof object formula steps           : 15
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 79
# ...of these trivial                  : 1
# ...subsumed                          : 3
# ...remaining for further processing  : 75
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 173
# ...of the previous two non-trivial   : 164
# Contextual simplify-reflections      : 0
# Paramodulations                      : 170
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 118
# ...number of literals in the above   : 736
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 653
# Rec. Clause-clause subsumption calls : 181
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 37
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4709

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28616 pages
