# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t20_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k6_waybel_0(X1,X2)=k6_waybel_0(X1,X3)=>X2=X3)))), file('data/MPTP/proofs/993b4c6b205f420e966c6f227a0b4c85.E_input', t20_waybel_0)).
fof(t18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/993b4c6b205f420e966c6f227a0b4c85.E_input', t18_waybel_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/993b4c6b205f420e966c6f227a0b4c85.E_input', d8_lattice3)).
fof(t31_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/993b4c6b205f420e966c6f227a0b4c85.E_input', t31_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/993b4c6b205f420e966c6f227a0b4c85.E_input', t24_orders_2)).
fof(c_0_5, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(k6_waybel_0(X1,X2)=k6_waybel_0(X1,X3)=>X2=X3))))), inference(assume_negation,[status(cth)],[t20_waybel_0])).
fof(c_0_6, plain, ![X8, X9, X10]:((~r2_hidden(X10,k6_waybel_0(X8,X9))|r1_orders_2(X8,X9,X10)|~m1_subset_1(X10,u1_struct_0(X8))|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l1_orders_2(X8)))&(~r1_orders_2(X8,X9,X10)|r2_hidden(X10,k6_waybel_0(X8,X9))|~m1_subset_1(X10,u1_struct_0(X8))|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l1_orders_2(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])).
fof(c_0_7, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(k6_waybel_0(esk1_0,esk2_0)=k6_waybel_0(esk1_0,esk3_0)&esk2_0!=esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])).
cnf(c_0_8, plain, (r1_orders_2(X2,X3,X1)|v2_struct_0(X2)|~r2_hidden(X1,k6_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, negated_conjecture, (k6_waybel_0(esk1_0,esk2_0)=k6_waybel_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_13, plain, ![X11, X12, X13, X14]:((~r1_lattice3(X11,X12,X13)|(~m1_subset_1(X14,u1_struct_0(X11))|(~r2_hidden(X14,X12)|r1_orders_2(X11,X13,X14)))|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&((m1_subset_1(esk4_3(X11,X12,X13),u1_struct_0(X11))|r1_lattice3(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&((r2_hidden(esk4_3(X11,X12,X13),X12)|r1_lattice3(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))&(~r1_orders_2(X11,X13,esk4_3(X11,X12,X13))|r1_lattice3(X11,X12,X13)|~m1_subset_1(X13,u1_struct_0(X11))|~l1_orders_2(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_14, negated_conjecture, (r1_orders_2(esk1_0,esk3_0,X1)|~r2_hidden(X1,k6_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_9]), c_0_10]), c_0_11])]), c_0_12])).
cnf(c_0_15, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_16, plain, ![X16, X17, X18, X19, X20]:(((r1_lattice3(X16,X18,X17)|(X17!=k2_yellow_0(X16,X18)|~r2_yellow_0(X16,X18))|~m1_subset_1(X17,u1_struct_0(X16))|(~v5_orders_2(X16)|~l1_orders_2(X16)))&(~m1_subset_1(X19,u1_struct_0(X16))|(~r1_lattice3(X16,X18,X19)|r1_orders_2(X16,X19,X17))|(X17!=k2_yellow_0(X16,X18)|~r2_yellow_0(X16,X18))|~m1_subset_1(X17,u1_struct_0(X16))|(~v5_orders_2(X16)|~l1_orders_2(X16))))&(((X17=k2_yellow_0(X16,X20)|(m1_subset_1(esk5_3(X16,X17,X20),u1_struct_0(X16))|~r1_lattice3(X16,X20,X17))|~m1_subset_1(X17,u1_struct_0(X16))|(~v5_orders_2(X16)|~l1_orders_2(X16)))&(r2_yellow_0(X16,X20)|(m1_subset_1(esk5_3(X16,X17,X20),u1_struct_0(X16))|~r1_lattice3(X16,X20,X17))|~m1_subset_1(X17,u1_struct_0(X16))|(~v5_orders_2(X16)|~l1_orders_2(X16))))&(((X17=k2_yellow_0(X16,X20)|(r1_lattice3(X16,X20,esk5_3(X16,X17,X20))|~r1_lattice3(X16,X20,X17))|~m1_subset_1(X17,u1_struct_0(X16))|(~v5_orders_2(X16)|~l1_orders_2(X16)))&(r2_yellow_0(X16,X20)|(r1_lattice3(X16,X20,esk5_3(X16,X17,X20))|~r1_lattice3(X16,X20,X17))|~m1_subset_1(X17,u1_struct_0(X16))|(~v5_orders_2(X16)|~l1_orders_2(X16))))&((X17=k2_yellow_0(X16,X20)|(~r1_orders_2(X16,esk5_3(X16,X17,X20),X17)|~r1_lattice3(X16,X20,X17))|~m1_subset_1(X17,u1_struct_0(X16))|(~v5_orders_2(X16)|~l1_orders_2(X16)))&(r2_yellow_0(X16,X20)|(~r1_orders_2(X16,esk5_3(X16,X17,X20),X17)|~r1_lattice3(X16,X20,X17))|~m1_subset_1(X17,u1_struct_0(X16))|(~v5_orders_2(X16)|~l1_orders_2(X16))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])).
cnf(c_0_17, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (r1_lattice3(X1,k6_waybel_0(esk1_0,esk2_0),X2)|r1_orders_2(esk1_0,esk3_0,esk4_3(X1,k6_waybel_0(esk1_0,esk2_0),X2))|~m1_subset_1(esk4_3(X1,k6_waybel_0(esk1_0,esk2_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (X1=k2_yellow_0(X2,X3)|r1_lattice3(X2,X3,esk5_3(X2,X1,X3))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21, negated_conjecture, (r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk3_0)|~m1_subset_1(esk4_3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_10]), c_0_11])])).
cnf(c_0_22, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, plain, (X1=k2_yellow_0(X2,X3)|m1_subset_1(esk5_3(X2,X1,X3),u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X22, X23]:(v2_struct_0(X22)|(~v3_orders_2(X22)|~l1_orders_2(X22))|(~m1_subset_1(X23,u1_struct_0(X22))|r1_orders_2(X22,X23,X23))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_25, negated_conjecture, (esk3_0=k2_yellow_0(esk1_0,X1)|r1_lattice3(esk1_0,X1,esk5_3(esk1_0,esk3_0,X1))|~r1_lattice3(esk1_0,X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_10]), c_0_11]), c_0_20])])).
cnf(c_0_26, negated_conjecture, (r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_10]), c_0_11])])).
cnf(c_0_27, negated_conjecture, (esk3_0=k2_yellow_0(esk1_0,X1)|m1_subset_1(esk5_3(esk1_0,esk3_0,X1),u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_10]), c_0_11]), c_0_20])])).
cnf(c_0_28, plain, (r2_hidden(X3,k6_waybel_0(X1,X2))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_29, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32, negated_conjecture, (esk3_0=k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))|r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk5_3(esk1_0,esk3_0,k6_waybel_0(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_33, negated_conjecture, (esk3_0=k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))|m1_subset_1(esk5_3(esk1_0,esk3_0,k6_waybel_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_26])).
cnf(c_0_34, negated_conjecture, (r2_hidden(X1,k6_waybel_0(esk1_0,esk2_0))|~r1_orders_2(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_9]), c_0_10]), c_0_11])]), c_0_12])).
cnf(c_0_35, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_11])]), c_0_12])).
cnf(c_0_36, negated_conjecture, (esk3_0=k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))|r1_orders_2(esk1_0,esk5_3(esk1_0,esk3_0,k6_waybel_0(esk1_0,esk2_0)),X1)|~r2_hidden(X1,k6_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_11])]), c_0_33])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk3_0,k6_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_10])])).
cnf(c_0_38, plain, (X1=k2_yellow_0(X2,X3)|~r1_orders_2(X2,esk5_3(X2,X1,X3),X1)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_39, negated_conjecture, (esk3_0=k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))|r1_orders_2(esk1_0,esk5_3(esk1_0,esk3_0,k6_waybel_0(esk1_0,esk2_0)),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_10])])).
cnf(c_0_40, negated_conjecture, (esk3_0=k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_26]), c_0_10]), c_0_11]), c_0_20])])).
cnf(c_0_41, negated_conjecture, (r2_hidden(X1,k6_waybel_0(esk1_0,esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_34, c_0_40])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_43, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_8, c_0_41]), c_0_42]), c_0_11])]), c_0_12])).
cnf(c_0_44, plain, (r1_lattice3(X1,k6_waybel_0(X2,X3),X4)|r1_orders_2(X2,X3,esk4_3(X1,k6_waybel_0(X2,X3),X4))|v2_struct_0(X2)|~m1_subset_1(esk4_3(X1,k6_waybel_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_8, c_0_15])).
cnf(c_0_45, negated_conjecture, (k6_waybel_0(esk1_0,k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0)))=k6_waybel_0(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_9, c_0_40])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_10, c_0_40])).
cnf(c_0_47, negated_conjecture, (r1_lattice3(X1,k6_waybel_0(esk1_0,esk2_0),X2)|r1_orders_2(esk1_0,esk2_0,esk4_3(X1,k6_waybel_0(esk1_0,esk2_0),X2))|~m1_subset_1(esk4_3(X1,k6_waybel_0(esk1_0,esk2_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_45]), c_0_45]), c_0_46]), c_0_11])]), c_0_12])).
cnf(c_0_48, negated_conjecture, (r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk2_0)|~m1_subset_1(esk4_3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_47]), c_0_42]), c_0_11])])).
cnf(c_0_49, negated_conjecture, (esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_50, negated_conjecture, (esk2_0=k2_yellow_0(esk1_0,X1)|r1_lattice3(esk1_0,X1,esk5_3(esk1_0,esk2_0,X1))|~r1_lattice3(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_42]), c_0_11]), c_0_20])])).
cnf(c_0_51, negated_conjecture, (r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_22]), c_0_42]), c_0_11])])).
cnf(c_0_52, negated_conjecture, (k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))!=esk2_0), inference(rw,[status(thm)],[c_0_49, c_0_40])).
cnf(c_0_53, negated_conjecture, (esk2_0=k2_yellow_0(esk1_0,X1)|m1_subset_1(esk5_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_lattice3(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_42]), c_0_11]), c_0_20])])).
cnf(c_0_54, negated_conjecture, (r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk5_3(esk1_0,esk2_0,k6_waybel_0(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,esk2_0,k6_waybel_0(esk1_0,esk2_0)),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_51]), c_0_52])).
cnf(c_0_56, negated_conjecture, (r1_orders_2(esk1_0,esk5_3(esk1_0,esk2_0,k6_waybel_0(esk1_0,esk2_0)),X1)|~r2_hidden(X1,k6_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_54]), c_0_55]), c_0_11])])).
cnf(c_0_57, negated_conjecture, (r1_orders_2(esk1_0,esk5_3(esk1_0,esk2_0,k6_waybel_0(esk1_0,esk2_0)),X1)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_28]), c_0_42]), c_0_11])]), c_0_12])).
cnf(c_0_58, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_57]), c_0_51]), c_0_42]), c_0_11]), c_0_20])]), c_0_52])).
cnf(c_0_59, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_35]), c_0_42])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 49
# Proof object formula steps           : 11
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 79
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 173
# ...of these trivial                  : 6
# ...subsumed                          : 33
# ...remaining for further processing  : 134
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 45
# Generated clauses                    : 170
# ...of the previous two non-trivial   : 171
# Contextual simplify-reflections      : 2
# Paramodulations                      : 167
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 65
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 46
# Current number of unprocessed clauses: 31
# ...number of literals in the above   : 140
# Current number of archived formulas  : 0
# Current number of archived clauses   : 69
# Clause-clause subsumption calls (NU) : 811
# Rec. Clause-clause subsumption calls : 350
# Non-unit clause-clause subsumptions  : 36
# Unit Clause-clause subsumption calls : 68
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7959

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28656 pages
