# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_relset_1, conjecture, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X5,X4),X3))))<=>k2_relset_1(X1,X2,X3)=X2)), file('data/MPTP/proofs/99eae35107ca496cb82a9d1262e3940e.E_input', t23_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/99eae35107ca496cb82a9d1262e3940e.E_input', t3_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/99eae35107ca496cb82a9d1262e3940e.E_input', d3_tarski)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/99eae35107ca496cb82a9d1262e3940e.E_input', redefinition_k2_relset_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/99eae35107ca496cb82a9d1262e3940e.E_input', l54_zfmisc_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/99eae35107ca496cb82a9d1262e3940e.E_input', d5_relat_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X5,X4),X3))))<=>k2_relset_1(X1,X2,X3)=X2))), inference(assume_negation,[status(cth)],[t23_relset_1])).
fof(c_0_7, plain, ![X13, X14]:((~m1_subset_1(X13,k1_zfmisc_1(X14))|r1_tarski(X13,X14))&(~r1_tarski(X13,X14)|m1_subset_1(X13,k1_zfmisc_1(X14)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_8, negated_conjecture, ![X10, X11]:(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))&(((r2_hidden(esk4_0,esk2_0)|k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0)&(~r2_hidden(k4_tarski(X10,esk4_0),esk3_0)|k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0))&(~r2_hidden(X11,esk2_0)|r2_hidden(k4_tarski(esk5_1(X11),X11),esk3_0)|k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])])).
fof(c_0_9, plain, ![X15, X16, X17, X18, X19]:((~r1_tarski(X15,X16)|(~r2_hidden(X17,X15)|r2_hidden(X17,X16)))&((r2_hidden(esk6_2(X18,X19),X18)|r1_tarski(X18,X19))&(~r2_hidden(esk6_2(X18,X19),X19)|r1_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_10, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X21, X22, X23]:(~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(X21,X22)))|k2_relset_1(X21,X22,X23)=k2_relat_1(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_13, plain, ![X24, X25, X26, X27]:(((r2_hidden(X24,X26)|~r2_hidden(k4_tarski(X24,X25),k2_zfmisc_1(X26,X27)))&(r2_hidden(X25,X27)|~r2_hidden(k4_tarski(X24,X25),k2_zfmisc_1(X26,X27))))&(~r2_hidden(X24,X26)|~r2_hidden(X25,X27)|r2_hidden(k4_tarski(X24,X25),k2_zfmisc_1(X26,X27)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_14, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_16, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_19, plain, ![X28, X29, X30, X32, X33, X34, X35, X37]:(((~r2_hidden(X30,X29)|r2_hidden(k4_tarski(esk7_3(X28,X29,X30),X30),X28)|X29!=k2_relat_1(X28))&(~r2_hidden(k4_tarski(X33,X32),X28)|r2_hidden(X32,X29)|X29!=k2_relat_1(X28)))&((~r2_hidden(esk8_2(X34,X35),X35)|~r2_hidden(k4_tarski(X37,esk8_2(X34,X35)),X34)|X35=k2_relat_1(X34))&(r2_hidden(esk8_2(X34,X35),X35)|r2_hidden(k4_tarski(esk9_2(X34,X35),esk8_2(X34,X35)),X34)|X35=k2_relat_1(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_20, negated_conjecture, (r2_hidden(k4_tarski(esk5_1(X1),X1),esk3_0)|k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)=k2_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_16, c_0_11])).
cnf(c_0_22, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(k4_tarski(X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, plain, (r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(k4_tarski(esk9_2(X1,X2),esk8_2(X1,X2)),X1)|X2=k2_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (~r2_hidden(k4_tarski(X1,esk4_0),esk3_0)|k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(esk7_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (X2=k2_relat_1(X1)|~r2_hidden(esk8_2(X1,X2),X2)|~r2_hidden(k4_tarski(X3,esk8_2(X1,X2)),X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (k2_relat_1(esk3_0)=esk2_0|r2_hidden(k4_tarski(esk5_1(X1),X1),esk3_0)|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, negated_conjecture, (X1=k2_relat_1(esk3_0)|r2_hidden(esk8_2(esk3_0,X1),esk2_0)|r2_hidden(esk8_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29, negated_conjecture, (k2_relat_1(esk3_0)!=esk2_0|~r2_hidden(k4_tarski(X1,esk4_0),esk3_0)), inference(rw,[status(thm)],[c_0_24, c_0_21])).
cnf(c_0_30, plain, (r2_hidden(k4_tarski(esk7_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (k2_relat_1(esk3_0)=esk2_0|X1=k2_relat_1(esk3_0)|~r2_hidden(esk8_2(esk3_0,X1),esk2_0)|~r2_hidden(esk8_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, negated_conjecture, (k2_relat_1(esk3_0)=esk2_0|r2_hidden(esk8_2(esk3_0,esk2_0),esk2_0)), inference(ef,[status(thm)],[c_0_28])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk4_0,esk2_0)|k2_relset_1(esk1_0,esk2_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_34, negated_conjecture, (k2_relat_1(esk3_0)!=esk2_0|~r2_hidden(esk4_0,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, negated_conjecture, (k2_relat_1(esk3_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_32])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk4_0,esk2_0)|k2_relat_1(esk3_0)!=esk2_0), inference(rw,[status(thm)],[c_0_33, c_0_21])).
cnf(c_0_37, negated_conjecture, (~r2_hidden(esk4_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35]), c_0_35])])).
cnf(c_0_38, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_35])]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 26
# Proof object formula steps           : 13
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 209
# ...of these trivial                  : 0
# ...subsumed                          : 48
# ...remaining for further processing  : 160
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 66
# Generated clauses                    : 1448
# ...of the previous two non-trivial   : 1453
# Contextual simplify-reflections      : 2
# Paramodulations                      : 1438
# Factorizations                       : 8
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 70
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 63
# Current number of unprocessed clauses: 1274
# ...number of literals in the above   : 4308
# Current number of archived formulas  : 0
# Current number of archived clauses   : 88
# Clause-clause subsumption calls (NU) : 2594
# Rec. Clause-clause subsumption calls : 2164
# Non-unit clause-clause subsumptions  : 54
# Unit Clause-clause subsumption calls : 36
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 21219

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.006 s
# Total time               : 0.023 s
# Maximum resident set size: 28608 pages
