# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t21_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t48_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t3_boole)).
fof(t40_xboole_1, axiom, ![X1, X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t40_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', d10_xboole_0)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t106_xboole_1)).
fof(t38_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,k4_xboole_0(X2,X1))=>X1=k1_xboole_0), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t38_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', d5_xboole_0)).
fof(t55_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_xboole_0(k2_tarski(X1,X2),X3)&r2_hidden(X1,X3))), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t55_zfmisc_1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', l1_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t69_enumset1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', commutativity_k3_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', commutativity_k2_xboole_0)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t39_xboole_1)).
fof(t63_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X2,X3))=>r1_xboole_0(X1,X3)), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t63_xboole_1)).
fof(t103_zfmisc_1, conjecture, ![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6))))))), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', t103_zfmisc_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/9a623cf10e204de381c3bd7bb1985af4.E_input', d2_zfmisc_1)).
fof(c_0_17, plain, ![X77, X78]:k3_xboole_0(X77,k2_xboole_0(X77,X78))=X77, inference(variable_rename,[status(thm)],[t21_xboole_1])).
fof(c_0_18, plain, ![X50, X51]:k4_xboole_0(X50,k4_xboole_0(X50,X51))=k3_xboole_0(X50,X51), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_19, plain, ![X43]:k4_xboole_0(X43,k1_xboole_0)=X43, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_20, plain, ![X75, X76]:k4_xboole_0(k2_xboole_0(X75,X76),X76)=k4_xboole_0(X75,X76), inference(variable_rename,[status(thm)],[t40_xboole_1])).
fof(c_0_21, plain, ![X46, X47]:(((r1_tarski(X46,X47)|X46!=X47)&(r1_tarski(X47,X46)|X46!=X47))&(~r1_tarski(X46,X47)|~r1_tarski(X47,X46)|X46=X47)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_22, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X57, X58, X59]:((r1_tarski(X57,X58)|~r1_tarski(X57,k4_xboole_0(X58,X59)))&(r1_xboole_0(X57,X59)|~r1_tarski(X57,k4_xboole_0(X58,X59)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X41, X42]:(~r1_tarski(X41,k4_xboole_0(X42,X41))|X41=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_xboole_1])])).
cnf(c_0_29, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k2_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_30, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_24])).
cnf(c_0_31, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (X1=k1_xboole_0|~r1_tarski(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X1))=X1), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_37, plain, ![X63, X64, X65, X66, X67, X68, X69, X70]:((((r2_hidden(X66,X63)|~r2_hidden(X66,X65)|X65!=k4_xboole_0(X63,X64))&(~r2_hidden(X66,X64)|~r2_hidden(X66,X65)|X65!=k4_xboole_0(X63,X64)))&(~r2_hidden(X67,X63)|r2_hidden(X67,X64)|r2_hidden(X67,X65)|X65!=k4_xboole_0(X63,X64)))&((~r2_hidden(esk6_3(X68,X69,X70),X70)|(~r2_hidden(esk6_3(X68,X69,X70),X68)|r2_hidden(esk6_3(X68,X69,X70),X69))|X70=k4_xboole_0(X68,X69))&((r2_hidden(esk6_3(X68,X69,X70),X68)|r2_hidden(esk6_3(X68,X69,X70),X70)|X70=k4_xboole_0(X68,X69))&(~r2_hidden(esk6_3(X68,X69,X70),X69)|r2_hidden(esk6_3(X68,X69,X70),X70)|X70=k4_xboole_0(X68,X69))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_38, plain, ![X60, X61, X62]:(~r1_xboole_0(k2_tarski(X60,X61),X62)|~r2_hidden(X60,X62)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_zfmisc_1])])).
cnf(c_0_39, plain, (r1_xboole_0(X1,k1_xboole_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_33, c_0_24])).
cnf(c_0_40, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_43, plain, (~r1_xboole_0(k2_tarski(X1,X2),X3)|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_39, c_0_32])).
cnf(c_0_45, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_34, c_0_40])).
fof(c_0_46, plain, ![X55, X56]:((~r1_tarski(k1_tarski(X55),X56)|r2_hidden(X55,X56))&(~r2_hidden(X55,X56)|r1_tarski(k1_tarski(X55),X56))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_47, plain, ![X52]:k2_tarski(X52,X52)=k1_tarski(X52), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_48, plain, ![X44, X45]:k3_xboole_0(X44,X45)=k3_xboole_0(X45,X44), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_49, plain, ![X53, X54]:k2_xboole_0(X53,X54)=k2_xboole_0(X54,X53), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_50, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_51, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_52, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_53, plain, (r2_hidden(esk6_3(X1,X2,X3),X1)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_54, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_45, c_0_36])).
fof(c_0_55, plain, ![X48, X49]:k2_xboole_0(X48,k4_xboole_0(X49,X48))=k2_xboole_0(X48,X49), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_56, plain, ![X72, X73, X74]:(~r1_tarski(X72,X73)|~r1_xboole_0(X73,X74)|r1_xboole_0(X72,X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_xboole_1])])).
cnf(c_0_57, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_59, negated_conjecture, ~(![X1, X2, X3, X4]:~((r1_tarski(X1,k2_zfmisc_1(X2,X3))&(r2_hidden(X4,X1)&![X5, X6]:~((r2_hidden(X5,X2)&(r2_hidden(X6,X3)&X4=k4_tarski(X5,X6)))))))), inference(assume_negation,[status(cth)],[t103_zfmisc_1])).
cnf(c_0_60, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_61, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_62, plain, (~r2_hidden(X1,k4_xboole_0(X2,k2_xboole_0(X2,X3)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_29]), c_0_51])).
cnf(c_0_63, plain, (X1=k1_xboole_0|r2_hidden(esk6_3(k1_xboole_0,X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_64, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_65, plain, (r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_66, plain, (r1_tarski(k2_tarski(X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_57, c_0_58])).
fof(c_0_67, negated_conjecture, ![X16, X17]:(r1_tarski(esk1_0,k2_zfmisc_1(esk2_0,esk3_0))&(r2_hidden(esk4_0,esk1_0)&(~r2_hidden(X16,esk2_0)|(~r2_hidden(X17,esk3_0)|esk4_0!=k4_tarski(X16,X17))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_59])])])])).
cnf(c_0_68, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_23]), c_0_23])).
cnf(c_0_69, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X1)=k4_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_25, c_0_61])).
cnf(c_0_70, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_71, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=k4_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_25, c_0_64])).
cnf(c_0_72, plain, (r1_xboole_0(k2_tarski(X1,X1),X2)|~r1_xboole_0(X3,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_74, negated_conjecture, (r1_tarski(esk1_0,k2_zfmisc_1(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_33, c_0_32])).
cnf(c_0_76, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_24]), c_0_71])).
fof(c_0_77, plain, ![X79, X80, X81, X82, X85, X86, X87, X88, X89, X90, X92, X93]:((((r2_hidden(esk7_4(X79,X80,X81,X82),X79)|~r2_hidden(X82,X81)|X81!=k2_zfmisc_1(X79,X80))&((r2_hidden(esk8_4(X79,X80,X81,X82),X80)|~r2_hidden(X82,X81)|X81!=k2_zfmisc_1(X79,X80))&(X82=k4_tarski(esk7_4(X79,X80,X81,X82),esk8_4(X79,X80,X81,X82))|~r2_hidden(X82,X81)|X81!=k2_zfmisc_1(X79,X80))))&(~r2_hidden(X86,X79)|(~r2_hidden(X87,X80)|X85!=k4_tarski(X86,X87))|r2_hidden(X85,X81)|X81!=k2_zfmisc_1(X79,X80)))&((~r2_hidden(esk9_3(X88,X89,X90),X90)|(~r2_hidden(X92,X88)|(~r2_hidden(X93,X89)|esk9_3(X88,X89,X90)!=k4_tarski(X92,X93)))|X90=k2_zfmisc_1(X88,X89))&((r2_hidden(esk10_3(X88,X89,X90),X88)|r2_hidden(esk9_3(X88,X89,X90),X90)|X90=k2_zfmisc_1(X88,X89))&((r2_hidden(esk11_3(X88,X89,X90),X89)|r2_hidden(esk9_3(X88,X89,X90),X90)|X90=k2_zfmisc_1(X88,X89))&(esk9_3(X88,X89,X90)=k4_tarski(esk10_3(X88,X89,X90),esk11_3(X88,X89,X90))|r2_hidden(esk9_3(X88,X89,X90),X90)|X90=k2_zfmisc_1(X88,X89)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_78, negated_conjecture, (r1_xboole_0(k2_tarski(esk4_0,esk4_0),X1)|~r1_xboole_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_79, negated_conjecture, (r1_xboole_0(esk1_0,X1)|~r1_xboole_0(k2_zfmisc_1(esk2_0,esk3_0),X1)), inference(spm,[status(thm)],[c_0_65, c_0_74])).
cnf(c_0_80, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_81, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_82, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_83, plain, (X1=k4_tarski(esk7_4(X2,X3,X4,X1),esk8_4(X2,X3,X4,X1))|~r2_hidden(X1,X4)|X4!=k2_zfmisc_1(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_84, negated_conjecture, (~r1_xboole_0(esk1_0,X1)|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_43, c_0_78])).
cnf(c_0_85, negated_conjecture, (r1_xboole_0(esk1_0,k4_xboole_0(X1,k2_zfmisc_1(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_86, plain, (~r2_hidden(X1,k4_xboole_0(X2,k2_tarski(X1,X3)))), inference(spm,[status(thm)],[c_0_43, c_0_80])).
cnf(c_0_87, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_81])).
cnf(c_0_88, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X1),X2)), inference(rw,[status(thm)],[c_0_82, c_0_58])).
cnf(c_0_89, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk3_0)|esk4_0!=k4_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_90, plain, (k4_tarski(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3),esk8_4(X1,X2,k2_zfmisc_1(X1,X2),X3))=X3|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_83])).
cnf(c_0_91, plain, (r2_hidden(esk8_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_92, negated_conjecture, (~r2_hidden(esk4_0,k4_xboole_0(X1,k2_zfmisc_1(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_93, plain, (r2_hidden(X1,k2_tarski(X1,X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_94, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_88, c_0_32])).
cnf(c_0_95, negated_conjecture, (~r2_hidden(esk8_4(X1,X2,k2_zfmisc_1(X1,X2),esk4_0),esk3_0)|~r2_hidden(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),esk4_0),esk2_0)|~r2_hidden(esk4_0,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_96, plain, (r2_hidden(esk8_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_91])).
cnf(c_0_97, plain, (r2_hidden(esk7_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_98, negated_conjecture, (r2_hidden(esk4_0,k2_zfmisc_1(esk2_0,esk3_0))|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_92, c_0_87])).
cnf(c_0_99, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_100, negated_conjecture, (~r2_hidden(esk7_4(X1,esk3_0,k2_zfmisc_1(X1,esk3_0),esk4_0),esk2_0)|~r2_hidden(esk4_0,k2_zfmisc_1(X1,esk3_0))), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_101, plain, (r2_hidden(esk7_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_97])).
cnf(c_0_102, negated_conjecture, (r2_hidden(esk4_0,k2_zfmisc_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_103, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 104
# Proof object clause steps            : 69
# Proof object formula steps           : 35
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 55
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 53
# Processed clauses                    : 1916
# ...of these trivial                  : 76
# ...subsumed                          : 1413
# ...remaining for further processing  : 427
# Other redundant clauses eliminated   : 30
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 67
# Generated clauses                    : 6859
# ...of the previous two non-trivial   : 5828
# Contextual simplify-reflections      : 2
# Paramodulations                      : 6833
# Factorizations                       : 6
# Equation resolutions                 : 30
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 283
#    Positive orientable unit clauses  : 81
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 22
#    Non-unit-clauses                  : 178
# Current number of unprocessed clauses: 3547
# ...number of literals in the above   : 8934
# Current number of archived formulas  : 0
# Current number of archived clauses   : 127
# Clause-clause subsumption calls (NU) : 10865
# Rec. Clause-clause subsumption calls : 8070
# Non-unit clause-clause subsumptions  : 578
# Unit Clause-clause subsumption calls : 797
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 217
# BW rewrite match successes           : 47
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 77053

# -------------------------------------------------
# User time                : 0.056 s
# System time              : 0.004 s
# Total time               : 0.060 s
# Maximum resident set size: 28632 pages
