# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t3_xboole_0)).
fof(t103_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t103_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t48_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', idempotence_k3_xboole_0)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t43_xboole_1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t98_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t7_xboole_1)).
fof(t67_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&(r1_tarski(X1,X3)&r1_xboole_0(X2,X3)))=>X1=k1_xboole_0), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t67_xboole_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t65_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t5_boole)).
fof(t9_zfmisc_1, conjecture, ![X1, X2, X3]:(k1_tarski(X1)=k2_tarski(X2,X3)=>X2=X3), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t9_zfmisc_1)).
fof(t41_enumset1, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2)), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t41_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t70_enumset1)).
fof(t63_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X2),k4_enumset1(X3,X4,X5,X6,X7,X8)), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t63_enumset1)).
fof(t100_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X2,X3,X1), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t100_enumset1)).
fof(t98_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t98_enumset1)).
fof(t95_enumset1, axiom, ![X1, X2]:k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t95_enumset1)).
fof(l57_enumset1, axiom, ![X1, X2, X3, X4, X5]:k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_tarski(X4,X5)), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', l57_enumset1)).
fof(t86_enumset1, axiom, ![X1, X2, X3, X4, X5]:k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t86_enumset1)).
fof(t136_enumset1, axiom, ![X1, X2, X3]:~((~(X1=X2)&(~(X1=X3)&~(k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3))))), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t136_enumset1)).
fof(t129_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9)), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t129_enumset1)).
fof(d7_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9, X10]:(X10=k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)<=>![X11]:(r2_hidden(X11,X10)<=>~((~(X11=X1)&(~(X11=X2)&(~(X11=X3)&(~(X11=X4)&(~(X11=X5)&(~(X11=X6)&(~(X11=X7)&(~(X11=X8)&~(X11=X9)))))))))))), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', d7_enumset1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t4_xboole_0)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t66_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/9afa6f54ec234543afee0ff6efa24ddf.E_input', t12_xboole_1)).
fof(c_0_26, plain, ![X19, X20, X22, X23, X24]:(((r2_hidden(esk1_2(X19,X20),X19)|r1_xboole_0(X19,X20))&(r2_hidden(esk1_2(X19,X20),X20)|r1_xboole_0(X19,X20)))&(~r2_hidden(X24,X22)|~r2_hidden(X24,X23)|~r1_xboole_0(X22,X23))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_27, plain, ![X25, X26]:r1_xboole_0(k3_xboole_0(X25,X26),k5_xboole_0(X25,X26)), inference(variable_rename,[status(thm)],[t103_xboole_1])).
fof(c_0_28, plain, ![X27, X28]:k4_xboole_0(X27,k4_xboole_0(X27,X28))=k3_xboole_0(X27,X28), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_29, plain, ![X29]:k3_xboole_0(X29,X29)=X29, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_30, plain, ![X12, X13, X14]:(~r1_tarski(X12,k2_xboole_0(X13,X14))|r1_tarski(k4_xboole_0(X12,X13),X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
fof(c_0_31, plain, ![X15, X16]:k2_xboole_0(X15,X16)=k5_xboole_0(X15,k4_xboole_0(X16,X15)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_32, plain, ![X17, X18]:r1_tarski(X17,k2_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_33, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_38, plain, ![X30, X31, X32]:(~r1_tarski(X30,X31)|(~r1_tarski(X30,X32)|~r1_xboole_0(X31,X32))|X30=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t67_xboole_1])])).
fof(c_0_39, plain, ![X33]:r1_xboole_0(X33,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_40, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk1_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_44, plain, (r2_hidden(esk1_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_45, plain, (r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_46, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X1))=X1), inference(rw,[status(thm)],[c_0_37, c_0_36])).
cnf(c_0_47, plain, (X1=k1_xboole_0|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k5_xboole_0(X2,k4_xboole_0(X3,X2)))), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50, plain, (r1_tarski(X1,k5_xboole_0(X1,k4_xboole_0(X2,X1)))), inference(rw,[status(thm)],[c_0_42, c_0_41])).
cnf(c_0_51, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52, plain, (r1_xboole_0(X1,k5_xboole_0(X1,X1))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_54, plain, (r1_tarski(k4_xboole_0(X1,X1),X2)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_55, plain, (r1_xboole_0(k5_xboole_0(X1,X1),X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
fof(c_0_56, plain, ![X34]:k5_xboole_0(X34,k1_xboole_0)=X34, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_57, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_54])])).
cnf(c_0_58, plain, (X1=k1_xboole_0|~r1_tarski(X1,k5_xboole_0(X2,X2))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_47, c_0_55])).
cnf(c_0_59, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_60, plain, (r1_tarski(k1_xboole_0,X1)), inference(rw,[status(thm)],[c_0_54, c_0_57])).
cnf(c_0_61, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_62, plain, (r1_tarski(k4_xboole_0(k1_xboole_0,X1),X2)), inference(spm,[status(thm)],[c_0_49, c_0_60])).
fof(c_0_63, negated_conjecture, ~(![X1, X2, X3]:(k1_tarski(X1)=k2_tarski(X2,X3)=>X2=X3)), inference(assume_negation,[status(cth)],[t9_zfmisc_1])).
cnf(c_0_64, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_62])).
fof(c_0_65, plain, ![X38, X39]:k2_tarski(X38,X39)=k2_xboole_0(k1_tarski(X38),k1_tarski(X39)), inference(variable_rename,[status(thm)],[t41_enumset1])).
fof(c_0_66, plain, ![X40]:k2_tarski(X40,X40)=k1_tarski(X40), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_67, plain, ![X41, X42]:k1_enumset1(X41,X41,X42)=k2_tarski(X41,X42), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_68, negated_conjecture, (k1_tarski(esk2_0)=k2_tarski(esk3_0,esk4_0)&esk3_0!=esk4_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_63])])])).
cnf(c_0_69, plain, (r1_tarski(k4_xboole_0(X1,X2),k1_xboole_0)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_64]), c_0_59])).
fof(c_0_70, plain, ![X43, X44, X45, X46, X47, X48, X49, X50]:k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50)=k2_xboole_0(k2_tarski(X43,X44),k4_enumset1(X45,X46,X47,X48,X49,X50)), inference(variable_rename,[status(thm)],[t63_enumset1])).
cnf(c_0_71, plain, (k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_72, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_73, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_74, negated_conjecture, (k1_tarski(esk2_0)=k2_tarski(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_75, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_61, c_0_69])).
fof(c_0_76, plain, ![X51, X52, X53]:k1_enumset1(X51,X52,X53)=k1_enumset1(X52,X53,X51), inference(variable_rename,[status(thm)],[t100_enumset1])).
fof(c_0_77, plain, ![X54, X55, X56]:k1_enumset1(X54,X55,X56)=k1_enumset1(X54,X56,X55), inference(variable_rename,[status(thm)],[t98_enumset1])).
fof(c_0_78, plain, ![X57, X58]:k6_enumset1(X57,X57,X57,X57,X57,X57,X57,X58)=k2_tarski(X57,X58), inference(variable_rename,[status(thm)],[t95_enumset1])).
cnf(c_0_79, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X2),k4_enumset1(X3,X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_80, plain, ![X59, X60, X61, X62, X63]:k3_enumset1(X59,X60,X61,X62,X63)=k2_xboole_0(k1_enumset1(X59,X60,X61),k2_tarski(X62,X63)), inference(variable_rename,[status(thm)],[l57_enumset1])).
fof(c_0_81, plain, ![X64, X65, X66, X67, X68]:k6_enumset1(X64,X64,X64,X64,X65,X66,X67,X68)=k3_enumset1(X64,X65,X66,X67,X68), inference(variable_rename,[status(thm)],[t86_enumset1])).
fof(c_0_82, plain, ![X69, X70, X71]:(X69=X70|(X69=X71|k4_xboole_0(k1_enumset1(X69,X70,X71),k1_tarski(X69))=k2_tarski(X70,X71))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_enumset1])])])).
cnf(c_0_83, plain, (k1_enumset1(X1,X1,X2)=k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k1_enumset1(X2,X2,X2),k1_enumset1(X1,X1,X1)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_72]), c_0_72]), c_0_73]), c_0_73]), c_0_73]), c_0_41])).
cnf(c_0_84, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk4_0)=k1_enumset1(esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_72]), c_0_73]), c_0_73])).
cnf(c_0_85, plain, (k4_xboole_0(X1,k5_xboole_0(X1,k4_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_50])).
cnf(c_0_86, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_87, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_88, plain, ![X72, X73, X74, X75, X76, X77, X78, X79, X80]:k7_enumset1(X72,X73,X74,X75,X76,X77,X78,X79,X80)=k2_xboole_0(k1_enumset1(X72,X73,X74),k4_enumset1(X75,X76,X77,X78,X79,X80)), inference(variable_rename,[status(thm)],[t129_enumset1])).
cnf(c_0_89, plain, (k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_90, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k1_enumset1(X1,X1,X2),k4_xboole_0(k4_enumset1(X3,X4,X5,X6,X7,X8),k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_73]), c_0_41])).
cnf(c_0_91, plain, (k3_enumset1(X1,X2,X3,X4,X5)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_tarski(X4,X5))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_92, plain, (k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_93, plain, (X1=X2|X1=X3|k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_94, negated_conjecture, (k5_xboole_0(k1_enumset1(esk3_0,esk3_0,esk4_0),k4_xboole_0(k1_enumset1(X1,X1,X1),k1_enumset1(esk3_0,esk3_0,esk4_0)))=k1_enumset1(esk2_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_95, plain, (k4_xboole_0(k1_enumset1(X1,X1,X1),k1_enumset1(X1,X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_85, c_0_83])).
cnf(c_0_96, plain, (k1_enumset1(X1,X2,X3)=k1_enumset1(X3,X2,X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
fof(c_0_97, plain, ![X81, X82, X83, X84, X85, X86, X87, X88, X89, X90, X91, X92, X93, X94, X95, X96, X97, X98, X99, X100, X101, X102]:(((~r2_hidden(X91,X90)|(X91=X81|(X91=X82|(X91=X83|(X91=X84|(X91=X85|(X91=X86|(X91=X87|(X91=X88|X91=X89))))))))|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&((X92!=X81|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&((X92!=X82|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&((X92!=X83|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&((X92!=X84|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&((X92!=X85|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&((X92!=X86|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&((X92!=X87|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&((X92!=X88|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89))&(X92!=X89|r2_hidden(X92,X90)|X90!=k7_enumset1(X81,X82,X83,X84,X85,X86,X87,X88,X89)))))))))))&(((esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X93|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))&((esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X94|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))&((esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X95|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))&((esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X96|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))&((esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X97|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))&((esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X98|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))&((esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X99|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))&((esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X100|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))&(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)!=X101|~r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101))))))))))&(r2_hidden(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102),X102)|(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X93|(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X94|(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X95|(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X96|(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X97|(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X98|(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X99|(esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X100|esk5_10(X93,X94,X95,X96,X97,X98,X99,X100,X101,X102)=X101))))))))|X102=k7_enumset1(X93,X94,X95,X96,X97,X98,X99,X100,X101)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_enumset1])])])])])])])).
cnf(c_0_98, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k1_enumset1(X1,X2,X3),k4_enumset1(X4,X5,X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_99, plain, ![X104, X105, X107, X108, X109]:((r1_xboole_0(X104,X105)|r2_hidden(esk6_2(X104,X105),k3_xboole_0(X104,X105)))&(~r2_hidden(X109,k3_xboole_0(X107,X108))|~r1_xboole_0(X107,X108))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_100, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k4_enumset1(X1,X1,X1,X1,X1,X2),k1_enumset1(X1,X1,X1)))=k1_enumset1(X1,X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_73]), c_0_90])).
cnf(c_0_101, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k4_enumset1(X1,X1,X2,X3,X4,X5),k1_enumset1(X1,X1,X1)))=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X4,X5),k1_enumset1(X1,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_73]), c_0_41]), c_0_92]), c_0_90])).
cnf(c_0_102, plain, (X1=X3|X1=X2|k4_xboole_0(k1_enumset1(X1,X2,X3),k1_enumset1(X1,X1,X1))=k1_enumset1(X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_93, c_0_72]), c_0_73]), c_0_73])).
cnf(c_0_103, negated_conjecture, (k1_enumset1(esk3_0,esk2_0,esk2_0)=k1_enumset1(esk3_0,esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_59]), c_0_96])).
cnf(c_0_104, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k7_enumset1(X4,X5,X6,X7,X2,X8,X9,X10,X11)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_105, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k4_enumset1(X4,X5,X6,X7,X8,X9),k1_enumset1(X1,X2,X3)))), inference(rw,[status(thm)],[c_0_98, c_0_41])).
cnf(c_0_106, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_107, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k1_enumset1(X1,X1,X2),k1_enumset1(X1,X1,X1)))=k1_enumset1(X1,X1,X2)), inference(rw,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_108, negated_conjecture, (k4_xboole_0(k1_enumset1(esk3_0,esk3_0,esk4_0),k1_enumset1(esk3_0,esk3_0,esk3_0))=k1_enumset1(esk3_0,esk3_0,esk4_0)|esk2_0=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_84])).
cnf(c_0_109, negated_conjecture, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k1_enumset1(esk3_0,esk3_0,esk4_0),k1_enumset1(X1,X1,X1)))=k1_enumset1(X1,X1,esk2_0)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_110, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k5_xboole_0(k1_enumset1(X4,X5,X6),k4_xboole_0(k4_enumset1(X7,X2,X8,X9,X10,X11),k1_enumset1(X4,X5,X6)))), inference(rw,[status(thm)],[c_0_104, c_0_105])).
fof(c_0_111, plain, ![X110]:((r1_xboole_0(X110,X110)|X110!=k1_xboole_0)&(X110=k1_xboole_0|~r1_xboole_0(X110,X110))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_112, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_106, c_0_36])).
cnf(c_0_113, negated_conjecture, (k5_xboole_0(k1_enumset1(esk3_0,esk3_0,esk3_0),k1_enumset1(esk3_0,esk3_0,esk4_0))=k1_enumset1(esk3_0,esk3_0,esk4_0)|esk2_0=esk3_0), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_114, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk4_0)=k1_enumset1(esk3_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_107, c_0_109])).
cnf(c_0_115, plain, (r2_hidden(X1,k5_xboole_0(k1_enumset1(X2,X3,X4),k4_xboole_0(k4_enumset1(X5,X1,X6,X7,X8,X9),k1_enumset1(X2,X3,X4))))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_110])])).
cnf(c_0_116, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_117, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_112, c_0_44])).
cnf(c_0_118, negated_conjecture, (k5_xboole_0(k1_enumset1(esk3_0,esk3_0,esk3_0),k1_enumset1(esk3_0,esk3_0,esk2_0))=k1_enumset1(esk3_0,esk3_0,esk2_0)|esk2_0=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_114]), c_0_114])).
cnf(c_0_119, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_46, c_0_57])).
cnf(c_0_120, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,k5_xboole_0(k1_enumset1(X3,X4,X5),k4_xboole_0(k4_enumset1(X6,X1,X7,X8,X9,X10),k1_enumset1(X3,X4,X5))))), inference(spm,[status(thm)],[c_0_33, c_0_115])).
cnf(c_0_121, plain, (k5_xboole_0(k1_enumset1(X1,X1,X1),k4_xboole_0(k4_enumset1(X1,X1,X1,X1,X2,X2),k1_enumset1(X1,X1,X1)))=k1_enumset1(X1,X1,X2)), inference(spm,[status(thm)],[c_0_83, c_0_101])).
cnf(c_0_122, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_123, negated_conjecture, (esk2_0=esk3_0|r1_xboole_0(k1_enumset1(esk3_0,esk3_0,esk3_0),k1_enumset1(esk3_0,esk3_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_118]), c_0_95]), c_0_119])).
cnf(c_0_124, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k7_enumset1(X2,X4,X5,X6,X7,X8,X9,X10,X11)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_125, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,k1_enumset1(X1,X1,X3))), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_126, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk3_0)=k1_xboole_0|esk2_0=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_95]), c_0_119])).
cnf(c_0_127, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k5_xboole_0(k1_enumset1(X2,X4,X5),k4_xboole_0(k4_enumset1(X6,X7,X8,X9,X10,X11),k1_enumset1(X2,X4,X5)))), inference(rw,[status(thm)],[c_0_124, c_0_105])).
cnf(c_0_128, plain, (r1_tarski(k1_enumset1(X1,X1,X1),k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k1_enumset1(X4,X4,X5),k1_enumset1(X1,X2,X3))))), inference(spm,[status(thm)],[c_0_50, c_0_101])).
cnf(c_0_129, plain, (k4_xboole_0(k1_enumset1(X1,X1,X1),k1_enumset1(X2,X1,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_130, negated_conjecture, (esk2_0=esk3_0|~r2_hidden(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_48])])).
cnf(c_0_131, plain, (r2_hidden(X1,k5_xboole_0(k1_enumset1(X1,X2,X3),k4_xboole_0(k4_enumset1(X4,X5,X6,X7,X8,X9),k1_enumset1(X1,X2,X3))))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_127])])).
fof(c_0_132, plain, ![X111, X112]:(~r1_tarski(X111,X112)|k2_xboole_0(X111,X112)=X112), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_133, plain, (r1_tarski(k1_enumset1(X1,X1,X1),k1_enumset1(X1,X2,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_59])).
cnf(c_0_134, negated_conjecture, (esk2_0=esk3_0), inference(spm,[status(thm)],[c_0_130, c_0_131])).
cnf(c_0_135, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k7_enumset1(X4,X5,X2,X6,X7,X8,X9,X10,X11)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_136, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_137, plain, (r1_tarski(k1_enumset1(X1,X1,X1),k1_enumset1(X2,X2,X1))), inference(spm,[status(thm)],[c_0_133, c_0_86])).
cnf(c_0_138, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk4_0)=k1_enumset1(esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[c_0_114, c_0_134])).
cnf(c_0_139, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k5_xboole_0(k1_enumset1(X4,X5,X2),k4_xboole_0(k4_enumset1(X6,X7,X8,X9,X10,X11),k1_enumset1(X4,X5,X2)))), inference(rw,[status(thm)],[c_0_135, c_0_105])).
cnf(c_0_140, plain, (k5_xboole_0(X1,k4_xboole_0(X2,X1))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_136, c_0_41])).
cnf(c_0_141, negated_conjecture, (r1_tarski(k1_enumset1(esk4_0,esk4_0,esk4_0),k1_enumset1(esk3_0,esk3_0,esk3_0))), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_142, plain, (r2_hidden(X1,k5_xboole_0(k1_enumset1(X2,X3,X1),k4_xboole_0(k4_enumset1(X4,X5,X6,X7,X8,X9),k1_enumset1(X2,X3,X1))))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_139])])).
cnf(c_0_143, negated_conjecture, (k1_enumset1(esk3_0,esk4_0,esk4_0)=k1_enumset1(esk3_0,esk3_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_83]), c_0_96])).
cnf(c_0_144, negated_conjecture, (esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_145, plain, (~r2_hidden(X1,k4_xboole_0(X2,X2))|~r1_xboole_0(X2,k4_xboole_0(X2,X2))), inference(spm,[status(thm)],[c_0_112, c_0_46])).
cnf(c_0_146, plain, (r2_hidden(X1,k1_enumset1(X1,X1,X2))), inference(spm,[status(thm)],[c_0_142, c_0_121])).
cnf(c_0_147, negated_conjecture, (k1_enumset1(esk4_0,esk4_0,esk4_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_143]), c_0_95]), c_0_144])).
cnf(c_0_148, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_145, c_0_57]), c_0_57]), c_0_48])])).
cnf(c_0_149, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_148]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 150
# Proof object clause steps            : 97
# Proof object formula steps           : 53
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 41
# Proof object simplifying inferences  : 64
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 43
# Processed clauses                    : 1929
# ...of these trivial                  : 152
# ...subsumed                          : 1415
# ...remaining for further processing  : 362
# Other redundant clauses eliminated   : 29
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 94
# Generated clauses                    : 10338
# ...of the previous two non-trivial   : 7787
# Contextual simplify-reflections      : 0
# Paramodulations                      : 10318
# Factorizations                       : 0
# Equation resolutions                 : 29
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 212
#    Positive orientable unit clauses  : 71
#    Positive unorientable unit clauses: 12
#    Negative unit clauses             : 20
#    Non-unit-clauses                  : 109
# Current number of unprocessed clauses: 5110
# ...number of literals in the above   : 19008
# Current number of archived formulas  : 0
# Current number of archived clauses   : 146
# Clause-clause subsumption calls (NU) : 6830
# Rec. Clause-clause subsumption calls : 6448
# Non-unit clause-clause subsumptions  : 1086
# Unit Clause-clause subsumption calls : 525
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 686
# BW rewrite match successes           : 129
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 141098

# -------------------------------------------------
# User time                : 0.118 s
# System time              : 0.000 s
# Total time               : 0.118 s
# Maximum resident set size: 28740 pages
