# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t71_yellow_0, conjecture, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v6_yellow_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>((X5=X3&X6=X4)=>k13_lattice3(X2,X3,X4)=k13_lattice3(X1,X5,X6)))))))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', t71_yellow_0)).
fof(redefinition_k7_domain_1, axiom, ![X1, X2, X3]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&m1_subset_1(X3,X1)))=>k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', redefinition_k7_domain_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', t70_enumset1)).
fof(t20_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>(v1_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r1_yellow_0(X1,k2_tarski(X2,X3)))))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', t20_yellow_0)).
fof(dt_m1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_yellow_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', dt_m1_yellow_0)).
fof(cc12_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_yellow_0(X2,X1)=>((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&v6_yellow_0(X2,X1)))=>(~(v2_struct_0(X2))&(v1_lattice3(X2)&(v4_yellow_0(X2,X1)&v6_yellow_0(X2,X1))))))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', cc12_yellow_0)).
fof(cc8_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v5_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', cc8_yellow_0)).
fof(cc7_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v4_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', cc7_yellow_0)).
fof(cc6_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v4_yellow_0(X2,X1)=>(v3_orders_2(X2)&v4_yellow_0(X2,X1))))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', cc6_yellow_0)).
fof(cc1_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v1_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', cc1_lattice3)).
fof(t41_yellow_0, axiom, ![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k13_lattice3(X1,X2,X3)))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', t41_yellow_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', fc2_struct_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', t2_subset)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', dt_l1_orders_2)).
fof(d17_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_yellow_0(X2,X1)=>(v6_yellow_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,u1_struct_0(X2))&(r2_hidden(X4,u1_struct_0(X2))&r1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))))=>r2_hidden(k1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X2)))))))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', d17_yellow_0)).
fof(t67_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&m1_yellow_0(X2,X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>((r1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4))&r2_hidden(k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X2)))=>(r1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))&k1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))=k1_yellow_0(X1,k7_domain_1(u1_struct_0(X2),X3,X4)))))))), file('data/MPTP/proofs/9b72011437d34c0bb0bd2c0b48eaf322.E_input', t67_yellow_0)).
fof(c_0_16, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))))=>![X2]:((~(v2_struct_0(X2))&(v4_yellow_0(X2,X1)&(v6_yellow_0(X2,X1)&m1_yellow_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>((X5=X3&X6=X4)=>k13_lattice3(X2,X3,X4)=k13_lattice3(X1,X5,X6))))))))), inference(assume_negation,[status(cth)],[t71_yellow_0])).
fof(c_0_17, plain, ![X39, X40, X41]:(v1_xboole_0(X39)|(~m1_subset_1(X40,X39)|~m1_subset_1(X41,X39))|k7_domain_1(X39,X40,X41)=k2_tarski(X40,X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k7_domain_1])])])).
fof(c_0_18, plain, ![X27, X28]:k1_enumset1(X27,X27,X28)=k2_tarski(X27,X28), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_19, negated_conjecture, ((v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_lattice3(esk1_0)&l1_orders_2(esk1_0)))))&((~v2_struct_0(esk2_0)&(v4_yellow_0(esk2_0,esk1_0)&(v6_yellow_0(esk2_0,esk1_0)&m1_yellow_0(esk2_0,esk1_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&(m1_subset_1(esk4_0,u1_struct_0(esk2_0))&(m1_subset_1(esk5_0,u1_struct_0(esk1_0))&(m1_subset_1(esk6_0,u1_struct_0(esk1_0))&((esk5_0=esk3_0&esk6_0=esk4_0)&k13_lattice3(esk2_0,esk3_0,esk4_0)!=k13_lattice3(esk1_0,esk5_0,esk6_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])).
fof(c_0_20, plain, ![X19, X20, X21]:((~v1_lattice3(X19)|(~m1_subset_1(X20,u1_struct_0(X19))|(~m1_subset_1(X21,u1_struct_0(X19))|r1_yellow_0(X19,k2_tarski(X20,X21))))|(~v5_orders_2(X19)|~l1_orders_2(X19)))&((m1_subset_1(esk9_1(X19),u1_struct_0(X19))|v1_lattice3(X19)|(~v5_orders_2(X19)|~l1_orders_2(X19)))&((m1_subset_1(esk10_1(X19),u1_struct_0(X19))|v1_lattice3(X19)|(~v5_orders_2(X19)|~l1_orders_2(X19)))&(~r1_yellow_0(X19,k2_tarski(esk9_1(X19),esk10_1(X19)))|v1_lattice3(X19)|(~v5_orders_2(X19)|~l1_orders_2(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_yellow_0])])])])])).
cnf(c_0_21, plain, (v1_xboole_0(X1)|k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (esk6_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (r1_yellow_0(X1,k2_tarski(X2,X3))|~v1_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (k1_enumset1(X2,X2,X3)=k7_domain_1(X1,X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk2_0))), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_30, negated_conjecture, (esk5_0=esk3_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (r1_yellow_0(X1,k1_enumset1(X2,X2,X3))|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_25, c_0_22])).
cnf(c_0_32, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, negated_conjecture, (v1_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_35, plain, ![X49, X50]:(~l1_orders_2(X49)|(~m1_yellow_0(X50,X49)|l1_orders_2(X50))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_yellow_0])])])).
fof(c_0_36, plain, ![X25, X26]:((~v2_struct_0(X26)|(v2_struct_0(X26)|(~v4_yellow_0(X26,X25)|~v6_yellow_0(X26,X25)))|~m1_yellow_0(X26,X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))&((v1_lattice3(X26)|(v2_struct_0(X26)|(~v4_yellow_0(X26,X25)|~v6_yellow_0(X26,X25)))|~m1_yellow_0(X26,X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))&((v4_yellow_0(X26,X25)|(v2_struct_0(X26)|(~v4_yellow_0(X26,X25)|~v6_yellow_0(X26,X25)))|~m1_yellow_0(X26,X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))&(v6_yellow_0(X26,X25)|(v2_struct_0(X26)|(~v4_yellow_0(X26,X25)|~v6_yellow_0(X26,X25)))|~m1_yellow_0(X26,X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|(~v1_lattice3(X25)|~l1_orders_2(X25)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc12_yellow_0])])])])])).
fof(c_0_37, plain, ![X34, X35]:((v5_orders_2(X35)|~v4_yellow_0(X35,X34)|~m1_yellow_0(X35,X34)|(~v5_orders_2(X34)|~l1_orders_2(X34)))&(v4_yellow_0(X35,X34)|~v4_yellow_0(X35,X34)|~m1_yellow_0(X35,X34)|(~v5_orders_2(X34)|~l1_orders_2(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_yellow_0])])])])).
fof(c_0_38, plain, ![X47, X48]:((v4_orders_2(X48)|~v4_yellow_0(X48,X47)|~m1_yellow_0(X48,X47)|(~v4_orders_2(X47)|~l1_orders_2(X47)))&(v4_yellow_0(X48,X47)|~v4_yellow_0(X48,X47)|~m1_yellow_0(X48,X47)|(~v4_orders_2(X47)|~l1_orders_2(X47)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_yellow_0])])])])).
fof(c_0_39, plain, ![X45, X46]:((v3_orders_2(X46)|~v4_yellow_0(X46,X45)|~m1_yellow_0(X46,X45)|(~v3_orders_2(X45)|~l1_orders_2(X45)))&(v4_yellow_0(X46,X45)|~v4_yellow_0(X46,X45)|~m1_yellow_0(X46,X45)|(~v3_orders_2(X45)|~l1_orders_2(X45)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_yellow_0])])])])).
cnf(c_0_40, negated_conjecture, (k1_enumset1(X1,X1,esk6_0)=k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0)|v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_42, negated_conjecture, (k1_enumset1(X1,X1,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_26, c_0_28])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_44, negated_conjecture, (r1_yellow_0(esk1_0,k1_enumset1(X1,X1,esk6_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_28]), c_0_32]), c_0_33]), c_0_34])])).
fof(c_0_45, plain, ![X24]:(~l1_orders_2(X24)|(~v1_lattice3(X24)|~v2_struct_0(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattice3])])])).
fof(c_0_46, plain, ![X42, X43, X44]:(~v3_orders_2(X42)|(~v4_orders_2(X42)|(~v5_orders_2(X42)|(~v1_lattice3(X42)|~l1_orders_2(X42))))|(~m1_subset_1(X43,u1_struct_0(X42))|(~m1_subset_1(X44,u1_struct_0(X42))|k1_yellow_0(X42,k7_domain_1(u1_struct_0(X42),X43,X44))=k13_lattice3(X42,X43,X44)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t41_yellow_0])])])).
cnf(c_0_47, plain, (l1_orders_2(X2)|~l1_orders_2(X1)|~m1_yellow_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, negated_conjecture, (m1_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_49, plain, (v1_lattice3(X1)|v2_struct_0(X1)|~v4_yellow_0(X1,X2)|~v6_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v1_lattice3(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50, negated_conjecture, (v4_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_51, negated_conjecture, (v6_yellow_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_52, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_53, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_54, plain, (v5_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_55, plain, (v4_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_56, plain, (v3_orders_2(X1)|~v4_yellow_0(X1,X2)|~m1_yellow_0(X1,X2)|~v3_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_57, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_58, plain, ![X36]:(v2_struct_0(X36)|~l1_struct_0(X36)|~v1_xboole_0(u1_struct_0(X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_59, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_60, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_61, negated_conjecture, (r1_yellow_0(esk1_0,k1_enumset1(esk3_0,esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_44, c_0_43])).
cnf(c_0_62, plain, (~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_63, plain, ![X37, X38]:(~m1_subset_1(X37,X38)|(v1_xboole_0(X38)|r2_hidden(X37,X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_64, plain, (k1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k13_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_65, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_32])])).
cnf(c_0_66, negated_conjecture, (v1_lattice3(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_48]), c_0_51]), c_0_32]), c_0_33]), c_0_34]), c_0_52])]), c_0_53])).
cnf(c_0_67, negated_conjecture, (v5_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_50]), c_0_48]), c_0_32]), c_0_34])])).
cnf(c_0_68, negated_conjecture, (v4_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_50]), c_0_48]), c_0_32]), c_0_52])])).
cnf(c_0_69, negated_conjecture, (v3_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_50]), c_0_48]), c_0_32]), c_0_57])])).
cnf(c_0_70, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_71, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
fof(c_0_72, plain, ![X33]:(~l1_orders_2(X33)|l1_struct_0(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_73, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_61, c_0_60])).
cnf(c_0_74, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_33]), c_0_32])])).
cnf(c_0_75, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_76, negated_conjecture, (k13_lattice3(esk2_0,esk3_0,esk4_0)!=k13_lattice3(esk1_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_77, negated_conjecture, (k13_lattice3(esk2_0,X1,esk6_0)=k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),X1,esk6_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_27]), c_0_65]), c_0_66]), c_0_67]), c_0_68]), c_0_69])])).
cnf(c_0_78, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_53])).
cnf(c_0_79, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_80, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_61, c_0_59])).
fof(c_0_81, plain, ![X13, X14, X15, X16]:((~v6_yellow_0(X14,X13)|(~m1_subset_1(X15,u1_struct_0(X13))|(~m1_subset_1(X16,u1_struct_0(X13))|(~r2_hidden(X15,u1_struct_0(X14))|(~r2_hidden(X16,u1_struct_0(X14))|~r1_yellow_0(X13,k7_domain_1(u1_struct_0(X13),X15,X16)))|r2_hidden(k1_yellow_0(X13,k7_domain_1(u1_struct_0(X13),X15,X16)),u1_struct_0(X14)))))|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&((m1_subset_1(esk7_2(X13,X14),u1_struct_0(X13))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&((m1_subset_1(esk8_2(X13,X14),u1_struct_0(X13))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&(((r2_hidden(esk7_2(X13,X14),u1_struct_0(X14))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&((r2_hidden(esk8_2(X13,X14),u1_struct_0(X14))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))&(r1_yellow_0(X13,k7_domain_1(u1_struct_0(X13),esk7_2(X13,X14),esk8_2(X13,X14)))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13)))))&(~r2_hidden(k1_yellow_0(X13,k7_domain_1(u1_struct_0(X13),esk7_2(X13,X14),esk8_2(X13,X14))),u1_struct_0(X14))|v6_yellow_0(X14,X13)|~m1_yellow_0(X14,X13)|(v2_struct_0(X13)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_yellow_0])])])])])])).
cnf(c_0_82, negated_conjecture, (r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_73]), c_0_74])).
cnf(c_0_83, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk6_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_75, c_0_27])).
cnf(c_0_84, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk3_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_75, c_0_41])).
cnf(c_0_85, negated_conjecture, (k13_lattice3(esk2_0,esk3_0,esk6_0)!=k13_lattice3(esk1_0,esk3_0,esk6_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_24]), c_0_30])).
cnf(c_0_86, negated_conjecture, (k13_lattice3(esk2_0,esk3_0,esk6_0)=k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_77, c_0_41])).
cnf(c_0_87, negated_conjecture, (k13_lattice3(esk1_0,X1,esk6_0)=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,esk6_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_28]), c_0_32]), c_0_33]), c_0_34]), c_0_52]), c_0_57])])).
cnf(c_0_88, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_65])])).
fof(c_0_89, plain, ![X29, X30, X31, X32]:((r1_yellow_0(X30,k7_domain_1(u1_struct_0(X30),X31,X32))|(~r1_yellow_0(X29,k7_domain_1(u1_struct_0(X30),X31,X32))|~r2_hidden(k1_yellow_0(X29,k7_domain_1(u1_struct_0(X30),X31,X32)),u1_struct_0(X30)))|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|(~v4_yellow_0(X30,X29)|~m1_yellow_0(X30,X29)))|(v2_struct_0(X29)|(~v4_orders_2(X29)|~l1_orders_2(X29))))&(k1_yellow_0(X30,k7_domain_1(u1_struct_0(X30),X31,X32))=k1_yellow_0(X29,k7_domain_1(u1_struct_0(X30),X31,X32))|(~r1_yellow_0(X29,k7_domain_1(u1_struct_0(X30),X31,X32))|~r2_hidden(k1_yellow_0(X29,k7_domain_1(u1_struct_0(X30),X31,X32)),u1_struct_0(X30)))|~m1_subset_1(X32,u1_struct_0(X30))|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|(~v4_yellow_0(X30,X29)|~m1_yellow_0(X30,X29)))|(v2_struct_0(X29)|(~v4_orders_2(X29)|~l1_orders_2(X29))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t67_yellow_0])])])])])).
cnf(c_0_90, negated_conjecture, (r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_80]), c_0_53])).
cnf(c_0_91, plain, (r2_hidden(k1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X1))|v2_struct_0(X2)|~v6_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,u1_struct_0(X1))|~r2_hidden(X4,u1_struct_0(X1))|~r1_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~m1_yellow_0(X1,X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_92, negated_conjecture, (r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_79]), c_0_32])])).
cnf(c_0_93, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk2_0))|~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_83]), c_0_53])).
cnf(c_0_94, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk2_0))|~l1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_84]), c_0_53])).
cnf(c_0_95, negated_conjecture, (k13_lattice3(esk1_0,esk3_0,esk6_0)!=k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))), inference(rw,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_96, negated_conjecture, (k13_lattice3(esk1_0,esk3_0,esk6_0)=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(spm,[status(thm)],[c_0_87, c_0_43])).
cnf(c_0_97, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_88]), c_0_74])).
cnf(c_0_98, plain, (k1_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k1_yellow_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3))|v2_struct_0(X1)|v2_struct_0(X4)|~r1_yellow_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_hidden(k1_yellow_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_yellow_0(X1,X4)|~m1_yellow_0(X1,X4)|~v4_orders_2(X4)|~l1_orders_2(X4)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_99, negated_conjecture, (r1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_79]), c_0_65])])).
cnf(c_0_100, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)),u1_struct_0(X1))|~r2_hidden(esk6_0,u1_struct_0(X1))|~r2_hidden(esk3_0,u1_struct_0(X1))|~m1_yellow_0(X1,esk1_0)|~v6_yellow_0(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_28]), c_0_43]), c_0_32])]), c_0_74])).
cnf(c_0_101, negated_conjecture, (r2_hidden(esk6_0,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_79]), c_0_65])])).
cnf(c_0_102, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_79]), c_0_65])])).
cnf(c_0_103, negated_conjecture, (k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))!=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(rw,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_104, negated_conjecture, (k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)=k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_79]), c_0_32])])).
cnf(c_0_105, negated_conjecture, (k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0))|~r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk2_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_27]), c_0_41]), c_0_48]), c_0_50]), c_0_32]), c_0_52])]), c_0_74]), c_0_53])).
cnf(c_0_106, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0)),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102]), c_0_48]), c_0_51])])).
cnf(c_0_107, negated_conjecture, (k1_yellow_0(esk2_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))!=k1_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),esk3_0,esk6_0))), inference(rw,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_108, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_104]), c_0_104]), c_0_104]), c_0_106])]), c_0_107]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 109
# Proof object clause steps            : 76
# Proof object formula steps           : 33
# Proof object conjectures             : 62
# Proof object clause conjectures      : 59
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 91
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 47
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 40
# Processed clauses                    : 248
# ...of these trivial                  : 3
# ...subsumed                          : 2
# ...remaining for further processing  : 242
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 76
# Generated clauses                    : 228
# ...of the previous two non-trivial   : 187
# Contextual simplify-reflections      : 0
# Paramodulations                      : 228
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 120
#    Positive orientable unit clauses  : 55
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 62
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 90
# Current number of archived formulas  : 0
# Current number of archived clauses   : 123
# Clause-clause subsumption calls (NU) : 1942
# Rec. Clause-clause subsumption calls : 332
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 180
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 71
# BW rewrite match successes           : 26
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10493

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 28712 pages
