# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t70_enumset1)).
fof(d5_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:((~(v2_struct_0(X4))&(v1_pre_topc(X4)&m1_pre_topc(X4,X1)))=>(X4=k2_tsep_1(X1,X2,X3)<=>u1_struct_0(X4)=k3_xboole_0(u1_struct_0(X2),u1_struct_0(X3)))))))), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', d5_tsep_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t73_enumset1)).
fof(dt_k2_tsep_1, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))&(~(v2_struct_0(X3))&m1_pre_topc(X3,X1))))=>(~(v2_struct_0(k2_tsep_1(X1,X2,X3)))&(v1_pre_topc(k2_tsep_1(X1,X2,X3))&m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)))), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', dt_k2_tsep_1)).
fof(t17_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:(m1_subset_1(X4,u1_struct_0(k2_tsep_1(X1,X2,X3)))=>(?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&X5=X4)&?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&X5=X4))))))), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t17_tmap_1)).
fof(t4_setfam_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(k1_setfam_1(X2),X1)), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t4_setfam_1)).
fof(d4_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:(X7=k4_enumset1(X1,X2,X3,X4,X5,X6)<=>![X8]:(r2_hidden(X8,X7)<=>~((~(X8=X1)&(~(X8=X2)&(~(X8=X3)&(~(X8=X4)&(~(X8=X5)&~(X8=X6))))))))), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', d4_enumset1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t17_xboole_1)).
fof(t4_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X1)=>(r1_tarski(u1_struct_0(X2),u1_struct_0(X3))<=>m1_pre_topc(X2,X3))))), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t4_tsep_1)).
fof(t55_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1))))), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', t55_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/9b84404b53814239b8579960ea415039.E_input', dt_m1_pre_topc)).
fof(c_0_14, plain, ![X9, X10]:k1_setfam_1(k2_tarski(X9,X10))=k3_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_15, plain, ![X11, X12]:k1_enumset1(X11,X11,X12)=k2_tarski(X11,X12), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X13, X14, X15, X16]:((X16!=k2_tsep_1(X13,X14,X15)|u1_struct_0(X16)=k3_xboole_0(u1_struct_0(X14),u1_struct_0(X15))|(v2_struct_0(X16)|(~v1_pre_topc(X16)|~m1_pre_topc(X16,X13)))|r1_tsep_1(X14,X15)|(v2_struct_0(X15)|~m1_pre_topc(X15,X13))|(v2_struct_0(X14)|~m1_pre_topc(X14,X13))|(v2_struct_0(X13)|~l1_pre_topc(X13)))&(u1_struct_0(X16)!=k3_xboole_0(u1_struct_0(X14),u1_struct_0(X15))|X16=k2_tsep_1(X13,X14,X15)|(v2_struct_0(X16)|(~v1_pre_topc(X16)|~m1_pre_topc(X16,X13)))|r1_tsep_1(X14,X15)|(v2_struct_0(X15)|~m1_pre_topc(X15,X13))|(v2_struct_0(X14)|~m1_pre_topc(X14,X13))|(v2_struct_0(X13)|~l1_pre_topc(X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_tsep_1])])])])])).
cnf(c_0_17, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X17, X18, X19]:k2_enumset1(X17,X17,X18,X19)=k1_enumset1(X17,X18,X19), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_20, plain, ![X20, X21, X22, X23]:k3_enumset1(X20,X20,X21,X22,X23)=k2_enumset1(X20,X21,X22,X23), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_21, plain, ![X24, X25, X26, X27, X28]:k4_enumset1(X24,X24,X25,X26,X27,X28)=k3_enumset1(X24,X25,X26,X27,X28), inference(variable_rename,[status(thm)],[t73_enumset1])).
cnf(c_0_22, plain, (u1_struct_0(X1)=k3_xboole_0(u1_struct_0(X3),u1_struct_0(X4))|v2_struct_0(X1)|r1_tsep_1(X3,X4)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|X1!=k2_tsep_1(X2,X3,X4)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_27, plain, ![X29, X30, X31]:((~v2_struct_0(k2_tsep_1(X29,X30,X31))|(v2_struct_0(X29)|~l1_pre_topc(X29)|(v2_struct_0(X30)|~m1_pre_topc(X30,X29)|(v2_struct_0(X31)|~m1_pre_topc(X31,X29)))))&((v1_pre_topc(k2_tsep_1(X29,X30,X31))|(v2_struct_0(X29)|~l1_pre_topc(X29)|(v2_struct_0(X30)|~m1_pre_topc(X30,X29)|(v2_struct_0(X31)|~m1_pre_topc(X31,X29)))))&(m1_pre_topc(k2_tsep_1(X29,X30,X31),X29)|(v2_struct_0(X29)|~l1_pre_topc(X29)|(v2_struct_0(X30)|~m1_pre_topc(X30,X29)|(v2_struct_0(X31)|~m1_pre_topc(X31,X29))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_tsep_1])])])])).
cnf(c_0_28, plain, (u1_struct_0(X1)=k1_setfam_1(k4_enumset1(u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X4)))|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|r1_tsep_1(X3,X4)|X1!=k2_tsep_1(X2,X3,X4)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X4,X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_25]), c_0_26])).
cnf(c_0_29, plain, (m1_pre_topc(k2_tsep_1(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30, plain, (v1_pre_topc(k2_tsep_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_31, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(k2_tsep_1(X1,X2,X3))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_32, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>(~(r1_tsep_1(X2,X3))=>![X4]:(m1_subset_1(X4,u1_struct_0(k2_tsep_1(X1,X2,X3)))=>(?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&X5=X4)&?[X5]:(m1_subset_1(X5,u1_struct_0(X3))&X5=X4)))))))), inference(assume_negation,[status(cth)],[t17_tmap_1])).
fof(c_0_33, plain, ![X38, X39]:(~r2_hidden(X38,X39)|r1_tarski(k1_setfam_1(X39),X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_setfam_1])])).
fof(c_0_34, plain, ![X40, X41, X42, X43, X44, X45, X46, X47, X48, X49, X50, X51, X52, X53, X54, X55]:(((~r2_hidden(X47,X46)|(X47=X40|(X47=X41|(X47=X42|(X47=X43|(X47=X44|X47=X45)))))|X46!=k4_enumset1(X40,X41,X42,X43,X44,X45))&((X48!=X40|r2_hidden(X48,X46)|X46!=k4_enumset1(X40,X41,X42,X43,X44,X45))&((X48!=X41|r2_hidden(X48,X46)|X46!=k4_enumset1(X40,X41,X42,X43,X44,X45))&((X48!=X42|r2_hidden(X48,X46)|X46!=k4_enumset1(X40,X41,X42,X43,X44,X45))&((X48!=X43|r2_hidden(X48,X46)|X46!=k4_enumset1(X40,X41,X42,X43,X44,X45))&((X48!=X44|r2_hidden(X48,X46)|X46!=k4_enumset1(X40,X41,X42,X43,X44,X45))&(X48!=X45|r2_hidden(X48,X46)|X46!=k4_enumset1(X40,X41,X42,X43,X44,X45))))))))&(((esk5_7(X49,X50,X51,X52,X53,X54,X55)!=X49|~r2_hidden(esk5_7(X49,X50,X51,X52,X53,X54,X55),X55)|X55=k4_enumset1(X49,X50,X51,X52,X53,X54))&((esk5_7(X49,X50,X51,X52,X53,X54,X55)!=X50|~r2_hidden(esk5_7(X49,X50,X51,X52,X53,X54,X55),X55)|X55=k4_enumset1(X49,X50,X51,X52,X53,X54))&((esk5_7(X49,X50,X51,X52,X53,X54,X55)!=X51|~r2_hidden(esk5_7(X49,X50,X51,X52,X53,X54,X55),X55)|X55=k4_enumset1(X49,X50,X51,X52,X53,X54))&((esk5_7(X49,X50,X51,X52,X53,X54,X55)!=X52|~r2_hidden(esk5_7(X49,X50,X51,X52,X53,X54,X55),X55)|X55=k4_enumset1(X49,X50,X51,X52,X53,X54))&((esk5_7(X49,X50,X51,X52,X53,X54,X55)!=X53|~r2_hidden(esk5_7(X49,X50,X51,X52,X53,X54,X55),X55)|X55=k4_enumset1(X49,X50,X51,X52,X53,X54))&(esk5_7(X49,X50,X51,X52,X53,X54,X55)!=X54|~r2_hidden(esk5_7(X49,X50,X51,X52,X53,X54,X55),X55)|X55=k4_enumset1(X49,X50,X51,X52,X53,X54)))))))&(r2_hidden(esk5_7(X49,X50,X51,X52,X53,X54,X55),X55)|(esk5_7(X49,X50,X51,X52,X53,X54,X55)=X49|(esk5_7(X49,X50,X51,X52,X53,X54,X55)=X50|(esk5_7(X49,X50,X51,X52,X53,X54,X55)=X51|(esk5_7(X49,X50,X51,X52,X53,X54,X55)=X52|(esk5_7(X49,X50,X51,X52,X53,X54,X55)=X53|esk5_7(X49,X50,X51,X52,X53,X54,X55)=X54)))))|X55=k4_enumset1(X49,X50,X51,X52,X53,X54)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_enumset1])])])])])])])).
cnf(c_0_35, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=k1_setfam_1(k4_enumset1(u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3)))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_28]), c_0_29]), c_0_30]), c_0_31])).
fof(c_0_36, negated_conjecture, ![X36, X37]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&(~r1_tsep_1(esk2_0,esk3_0)&(m1_subset_1(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))&(~m1_subset_1(X36,u1_struct_0(esk2_0))|X36!=esk4_0|(~m1_subset_1(X37,u1_struct_0(esk3_0))|X37!=esk4_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_32])])])])])).
fof(c_0_37, plain, ![X57, X58]:r1_tarski(k3_xboole_0(X57,X58),X57), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_38, plain, (r1_tarski(k1_setfam_1(X2),X1)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k4_enumset1(X4,X5,X6,X7,X8,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (u1_struct_0(k2_tsep_1(X1,X2,X3))=u1_struct_0(k2_tsep_1(X4,X2,X3))|r1_tsep_1(X2,X3)|v2_struct_0(X4)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X4)|~m1_pre_topc(X2,X4)|~l1_pre_topc(X1)|~l1_pre_topc(X4)), inference(spm,[status(thm)],[c_0_35, c_0_35])).
cnf(c_0_41, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_46, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X1,X2,X3)),X4)|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(X4,k4_enumset1(u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3)))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_38, c_0_35])).
cnf(c_0_47, plain, (r2_hidden(X1,k4_enumset1(X2,X3,X4,X5,X6,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_39])])).
cnf(c_0_48, negated_conjecture, (u1_struct_0(k2_tsep_1(esk1_0,X1,esk3_0))=u1_struct_0(k2_tsep_1(X2,X1,esk3_0))|r1_tsep_1(X1,esk3_0)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_pre_topc(X1,esk1_0)|~m1_pre_topc(esk3_0,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])]), c_0_43]), c_0_44])).
cnf(c_0_49, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50, negated_conjecture, (~r1_tsep_1(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_51, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_52, plain, ![X59, X60, X61]:((~r1_tarski(u1_struct_0(X60),u1_struct_0(X61))|m1_pre_topc(X60,X61)|~m1_pre_topc(X61,X59)|~m1_pre_topc(X60,X59)|(~v2_pre_topc(X59)|~l1_pre_topc(X59)))&(~m1_pre_topc(X60,X61)|r1_tarski(u1_struct_0(X60),u1_struct_0(X61))|~m1_pre_topc(X61,X59)|~m1_pre_topc(X60,X59)|(~v2_pre_topc(X59)|~l1_pre_topc(X59)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_tsep_1])])])])).
cnf(c_0_53, plain, (r1_tarski(k1_setfam_1(k4_enumset1(X1,X1,X1,X1,X1,X2)),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_23]), c_0_24]), c_0_25]), c_0_26])).
cnf(c_0_54, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X1,X2,X3)),u1_struct_0(X3))|r1_tsep_1(X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55, negated_conjecture, (u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))=u1_struct_0(k2_tsep_1(X1,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_56, plain, (m1_pre_topc(X1,X2)|~r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_57, plain, (r1_tarski(u1_struct_0(k2_tsep_1(X1,X2,X3)),u1_struct_0(X2))|r1_tsep_1(X2,X3)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_53, c_0_35])).
cnf(c_0_58, negated_conjecture, (r1_tarski(u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk3_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_50]), c_0_51]), c_0_43])).
cnf(c_0_59, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),X1)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X2)|~v2_pre_topc(X4)|~m1_pre_topc(k2_tsep_1(X3,X1,X2),X4)|~m1_pre_topc(X1,X4)|~m1_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X4)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
fof(c_0_60, plain, ![X62, X63, X64]:(v2_struct_0(X62)|~l1_pre_topc(X62)|(v2_struct_0(X63)|~m1_pre_topc(X63,X62)|(~m1_subset_1(X64,u1_struct_0(X63))|m1_subset_1(X64,u1_struct_0(X62))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_pre_topc])])])])).
cnf(c_0_61, negated_conjecture, (r1_tarski(u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_42]), c_0_41]), c_0_49])]), c_0_44])).
fof(c_0_62, plain, ![X65, X66]:(~l1_pre_topc(X65)|(~m1_pre_topc(X66,X65)|l1_pre_topc(X66))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_63, plain, (r1_tsep_1(X1,X2)|m1_pre_topc(k2_tsep_1(X3,X1,X2),X1)|v2_struct_0(X2)|v2_struct_0(X1)|v2_struct_0(X3)|~v2_pre_topc(X3)|~m1_pre_topc(X1,X3)|~m1_pre_topc(X2,X3)|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_59, c_0_29])).
cnf(c_0_64, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_65, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_67, negated_conjecture, (m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk3_0)|~v2_pre_topc(X1)|~m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_56, c_0_61])).
cnf(c_0_68, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69, negated_conjecture, (r1_tsep_1(esk2_0,X1)|m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,X1),esk2_0)|v2_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_49]), c_0_64]), c_0_42])]), c_0_51]), c_0_44])).
cnf(c_0_70, negated_conjecture, (~m1_subset_1(X1,u1_struct_0(esk2_0))|X1!=esk4_0|~m1_subset_1(X2,u1_struct_0(esk3_0))|X2!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(X1))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_72, negated_conjecture, (m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_29]), c_0_64]), c_0_41]), c_0_42]), c_0_49])]), c_0_43]), c_0_51]), c_0_44])).
cnf(c_0_73, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_41]), c_0_42])])).
cnf(c_0_74, negated_conjecture, (m1_pre_topc(k2_tsep_1(esk1_0,esk2_0,esk3_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_69]), c_0_41])]), c_0_43])).
cnf(c_0_75, negated_conjecture, (l1_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_49]), c_0_42])])).
cnf(c_0_76, negated_conjecture, (~m1_subset_1(esk4_0,u1_struct_0(esk3_0))|~m1_subset_1(esk4_0,u1_struct_0(esk2_0))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_70])])).
cnf(c_0_77, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk3_0))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])]), c_0_43])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk2_0))|v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_74]), c_0_75])]), c_0_51])).
cnf(c_0_79, negated_conjecture, (v2_struct_0(k2_tsep_1(esk1_0,esk2_0,esk3_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_79]), c_0_41]), c_0_49]), c_0_42])]), c_0_43]), c_0_51]), c_0_44]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 52
# Proof object formula steps           : 29
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 64
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 35
# Processed clauses                    : 139
# ...of these trivial                  : 0
# ...subsumed                          : 7
# ...remaining for further processing  : 132
# Other redundant clauses eliminated   : 16
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 19
# Generated clauses                    : 251
# ...of the previous two non-trivial   : 227
# Contextual simplify-reflections      : 12
# Paramodulations                      : 241
# Factorizations                       : 0
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 65
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 152
# ...number of literals in the above   : 2215
# Current number of archived formulas  : 0
# Current number of archived clauses   : 63
# Clause-clause subsumption calls (NU) : 2601
# Rec. Clause-clause subsumption calls : 167
# Non-unit clause-clause subsumptions  : 21
# Unit Clause-clause subsumption calls : 132
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 35
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14460

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.003 s
# Total time               : 0.019 s
# Maximum resident set size: 28644 pages
