# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t1_boole)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', commutativity_k2_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', commutativity_k3_xboole_0)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t22_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', d5_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', l32_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t4_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', idempotence_k3_xboole_0)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t16_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t28_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t3_xboole_0)).
fof(t29_xboole_1, axiom, ![X1, X2, X3]:r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X1,X3)), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t29_xboole_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t39_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t41_xboole_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t7_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t36_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t3_boole)).
fof(t47_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/9bfbecb9f3a84978bbe4902336bff9b7.E_input', t47_xboole_1)).
fof(c_0_19, plain, ![X26]:k2_xboole_0(X26,k1_xboole_0)=X26, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_20, plain, ![X10, X11]:k2_xboole_0(X10,X11)=k2_xboole_0(X11,X10), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_21, plain, ![X52, X53]:r1_tarski(k3_xboole_0(X52,X53),X52), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_22, plain, ![X44, X45]:k3_xboole_0(X44,X45)=k3_xboole_0(X45,X44), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_23, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_25, plain, ![X40, X41]:k2_xboole_0(X40,k3_xboole_0(X40,X41))=X40, inference(variable_rename,[status(thm)],[t22_xboole_1])).
cnf(c_0_26, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_30, plain, ![X27, X28, X29, X30, X31, X32, X33, X34]:((((r2_hidden(X30,X27)|~r2_hidden(X30,X29)|X29!=k4_xboole_0(X27,X28))&(~r2_hidden(X30,X28)|~r2_hidden(X30,X29)|X29!=k4_xboole_0(X27,X28)))&(~r2_hidden(X31,X27)|r2_hidden(X31,X28)|r2_hidden(X31,X29)|X29!=k4_xboole_0(X27,X28)))&((~r2_hidden(esk4_3(X32,X33,X34),X34)|(~r2_hidden(esk4_3(X32,X33,X34),X32)|r2_hidden(esk4_3(X32,X33,X34),X33))|X34=k4_xboole_0(X32,X33))&((r2_hidden(esk4_3(X32,X33,X34),X32)|r2_hidden(esk4_3(X32,X33,X34),X34)|X34=k4_xboole_0(X32,X33))&(~r2_hidden(esk4_3(X32,X33,X34),X33)|r2_hidden(esk4_3(X32,X33,X34),X34)|X34=k4_xboole_0(X32,X33))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_31, plain, ![X12, X13]:((k4_xboole_0(X12,X13)!=k1_xboole_0|r1_tarski(X12,X13))&(~r1_tarski(X12,X13)|k4_xboole_0(X12,X13)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_32, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_37, plain, ![X46, X47, X49, X50, X51]:((r1_xboole_0(X46,X47)|r2_hidden(esk5_2(X46,X47),k3_xboole_0(X46,X47)))&(~r2_hidden(X51,k3_xboole_0(X49,X50))|~r1_xboole_0(X49,X50))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_38, plain, ![X39]:k3_xboole_0(X39,X39)=X39, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_39, plain, ![X23, X24, X25]:k3_xboole_0(k3_xboole_0(X23,X24),X25)=k3_xboole_0(X23,k3_xboole_0(X24,X25)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_40, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_43, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_44, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_45, plain, ![X56, X57]:(~r1_tarski(X56,X57)|k3_xboole_0(X56,X57)=X56), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_46, plain, ![X17, X18, X20, X21, X22]:(((r2_hidden(esk3_2(X17,X18),X17)|r1_xboole_0(X17,X18))&(r2_hidden(esk3_2(X17,X18),X18)|r1_xboole_0(X17,X18)))&(~r2_hidden(X22,X20)|~r2_hidden(X22,X21)|~r1_xboole_0(X20,X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_47, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_49, plain, (r2_hidden(esk5_2(X1,X1),X1)|r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_33])).
fof(c_0_51, plain, ![X7, X8, X9]:r1_tarski(k3_xboole_0(X7,X8),k2_xboole_0(X7,X9)), inference(variable_rename,[status(thm)],[t29_xboole_1])).
fof(c_0_52, plain, ![X42, X43]:k2_xboole_0(X42,k4_xboole_0(X43,X42))=k2_xboole_0(X42,X43), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_53, plain, ![X36, X37, X38]:k4_xboole_0(k4_xboole_0(X36,X37),X38)=k4_xboole_0(X36,k2_xboole_0(X37,X38)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_54, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_44, c_0_43])).
cnf(c_0_55, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_57, plain, ![X54]:(X54=k1_xboole_0|r2_hidden(esk6_1(X54),X54)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_58, plain, (r2_hidden(esk3_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_59, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_47, c_0_27])).
cnf(c_0_60, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_47, c_0_43])).
cnf(c_0_61, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)|~r2_hidden(esk5_2(k1_xboole_0,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_62, plain, (r2_hidden(esk5_2(X1,k1_xboole_0),k1_xboole_0)|r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_42, c_0_50])).
fof(c_0_63, plain, ![X15, X16]:r1_tarski(k4_xboole_0(X15,X16),X15), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_64, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_65, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_66, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_67, plain, (~r2_hidden(X1,k3_xboole_0(X2,k3_xboole_0(X3,X4)))|~r1_xboole_0(k3_xboole_0(X2,X3),X4)), inference(spm,[status(thm)],[c_0_47, c_0_44])).
cnf(c_0_68, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_54, c_0_27])).
cnf(c_0_69, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_70, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_71, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk3_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_40, c_0_58])).
cnf(c_0_72, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_73, plain, (~r2_hidden(X1,k1_xboole_0)|~r1_xboole_0(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_59, c_0_33])).
cnf(c_0_74, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_60, c_0_58])).
cnf(c_0_75, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_76, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_77, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_64, c_0_43])).
cnf(c_0_78, plain, (k2_xboole_0(X1,k4_xboole_0(X2,k2_xboole_0(X3,X1)))=k2_xboole_0(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_79, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_26])).
cnf(c_0_80, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(k3_xboole_0(X3,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_43]), c_0_44]), c_0_68])).
cnf(c_0_81, plain, (k3_xboole_0(X1,X2)=X2|k4_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_69])).
cnf(c_0_82, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_47, c_0_70])).
cnf(c_0_83, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_84, plain, (r1_xboole_0(k1_xboole_0,X1)|~r1_xboole_0(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_73, c_0_72])).
cnf(c_0_85, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_86, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_26, c_0_44])).
cnf(c_0_87, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_76]), c_0_27])).
cnf(c_0_88, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_77, c_0_24])).
cnf(c_0_89, plain, (k2_xboole_0(X1,k4_xboole_0(k3_xboole_0(k2_xboole_0(X2,X1),X3),X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_23])).
cnf(c_0_90, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_64, c_0_24])).
cnf(c_0_91, plain, (k4_xboole_0(X1,X2)!=k1_xboole_0|~r2_hidden(X3,X1)|~r1_xboole_0(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_92, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_93, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_85])])).
cnf(c_0_94, plain, (r1_tarski(k3_xboole_0(X1,k4_xboole_0(X2,X3)),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
fof(c_0_95, plain, ![X14]:k4_xboole_0(X14,k1_xboole_0)=X14, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_96, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X1),X2)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_97, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_90, c_0_87])).
cnf(c_0_98, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))!=k1_xboole_0|~r2_hidden(X3,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93])])).
cnf(c_0_99, plain, (k4_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X3)),k3_xboole_0(X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_94])).
cnf(c_0_100, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_101, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_102, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X2),X1)), inference(spm,[status(thm)],[c_0_96, c_0_24])).
cnf(c_0_103, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k2_xboole_0(X3,X1))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_55, c_0_97])).
cnf(c_0_104, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_105, plain, (k3_xboole_0(X1,X2)!=k1_xboole_0|~r2_hidden(X3,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_100])).
cnf(c_0_106, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_101])).
cnf(c_0_107, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),k4_xboole_0(X2,X1)),X1)), inference(spm,[status(thm)],[c_0_102, c_0_65])).
cnf(c_0_108, plain, (k3_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X3))=k4_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_27, c_0_103])).
cnf(c_0_109, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_104])).
cnf(c_0_110, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X2,k3_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_27]), c_0_44])).
cnf(c_0_111, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,k3_xboole_0(X1,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_54, c_0_69])).
cnf(c_0_112, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_113, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X3))),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_66])).
cnf(c_0_114, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_67, c_0_42])).
cnf(c_0_115, plain, (r2_hidden(esk3_2(k4_xboole_0(X1,X2),X3),X1)|r1_xboole_0(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_109, c_0_72])).
cnf(c_0_116, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X1,X3)))=k3_xboole_0(X2,k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_110, c_0_87])).
cnf(c_0_117, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_44])).
cnf(c_0_118, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_113, c_0_28])).
cnf(c_0_119, plain, (r1_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,k3_xboole_0(X1,X2))))), inference(spm,[status(thm)],[c_0_114, c_0_83])).
cnf(c_0_120, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_59, c_0_70])).
cnf(c_0_121, plain, (r1_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_71, c_0_115])).
cnf(c_0_122, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k3_xboole_0(X2,X1)|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_123, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_118]), c_0_27])).
cnf(c_0_124, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_119, c_0_87])).
cnf(c_0_125, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))=k1_xboole_0), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_126, plain, (X1=X2|k4_xboole_0(X2,X1)!=k1_xboole_0|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_69, c_0_81])).
cnf(c_0_127, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_128, plain, (k4_xboole_0(X1,k2_xboole_0(X2,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_76]), c_0_66])).
cnf(c_0_129, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_124, c_0_27])).
fof(c_0_130, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t47_xboole_1])).
cnf(c_0_131, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_125, c_0_123])).
cnf(c_0_132, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_77])).
cnf(c_0_133, plain, (k4_xboole_0(X1,X2)=X1|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_66]), c_0_128])])).
cnf(c_0_134, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_82, c_0_129])).
fof(c_0_135, negated_conjecture, k4_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(esk1_0,esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_130])])])).
cnf(c_0_136, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_131]), c_0_66]), c_0_24]), c_0_65]), c_0_132])])).
cnf(c_0_137, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=X1), inference(spm,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_138, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_29, c_0_27])).
cnf(c_0_139, negated_conjecture, (k4_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_140, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_66]), c_0_24]), c_0_138])).
cnf(c_0_141, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_139, c_0_140])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 142
# Proof object clause steps            : 103
# Proof object formula steps           : 39
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 75
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 18719
# ...of these trivial                  : 4920
# ...subsumed                          : 11571
# ...remaining for further processing  : 2228
# Other redundant clauses eliminated   : 9064
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 115
# Generated clauses                    : 947522
# ...of the previous two non-trivial   : 508895
# Contextual simplify-reflections      : 4
# Paramodulations                      : 938452
# Factorizations                       : 6
# Equation resolutions                 : 9064
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2106
#    Positive orientable unit clauses  : 1353
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 745
# Current number of unprocessed clauses: 489588
# ...number of literals in the above   : 1081936
# Current number of archived formulas  : 0
# Current number of archived clauses   : 119
# Clause-clause subsumption calls (NU) : 100973
# Rec. Clause-clause subsumption calls : 84536
# Non-unit clause-clause subsumptions  : 5176
# Unit Clause-clause subsumption calls : 1804
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14137
# BW rewrite match successes           : 190
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7963479

# -------------------------------------------------
# User time                : 5.323 s
# System time              : 0.120 s
# Total time               : 5.443 s
# Maximum resident set size: 341716 pages
