# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_tmap_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>(r1_subset_1(X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))&(k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X4)=X6))))))))), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', t6_tmap_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', d10_xboole_0)).
fof(redefinition_r1_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))=>(r1_subset_1(X1,X2)<=>r1_xboole_0(X1,X2))), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', redefinition_r1_subset_1)).
fof(cc5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X2))=>v4_relat_1(X3,X1))), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', cc5_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', fc6_relat_1)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', d18_relat_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', d7_xboole_0)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', redefinition_k9_subset_1)).
fof(t207_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_xboole_0(X1,X2)=>k7_relat_1(k7_relat_1(X3,X1),X2)=k1_xboole_0)), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', t207_relat_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', t65_xboole_1)).
fof(redefinition_k2_partfun1, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k2_partfun1(X1,X2,X3,X4)=k7_relat_1(X3,X4)), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', redefinition_k2_partfun1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', cc1_relset_1)).
fof(dt_k1_tmap_1, axiom, ![X1, X2, X3, X4, X5, X6]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))&((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))&((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))&(v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2))))))))))=>(v1_funct_1(k1_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k1_tmap_1(X1,X2,X3,X4,X5,X6),k4_subset_1(X1,X3,X4),X2)&m1_subset_1(k1_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))))), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', dt_k1_tmap_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', t3_subset)).
fof(t209_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>X2=k7_relat_1(X2,X1)), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', t209_relat_1)).
fof(d1_tmap_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))=>![X7]:((v1_funct_1(X7)&(v1_funct_2(X7,k4_subset_1(X1,X3,X4),X2)&m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))))=>(X7=k1_tmap_1(X1,X2,X3,X4,X5,X6)<=>(k2_partfun1(k4_subset_1(X1,X3,X4),X2,X7,X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,X7,X4)=X6)))))))))), file('data/MPTP/proofs/9d4d636dd6334c33aba6ee2507ae73ba.E_input', d1_tmap_1)).
fof(c_0_16, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:((~(v1_xboole_0(X3))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>![X4]:((~(v1_xboole_0(X4))&m1_subset_1(X4,k1_zfmisc_1(X1)))=>(r1_subset_1(X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))))=>(k2_partfun1(X3,X2,X5,k9_subset_1(X1,X3,X4))=k2_partfun1(X4,X2,X6,k9_subset_1(X1,X3,X4))&(k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X3)=X5&k2_partfun1(k4_subset_1(X1,X3,X4),X2,k1_tmap_1(X1,X2,X3,X4,X5,X6),X4)=X6)))))))))), inference(assume_negation,[status(cth)],[t6_tmap_1])).
fof(c_0_17, plain, ![X14, X15]:(((r1_tarski(X14,X15)|X14!=X15)&(r1_tarski(X15,X14)|X14!=X15))&(~r1_tarski(X14,X15)|~r1_tarski(X15,X14)|X14=X15)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_18, plain, ![X16, X17]:((~r1_subset_1(X16,X17)|r1_xboole_0(X16,X17)|(v1_xboole_0(X16)|v1_xboole_0(X17)))&(~r1_xboole_0(X16,X17)|r1_subset_1(X16,X17)|(v1_xboole_0(X16)|v1_xboole_0(X17)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r1_subset_1])])])])).
fof(c_0_19, negated_conjecture, (~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&((~v1_xboole_0(esk3_0)&m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0)))&((~v1_xboole_0(esk4_0)&m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0)))&(r1_subset_1(esk3_0,esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,esk3_0,esk2_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk2_0)))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,esk4_0,esk2_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0)))))&(k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk1_0,esk3_0,esk4_0))!=k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk1_0,esk3_0,esk4_0))|(k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])).
fof(c_0_20, plain, ![X18, X19, X20]:(~v1_relat_1(X19)|~v4_relat_1(X19,X18)|(~m1_subset_1(X20,k1_zfmisc_1(X19))|v4_relat_1(X20,X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_relat_1])])])).
fof(c_0_21, plain, ![X21, X22]:v1_relat_1(k2_zfmisc_1(X21,X22)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_22, plain, ![X23, X24]:((~v4_relat_1(X24,X23)|r1_tarski(k1_relat_1(X24),X23)|~v1_relat_1(X24))&(~r1_tarski(k1_relat_1(X24),X23)|v4_relat_1(X24,X23)|~v1_relat_1(X24))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
cnf(c_0_23, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X25, X26]:((~r1_xboole_0(X25,X26)|k3_xboole_0(X25,X26)=k1_xboole_0)&(k3_xboole_0(X25,X26)!=k1_xboole_0|r1_xboole_0(X25,X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_25, plain, (r1_xboole_0(X1,X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (r1_subset_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_29, plain, ![X27, X28, X29]:(~m1_subset_1(X29,k1_zfmisc_1(X27))|k9_subset_1(X27,X28,X29)=k3_xboole_0(X28,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
fof(c_0_30, plain, ![X30, X31, X32]:(~v1_relat_1(X32)|(~r1_xboole_0(X30,X31)|k7_relat_1(k7_relat_1(X32,X30),X31)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t207_relat_1])])).
fof(c_0_31, plain, ![X33]:r1_xboole_0(X33,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_32, plain, ![X34, X35, X36, X37]:(~v1_funct_1(X36)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X34,X35)))|k2_partfun1(X34,X35,X36,X37)=k7_relat_1(X36,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_partfun1])])).
fof(c_0_33, plain, ![X38, X39, X40]:(~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))|v1_relat_1(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_34, plain, (v4_relat_1(X3,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_37, plain, (v4_relat_1(X1,X2)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_38, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_23])).
fof(c_0_39, plain, ![X41, X42, X43, X44, X45, X46]:((v1_funct_1(k1_tmap_1(X41,X42,X43,X44,X45,X46))|(v1_xboole_0(X41)|(v1_xboole_0(X42)|(v1_xboole_0(X43)|~m1_subset_1(X43,k1_zfmisc_1(X41))|(v1_xboole_0(X44)|~m1_subset_1(X44,k1_zfmisc_1(X41))|(~v1_funct_1(X45)|(~v1_funct_2(X45,X43,X42)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X42))))|(~v1_funct_1(X46)|(~v1_funct_2(X46,X44,X42)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X42)))))))))))&((v1_funct_2(k1_tmap_1(X41,X42,X43,X44,X45,X46),k4_subset_1(X41,X43,X44),X42)|(v1_xboole_0(X41)|(v1_xboole_0(X42)|(v1_xboole_0(X43)|~m1_subset_1(X43,k1_zfmisc_1(X41))|(v1_xboole_0(X44)|~m1_subset_1(X44,k1_zfmisc_1(X41))|(~v1_funct_1(X45)|(~v1_funct_2(X45,X43,X42)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X42))))|(~v1_funct_1(X46)|(~v1_funct_2(X46,X44,X42)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X42)))))))))))&(m1_subset_1(k1_tmap_1(X41,X42,X43,X44,X45,X46),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X41,X43,X44),X42)))|(v1_xboole_0(X41)|(v1_xboole_0(X42)|(v1_xboole_0(X43)|~m1_subset_1(X43,k1_zfmisc_1(X41))|(v1_xboole_0(X44)|~m1_subset_1(X44,k1_zfmisc_1(X41))|(~v1_funct_1(X45)|(~v1_funct_2(X45,X43,X42)|~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X43,X42))))|(~v1_funct_1(X46)|(~v1_funct_2(X46,X44,X42)|~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X42))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tmap_1])])])])).
fof(c_0_40, plain, ![X47, X48]:((~m1_subset_1(X47,k1_zfmisc_1(X48))|r1_tarski(X47,X48))&(~r1_tarski(X47,X48)|m1_subset_1(X47,k1_zfmisc_1(X48)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_41, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42, negated_conjecture, (r1_xboole_0(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_28])).
cnf(c_0_43, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_45, plain, (k7_relat_1(k7_relat_1(X1,X2),X3)=k1_xboole_0|~v1_relat_1(X1)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_46, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_47, plain, (k2_partfun1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_48, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_49, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_50, plain, ![X49, X50]:(~v1_relat_1(X50)|~v4_relat_1(X50,X49)|X50=k7_relat_1(X50,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t209_relat_1])])).
cnf(c_0_51, negated_conjecture, (v4_relat_1(esk6_0,X1)|~v4_relat_1(k2_zfmisc_1(esk4_0,esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_52, plain, (v4_relat_1(X1,k1_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_53, plain, (m1_subset_1(k1_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X3,X4),X2)))|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_54, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_55, plain, (v1_funct_1(k1_tmap_1(X1,X2,X3,X4,X5,X6))|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_57, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_58, negated_conjecture, (k3_xboole_0(esk3_0,esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_59, negated_conjecture, (k3_xboole_0(X1,esk4_0)=k9_subset_1(esk1_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_60, plain, (k7_relat_1(k7_relat_1(X1,X2),k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_61, negated_conjecture, (k7_relat_1(esk6_0,X1)=k2_partfun1(esk4_0,esk2_0,esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_35]), c_0_48])])).
cnf(c_0_62, negated_conjecture, (v1_relat_1(esk6_0)), inference(spm,[status(thm)],[c_0_49, c_0_35])).
cnf(c_0_63, plain, (X1=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_64, negated_conjecture, (v4_relat_1(esk6_0,k1_relat_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_36])])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(k1_tmap_1(esk1_0,X1,X2,esk4_0,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(esk1_0,X2,esk4_0),X1)))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_2(X4,esk4_0,X1)|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk4_0,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~m1_subset_1(X2,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_44]), c_0_27]), c_0_54])).
cnf(c_0_66, negated_conjecture, (v1_funct_2(esk6_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_67, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_68, negated_conjecture, (v1_funct_1(k1_tmap_1(esk1_0,X1,X2,esk4_0,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_2(X4,esk4_0,X1)|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk4_0,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~m1_subset_1(X2,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_44]), c_0_27]), c_0_54])).
cnf(c_0_69, negated_conjecture, (v4_relat_1(esk5_0,X1)|~v4_relat_1(k2_zfmisc_1(esk3_0,esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_56]), c_0_36])])).
fof(c_0_70, plain, ![X51, X52, X53, X54, X55, X56, X57]:(((k2_partfun1(k4_subset_1(X51,X53,X54),X52,X57,X53)=X55|X57!=k1_tmap_1(X51,X52,X53,X54,X55,X56)|(~v1_funct_1(X57)|(~v1_funct_2(X57,k4_subset_1(X51,X53,X54),X52)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X51,X53,X54),X52)))))|k2_partfun1(X53,X52,X55,k9_subset_1(X51,X53,X54))!=k2_partfun1(X54,X52,X56,k9_subset_1(X51,X53,X54))|(~v1_funct_1(X56)|(~v1_funct_2(X56,X54,X52)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X52)))))|(~v1_funct_1(X55)|(~v1_funct_2(X55,X53,X52)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X52)))))|(v1_xboole_0(X54)|~m1_subset_1(X54,k1_zfmisc_1(X51)))|(v1_xboole_0(X53)|~m1_subset_1(X53,k1_zfmisc_1(X51)))|v1_xboole_0(X52)|v1_xboole_0(X51))&(k2_partfun1(k4_subset_1(X51,X53,X54),X52,X57,X54)=X56|X57!=k1_tmap_1(X51,X52,X53,X54,X55,X56)|(~v1_funct_1(X57)|(~v1_funct_2(X57,k4_subset_1(X51,X53,X54),X52)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X51,X53,X54),X52)))))|k2_partfun1(X53,X52,X55,k9_subset_1(X51,X53,X54))!=k2_partfun1(X54,X52,X56,k9_subset_1(X51,X53,X54))|(~v1_funct_1(X56)|(~v1_funct_2(X56,X54,X52)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X52)))))|(~v1_funct_1(X55)|(~v1_funct_2(X55,X53,X52)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X52)))))|(v1_xboole_0(X54)|~m1_subset_1(X54,k1_zfmisc_1(X51)))|(v1_xboole_0(X53)|~m1_subset_1(X53,k1_zfmisc_1(X51)))|v1_xboole_0(X52)|v1_xboole_0(X51)))&(k2_partfun1(k4_subset_1(X51,X53,X54),X52,X57,X53)!=X55|k2_partfun1(k4_subset_1(X51,X53,X54),X52,X57,X54)!=X56|X57=k1_tmap_1(X51,X52,X53,X54,X55,X56)|(~v1_funct_1(X57)|(~v1_funct_2(X57,k4_subset_1(X51,X53,X54),X52)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X51,X53,X54),X52)))))|k2_partfun1(X53,X52,X55,k9_subset_1(X51,X53,X54))!=k2_partfun1(X54,X52,X56,k9_subset_1(X51,X53,X54))|(~v1_funct_1(X56)|(~v1_funct_2(X56,X54,X52)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X52)))))|(~v1_funct_1(X55)|(~v1_funct_2(X55,X53,X52)|~m1_subset_1(X55,k1_zfmisc_1(k2_zfmisc_1(X53,X52)))))|(v1_xboole_0(X54)|~m1_subset_1(X54,k1_zfmisc_1(X51)))|(v1_xboole_0(X53)|~m1_subset_1(X53,k1_zfmisc_1(X51)))|v1_xboole_0(X52)|v1_xboole_0(X51))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tmap_1])])])])])).
cnf(c_0_71, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_57, c_0_38])).
cnf(c_0_72, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k9_subset_1(esk1_0,esk3_0,esk4_0))!=k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk1_0,esk3_0,esk4_0))|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_73, negated_conjecture, (k9_subset_1(esk1_0,esk3_0,esk4_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_74, negated_conjecture, (k7_relat_1(k2_partfun1(esk4_0,esk2_0,esk6_0,X1),k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])])).
cnf(c_0_75, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,k1_relat_1(k2_zfmisc_1(esk4_0,esk2_0)))=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_61]), c_0_62])])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(k1_tmap_1(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(esk1_0,X1,esk4_0),esk2_0)))|v1_xboole_0(X1)|~v1_funct_2(X2,X1,esk2_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_48]), c_0_35])]), c_0_67])).
cnf(c_0_77, negated_conjecture, (v1_funct_2(esk5_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_78, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_80, negated_conjecture, (v1_funct_1(k1_tmap_1(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0))|v1_xboole_0(X1)|~v1_funct_2(X2,X1,esk2_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_66]), c_0_48]), c_0_35])]), c_0_67])).
cnf(c_0_81, negated_conjecture, (v4_relat_1(esk5_0,k1_relat_1(k2_zfmisc_1(esk3_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_52]), c_0_36])])).
cnf(c_0_82, negated_conjecture, (v1_relat_1(esk5_0)), inference(spm,[status(thm)],[c_0_49, c_0_56])).
cnf(c_0_83, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,X5,X3)=X6|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X4)|v1_xboole_0(X1)|X5!=k1_tmap_1(X1,X4,X2,X3,X7,X6)|~v1_funct_1(X5)|~v1_funct_2(X5,k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X2,X3),X4)))|k2_partfun1(X2,X4,X7,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_1(X6)|~v1_funct_2(X6,X3,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X7)|~v1_funct_2(X7,X2,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_84, plain, (v1_funct_2(k1_tmap_1(X1,X2,X3,X4,X5,X6),k4_subset_1(X1,X3,X4),X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~v1_funct_1(X6)|~v1_funct_2(X6,X4,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_85, plain, (k9_subset_1(X1,X2,X1)=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_43, c_0_71])).
cnf(c_0_86, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk4_0,esk2_0,esk6_0,k1_xboole_0)!=k2_partfun1(esk3_0,esk2_0,esk5_0,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73]), c_0_73])).
cnf(c_0_87, negated_conjecture, (k2_partfun1(esk4_0,esk2_0,esk6_0,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_61])).
cnf(c_0_88, negated_conjecture, (m1_subset_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), c_0_56]), c_0_79])]), c_0_28])).
cnf(c_0_89, negated_conjecture, (v1_funct_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_77]), c_0_78]), c_0_56]), c_0_79])]), c_0_28])).
cnf(c_0_90, negated_conjecture, (k7_relat_1(esk5_0,X1)=k2_partfun1(esk3_0,esk2_0,esk5_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_56]), c_0_78])])).
cnf(c_0_91, negated_conjecture, (k7_relat_1(esk5_0,k1_relat_1(k2_zfmisc_1(esk3_0,esk2_0)))=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_81]), c_0_82])])).
cnf(c_0_92, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,k1_tmap_1(X1,X4,X2,X3,X5,X6),X3)=X6|v1_xboole_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|k2_partfun1(X2,X4,X5,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_2(X5,X2,X4)|~v1_funct_2(X6,X3,X4)|~v1_funct_1(X5)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_83]), c_0_53]), c_0_55]), c_0_84])).
cnf(c_0_93, negated_conjecture, (k9_subset_1(esk1_0,X1,esk4_0)=k9_subset_1(esk4_0,X1,esk4_0)), inference(rw,[status(thm)],[c_0_59, c_0_85])).
cnf(c_0_94, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,X5,X2)=X6|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X4)|v1_xboole_0(X1)|X5!=k1_tmap_1(X1,X4,X2,X3,X6,X7)|~v1_funct_1(X5)|~v1_funct_2(X5,k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k4_subset_1(X1,X2,X3),X4)))|k2_partfun1(X2,X4,X6,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X7,k9_subset_1(X1,X2,X3))|~v1_funct_1(X7)|~v1_funct_2(X7,X3,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X6)|~v1_funct_2(X6,X2,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_95, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk3_0,esk2_0,esk5_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_96, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,esk3_0,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),X1)=k7_relat_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_88]), c_0_89])])).
cnf(c_0_97, negated_conjecture, (k7_relat_1(k2_partfun1(esk3_0,esk2_0,esk5_0,X1),k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_90]), c_0_82])])).
cnf(c_0_98, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k1_relat_1(k2_zfmisc_1(esk3_0,esk2_0)))=esk5_0), inference(rw,[status(thm)],[c_0_91, c_0_90])).
cnf(c_0_99, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,X1,esk4_0),X2,k1_tmap_1(esk1_0,X2,X1,esk4_0,X3,X4),esk4_0)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k2_partfun1(X1,X2,X3,k9_subset_1(esk4_0,X1,esk4_0))!=k2_partfun1(esk4_0,X2,X4,k9_subset_1(esk4_0,X1,esk4_0))|~v1_funct_2(X4,esk4_0,X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk4_0,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_44]), c_0_54]), c_0_27]), c_0_93]), c_0_93])).
cnf(c_0_100, plain, (k2_partfun1(k4_subset_1(X1,X2,X3),X4,k1_tmap_1(X1,X4,X2,X3,X5,X6),X2)=X5|v1_xboole_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|k2_partfun1(X2,X4,X5,k9_subset_1(X1,X2,X3))!=k2_partfun1(X3,X4,X6,k9_subset_1(X1,X2,X3))|~v1_funct_2(X6,X3,X4)|~v1_funct_2(X5,X2,X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_94]), c_0_53]), c_0_55]), c_0_84])).
cnf(c_0_101, negated_conjecture, (k7_relat_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k7_relat_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0|k2_partfun1(esk3_0,esk2_0,esk5_0,k1_xboole_0)!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_96]), c_0_96])).
cnf(c_0_102, negated_conjecture, (k2_partfun1(esk3_0,esk2_0,esk5_0,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_90])).
cnf(c_0_103, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,X1,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0),esk4_0)=esk6_0|v1_xboole_0(X1)|k2_partfun1(X1,esk2_0,X2,k9_subset_1(esk4_0,X1,esk4_0))!=k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk4_0,X1,esk4_0))|~v1_funct_2(X2,X1,esk2_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_66]), c_0_48]), c_0_35])]), c_0_67])).
cnf(c_0_104, negated_conjecture, (k9_subset_1(esk4_0,esk3_0,esk4_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_73, c_0_93])).
cnf(c_0_105, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,X1,esk4_0),X2,k1_tmap_1(esk1_0,X2,X1,esk4_0,X3,X4),X1)=X3|v1_xboole_0(X2)|v1_xboole_0(X1)|k2_partfun1(X1,X2,X3,k9_subset_1(esk4_0,X1,esk4_0))!=k2_partfun1(esk4_0,X2,X4,k9_subset_1(esk4_0,X1,esk4_0))|~v1_funct_2(X4,esk4_0,X2)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk4_0,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_44]), c_0_54]), c_0_27]), c_0_93]), c_0_93])).
cnf(c_0_106, negated_conjecture, (k7_relat_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0|k7_relat_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)!=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_107, negated_conjecture, (k7_relat_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk4_0)=esk6_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_77]), c_0_96]), c_0_104]), c_0_102]), c_0_104]), c_0_87]), c_0_78]), c_0_56]), c_0_79])]), c_0_28])).
cnf(c_0_108, negated_conjecture, (k2_partfun1(k4_subset_1(esk1_0,X1,esk4_0),esk2_0,k1_tmap_1(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0),X1)=X2|v1_xboole_0(X1)|k2_partfun1(X1,esk2_0,X2,k9_subset_1(esk4_0,X1,esk4_0))!=k2_partfun1(esk4_0,esk2_0,esk6_0,k9_subset_1(esk4_0,X1,esk4_0))|~v1_funct_2(X2,X1,esk2_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_66]), c_0_48]), c_0_35])]), c_0_67])).
cnf(c_0_109, negated_conjecture, (k7_relat_1(k1_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk3_0)!=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_107])])).
cnf(c_0_110, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_77]), c_0_96]), c_0_104]), c_0_102]), c_0_104]), c_0_87]), c_0_78]), c_0_56]), c_0_79])]), c_0_109]), c_0_28]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 111
# Proof object clause steps            : 78
# Proof object formula steps           : 33
# Proof object conjectures             : 56
# Proof object clause conjectures      : 53
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 108
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 39
# Processed clauses                    : 306
# ...of these trivial                  : 4
# ...subsumed                          : 3
# ...remaining for further processing  : 299
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 30
# Generated clauses                    : 333
# ...of the previous two non-trivial   : 307
# Contextual simplify-reflections      : 6
# Paramodulations                      : 328
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 222
#    Positive orientable unit clauses  : 103
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 113
# Current number of unprocessed clauses: 78
# ...number of literals in the above   : 975
# Current number of archived formulas  : 0
# Current number of archived clauses   : 72
# Clause-clause subsumption calls (NU) : 12523
# Rec. Clause-clause subsumption calls : 1839
# Non-unit clause-clause subsumptions  : 13
# Unit Clause-clause subsumption calls : 784
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 203
# BW rewrite match successes           : 29
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 23640

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.003 s
# Total time               : 0.028 s
# Maximum resident set size: 28792 pages
