# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t113_zfmisc_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', cc1_relset_1)).
fof(t18_funct_1, axiom, ![X1, X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k1_relat_1(X3)&r1_tarski(k2_relat_1(X3),X1)))))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t18_funct_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t3_subset)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', d5_funct_1)).
fof(t64_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0)), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t64_relat_1)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', d18_relat_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t7_ordinal1)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t3_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t2_xboole_1)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t60_relat_1)).
fof(t65_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k1_relat_1(X1)=k1_xboole_0<=>k2_relat_1(X1)=k1_xboole_0)), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t65_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', d3_tarski)).
fof(t4_funct_2, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(k2_relat_1(X2),X1)=>(v1_funct_1(X2)&(v1_funct_2(X2,k1_relat_1(X2),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1))))))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t4_funct_2)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', d10_xboole_0)).
fof(t118_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>(r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&r1_tarski(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2)))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t118_zfmisc_1)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t3_funct_2)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', t1_xboole_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', d1_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/9d8a62dc3b2f4a46b899ec34f67583ff.E_input', redefinition_k1_relset_1)).
fof(c_0_21, plain, ![X23, X24]:((k2_zfmisc_1(X23,X24)!=k1_xboole_0|(X23=k1_xboole_0|X24=k1_xboole_0))&((X23!=k1_xboole_0|k2_zfmisc_1(X23,X24)=k1_xboole_0)&(X24!=k1_xboole_0|k2_zfmisc_1(X23,X24)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_22, plain, ![X36, X37, X38]:((v4_relat_1(X38,X36)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X36,X37))))&(v5_relat_1(X38,X37)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X36,X37))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_23, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_24, plain, ![X39, X40, X41]:(~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X39,X40)))|v1_relat_1(X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_25, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_26, plain, ![X15, X16]:((((v1_relat_1(esk4_2(X15,X16))|X15=k1_xboole_0)&(v1_funct_1(esk4_2(X15,X16))|X15=k1_xboole_0))&((X16=k1_relat_1(esk4_2(X15,X16))|X15=k1_xboole_0)&(r1_tarski(k2_relat_1(esk4_2(X15,X16)),X15)|X15=k1_xboole_0)))&(((v1_relat_1(esk4_2(X15,X16))|X16!=k1_xboole_0)&(v1_funct_1(esk4_2(X15,X16))|X16!=k1_xboole_0))&((X16=k1_relat_1(esk4_2(X15,X16))|X16!=k1_xboole_0)&(r1_tarski(k2_relat_1(esk4_2(X15,X16)),X15)|X16!=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_funct_1])])])])])).
cnf(c_0_27, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_23])).
fof(c_0_29, plain, ![X42, X43]:((~m1_subset_1(X42,k1_zfmisc_1(X43))|r1_tarski(X42,X43))&(~r1_tarski(X42,X43)|m1_subset_1(X42,k1_zfmisc_1(X43)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_30, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X5, X6, X7, X9, X10, X11, X13]:((((r2_hidden(esk1_3(X5,X6,X7),k1_relat_1(X5))|~r2_hidden(X7,X6)|X6!=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(X7=k1_funct_1(X5,esk1_3(X5,X6,X7))|~r2_hidden(X7,X6)|X6!=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5))))&(~r2_hidden(X10,k1_relat_1(X5))|X9!=k1_funct_1(X5,X10)|r2_hidden(X9,X6)|X6!=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5))))&((~r2_hidden(esk2_2(X5,X11),X11)|(~r2_hidden(X13,k1_relat_1(X5))|esk2_2(X5,X11)!=k1_funct_1(X5,X13))|X11=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))&((r2_hidden(esk3_2(X5,X11),k1_relat_1(X5))|r2_hidden(esk2_2(X5,X11),X11)|X11=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))&(esk2_2(X5,X11)=k1_funct_1(X5,esk3_2(X5,X11))|r2_hidden(esk2_2(X5,X11),X11)|X11=k2_relat_1(X5)|(~v1_relat_1(X5)|~v1_funct_1(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
fof(c_0_33, plain, ![X20]:((k1_relat_1(X20)!=k1_xboole_0|X20=k1_xboole_0|~v1_relat_1(X20))&(k2_relat_1(X20)!=k1_xboole_0|X20=k1_xboole_0|~v1_relat_1(X20))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
cnf(c_0_34, plain, (X1=k1_relat_1(esk4_2(X2,X1))|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (v1_relat_1(esk4_2(X1,X2))|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_36, plain, ![X49, X50]:((~v4_relat_1(X50,X49)|r1_tarski(k1_relat_1(X50),X49)|~v1_relat_1(X50))&(~r1_tarski(k1_relat_1(X50),X49)|v4_relat_1(X50,X49)|~v1_relat_1(X50))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
cnf(c_0_37, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_38, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_40, plain, ![X18, X19]:(~r2_hidden(X18,X19)|~r1_tarski(X19,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_41, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (v1_funct_1(esk4_2(X1,X2))|X2!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_43, plain, (X1=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (k1_relat_1(esk4_2(X1,k1_xboole_0))=k1_xboole_0), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_45, plain, (v1_relat_1(esk4_2(X1,k1_xboole_0))), inference(er,[status(thm)],[c_0_35])).
fof(c_0_46, plain, ![X44]:(~r1_tarski(X44,k1_xboole_0)|X44=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_47, plain, (r1_tarski(k1_relat_1(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, plain, (v4_relat_1(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_49, plain, (v1_relat_1(X1)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_39, c_0_38])).
cnf(c_0_50, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_41])])).
fof(c_0_52, plain, ![X22]:r1_tarski(k1_xboole_0,X22), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_53, plain, (v1_funct_1(esk4_2(X1,k1_xboole_0))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_54, plain, (esk4_2(X1,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])])).
cnf(c_0_55, plain, (r2_hidden(esk1_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_56, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_57, plain, (r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_58, plain, (~r1_tarski(k2_relat_1(X1),k1_funct_1(X1,X2))|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_59, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_60, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_62, plain, (v1_funct_1(k1_xboole_0)), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_63, plain, (v1_relat_1(k1_xboole_0)), inference(rw,[status(thm)],[c_0_45, c_0_54])).
fof(c_0_64, plain, ![X21]:((k1_relat_1(X21)!=k1_xboole_0|k2_relat_1(X21)=k1_xboole_0|~v1_relat_1(X21))&(k2_relat_1(X21)!=k1_xboole_0|k1_relat_1(X21)=k1_xboole_0|~v1_relat_1(X21))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_relat_1])])])).
cnf(c_0_65, plain, (r2_hidden(esk1_3(X1,k2_relat_1(X1),X2),k1_relat_1(X1))|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_66, plain, (k1_relat_1(X1)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_67, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60]), c_0_61]), c_0_62]), c_0_63])])).
fof(c_0_68, plain, ![X28, X29, X30, X31, X32]:((~r1_tarski(X28,X29)|(~r2_hidden(X30,X28)|r2_hidden(X30,X29)))&((r2_hidden(esk5_2(X31,X32),X31)|r1_tarski(X31,X32))&(~r2_hidden(esk5_2(X31,X32),X32)|r1_tarski(X31,X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_69, plain, (k2_relat_1(X1)=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_70, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(k2_relat_1(X2),X1)=>(v1_funct_1(X2)&(v1_funct_2(X2,k1_relat_1(X2),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1)))))))), inference(assume_negation,[status(cth)],[t4_funct_2])).
fof(c_0_71, plain, ![X57, X58]:(((r1_tarski(X57,X58)|X57!=X58)&(r1_tarski(X58,X57)|X57!=X58))&(~r1_tarski(X57,X58)|~r1_tarski(X58,X57)|X57=X58)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_72, plain, (~r1_tarski(X1,k1_xboole_0)|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_49])).
cnf(c_0_73, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_74, plain, (k1_relat_1(X1)!=k1_xboole_0|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_69]), c_0_60])])).
fof(c_0_75, negated_conjecture, ((v1_relat_1(esk7_0)&v1_funct_1(esk7_0))&(r1_tarski(k2_relat_1(esk7_0),esk6_0)&(~v1_funct_1(esk7_0)|(~v1_funct_2(esk7_0,k1_relat_1(esk7_0),esk6_0)|~m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_70])])])).
cnf(c_0_76, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_77, plain, (r1_tarski(k2_relat_1(X1),X2)|~r1_tarski(X1,k1_xboole_0)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_78, plain, (r1_tarski(k1_relat_1(X1),X2)|k1_relat_1(X1)!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_74, c_0_73])).
cnf(c_0_79, negated_conjecture, (v1_funct_1(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_80, negated_conjecture, (v1_relat_1(esk7_0)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_81, plain, (v4_relat_1(X1,X2)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_27, c_0_38])).
cnf(c_0_82, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_83, plain, (v1_relat_1(X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_30, c_0_38])).
fof(c_0_84, plain, ![X25, X26, X27]:((r1_tarski(k2_zfmisc_1(X25,X27),k2_zfmisc_1(X26,X27))|~r1_tarski(X25,X26))&(r1_tarski(k2_zfmisc_1(X27,X25),k2_zfmisc_1(X27,X26))|~r1_tarski(X25,X26))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t118_zfmisc_1])])])).
cnf(c_0_85, plain, (k2_relat_1(X1)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_56, c_0_77])).
cnf(c_0_86, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_87, negated_conjecture, (r1_tarski(k1_relat_1(esk7_0),X1)|k1_relat_1(esk7_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])])).
cnf(c_0_88, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_89, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_83, c_0_82])).
cnf(c_0_90, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_91, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_92, negated_conjecture, (k2_relat_1(esk7_0)=k1_xboole_0|~r1_tarski(esk7_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_85, c_0_79])).
cnf(c_0_93, negated_conjecture, (X1=k1_relat_1(esk7_0)|k1_relat_1(esk7_0)!=k1_xboole_0|~r1_tarski(X1,k1_relat_1(esk7_0))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_94, plain, (r1_tarski(k1_relat_1(k2_zfmisc_1(X1,X2)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_88]), c_0_89])])).
cnf(c_0_95, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_96, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k1_xboole_0)|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_90, c_0_28])).
fof(c_0_97, plain, ![X48]:((v1_funct_1(X48)|(~v1_relat_1(X48)|~v1_funct_1(X48)))&((v1_funct_2(X48,k1_relat_1(X48),k2_relat_1(X48))|(~v1_relat_1(X48)|~v1_funct_1(X48)))&(m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X48),k2_relat_1(X48))))|(~v1_relat_1(X48)|~v1_funct_1(X48))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_98, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X1,X3)|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_86, c_0_90])).
cnf(c_0_99, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X2)|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_86, c_0_91])).
cnf(c_0_100, negated_conjecture, (~r1_tarski(esk7_0,k1_xboole_0)|~r2_hidden(X1,k1_relat_1(esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_92]), c_0_79]), c_0_80])]), c_0_67])).
cnf(c_0_101, negated_conjecture, (k1_relat_1(k2_zfmisc_1(k1_relat_1(esk7_0),X1))=k1_relat_1(esk7_0)|k1_relat_1(esk7_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_102, plain, (~r1_tarski(X1,k1_xboole_0)|~r2_hidden(X2,k2_zfmisc_1(X3,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_67])).
cnf(c_0_103, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_104, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_105, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X1,X3)|~r1_tarski(X3,X2)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_98, c_0_90])).
cnf(c_0_106, negated_conjecture, (r1_tarski(k2_relat_1(esk7_0),esk6_0)), inference(split_conjunct,[status(thm)],[c_0_75])).
fof(c_0_107, plain, ![X45, X46, X47]:(~r1_tarski(X45,X46)|~r1_tarski(X46,X47)|r1_tarski(X45,X47)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_108, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X2)|~r1_tarski(X3,X1)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_99, c_0_91])).
cnf(c_0_109, negated_conjecture, (~r1_tarski(esk7_0,k1_xboole_0)|~r2_hidden(X1,k2_relat_1(esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_65]), c_0_79]), c_0_80])])).
cnf(c_0_110, negated_conjecture, (k1_relat_1(esk7_0)!=k1_xboole_0|~r2_hidden(X1,k1_relat_1(esk7_0))|~v1_funct_1(k2_zfmisc_1(k1_relat_1(esk7_0),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_101]), c_0_89])])).
cnf(c_0_111, plain, (r1_tarski(k2_zfmisc_1(X1,X2),X3)|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_102, c_0_73])).
cnf(c_0_112, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_113, negated_conjecture, (k2_zfmisc_1(X1,k2_relat_1(esk7_0))=k2_zfmisc_1(X1,esk6_0)|~r1_tarski(esk6_0,k2_relat_1(esk7_0))), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_114, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_115, negated_conjecture, (k2_zfmisc_1(k2_relat_1(esk7_0),X1)=k2_zfmisc_1(esk6_0,X1)|~r1_tarski(esk6_0,k2_relat_1(esk7_0))), inference(spm,[status(thm)],[c_0_108, c_0_106])).
cnf(c_0_116, negated_conjecture, (r1_tarski(k2_relat_1(esk7_0),X1)|~r1_tarski(esk7_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_109, c_0_73])).
cnf(c_0_117, negated_conjecture, (k1_relat_1(esk7_0)!=k1_xboole_0|~r2_hidden(X1,k1_relat_1(esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_28]), c_0_62])])).
cnf(c_0_118, plain, (X1=k2_zfmisc_1(X2,X3)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_86, c_0_111])).
cnf(c_0_119, negated_conjecture, (r1_tarski(esk7_0,k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0))|~r1_tarski(esk6_0,k2_relat_1(esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_79]), c_0_80])])).
cnf(c_0_120, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_114, c_0_60])).
cnf(c_0_121, negated_conjecture, (k2_zfmisc_1(k2_relat_1(esk7_0),esk6_0)=k2_zfmisc_1(esk6_0,k2_relat_1(esk7_0))|~r1_tarski(esk6_0,k2_relat_1(esk7_0))), inference(spm,[status(thm)],[c_0_115, c_0_113])).
cnf(c_0_122, negated_conjecture, (k2_relat_1(esk7_0)=k1_relat_1(esk7_0)|~r1_tarski(esk7_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_116]), c_0_66])).
cnf(c_0_123, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_zfmisc_1(X3,X4))|~r1_tarski(X4,k1_xboole_0)), inference(spm,[status(thm)],[c_0_114, c_0_111])).
cnf(c_0_124, negated_conjecture, (k1_relat_1(esk7_0)!=k1_xboole_0|~r2_hidden(X1,k2_relat_1(esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_65]), c_0_79]), c_0_80])])).
cnf(c_0_125, negated_conjecture, (k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0)=esk7_0|~r1_tarski(esk6_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_120])).
cnf(c_0_126, negated_conjecture, (k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0)=k2_zfmisc_1(esk6_0,k1_relat_1(esk7_0))|~r1_tarski(esk6_0,k1_relat_1(esk7_0))|~r1_tarski(esk7_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_127, negated_conjecture, (r1_tarski(esk7_0,X1)|~r1_tarski(esk6_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_119]), c_0_120])).
cnf(c_0_128, negated_conjecture, (r1_tarski(k2_relat_1(esk7_0),X1)|k1_relat_1(esk7_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_124, c_0_73])).
cnf(c_0_129, negated_conjecture, (v4_relat_1(k1_relat_1(esk7_0),X1)|k1_relat_1(esk7_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_81, c_0_87])).
cnf(c_0_130, negated_conjecture, (v1_relat_1(k1_relat_1(esk7_0))|k1_relat_1(esk7_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_83, c_0_87])).
fof(c_0_131, plain, ![X51, X52, X53]:((((~v1_funct_2(X53,X51,X52)|X51=k1_relset_1(X51,X52,X53)|X52=k1_xboole_0|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52))))&(X51!=k1_relset_1(X51,X52,X53)|v1_funct_2(X53,X51,X52)|X52=k1_xboole_0|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52)))))&((~v1_funct_2(X53,X51,X52)|X51=k1_relset_1(X51,X52,X53)|X51!=k1_xboole_0|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52))))&(X51!=k1_relset_1(X51,X52,X53)|v1_funct_2(X53,X51,X52)|X51!=k1_xboole_0|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52))))))&((~v1_funct_2(X53,X51,X52)|X53=k1_xboole_0|X51=k1_xboole_0|X52!=k1_xboole_0|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52))))&(X53!=k1_xboole_0|v1_funct_2(X53,X51,X52)|X51=k1_xboole_0|X52!=k1_xboole_0|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X51,X52)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_132, plain, ![X54, X55, X56]:(~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X55)))|k1_relset_1(X54,X55,X56)=k1_relat_1(X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_133, plain, (k1_relat_1(k2_zfmisc_1(X1,X2))=X1|~r1_tarski(X1,k1_relat_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_86, c_0_94])).
cnf(c_0_134, negated_conjecture, (k2_zfmisc_1(esk6_0,k1_relat_1(esk7_0))=esk7_0|~r1_tarski(esk6_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_127]), c_0_120])).
cnf(c_0_135, negated_conjecture, (X1=k2_relat_1(esk7_0)|k1_relat_1(esk7_0)!=k1_xboole_0|~r1_tarski(X1,k2_relat_1(esk7_0))), inference(spm,[status(thm)],[c_0_86, c_0_128])).
cnf(c_0_136, negated_conjecture, (r1_tarski(k1_relat_1(k1_relat_1(esk7_0)),X1)|k1_relat_1(esk7_0)!=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_129]), c_0_130])).
cnf(c_0_137, negated_conjecture, (~v1_funct_1(esk7_0)|~v1_funct_2(esk7_0,k1_relat_1(esk7_0),esk6_0)|~m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_138, plain, (v1_funct_2(X3,X1,X2)|X2=k1_xboole_0|X1!=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_139, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_140, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_141, negated_conjecture, (k1_relat_1(esk7_0)=esk6_0|~r1_tarski(esk6_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_120])).
cnf(c_0_142, negated_conjecture, (k1_relat_1(k1_relat_1(esk7_0))=k2_relat_1(esk7_0)|k1_relat_1(esk7_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_143, negated_conjecture, (~m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0)))|~v1_funct_2(esk7_0,k1_relat_1(esk7_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_137, c_0_79])])).
cnf(c_0_144, negated_conjecture, (m1_subset_1(esk7_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0)))|~r1_tarski(esk6_0,k2_relat_1(esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_113]), c_0_79]), c_0_80])])).
cnf(c_0_145, plain, (X1=k1_xboole_0|v1_funct_2(X2,X3,X1)|k1_relset_1(X3,X1,X2)!=X3|~r1_tarski(X2,k2_zfmisc_1(X3,X1))), inference(spm,[status(thm)],[c_0_138, c_0_38])).
cnf(c_0_146, plain, (k1_relset_1(X1,X2,X3)=k1_relat_1(X3)|~r1_tarski(X3,k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_139, c_0_38])).
cnf(c_0_147, negated_conjecture, (v1_funct_2(esk7_0,esk6_0,k2_relat_1(esk7_0))|~r1_tarski(esk6_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_79]), c_0_80])])).
cnf(c_0_148, negated_conjecture, (k2_relat_1(esk7_0)=k1_relat_1(esk6_0)|~r1_tarski(esk6_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_141]), c_0_56])).
cnf(c_0_149, negated_conjecture, (k1_relat_1(k1_relat_1(esk7_0))=k1_relat_1(esk7_0)|k1_relat_1(esk7_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_93, c_0_136])).
cnf(c_0_150, negated_conjecture, (~v1_funct_2(esk7_0,k1_relat_1(esk7_0),esk6_0)|~r1_tarski(esk6_0,k2_relat_1(esk7_0))), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_151, negated_conjecture, (~v1_funct_2(esk7_0,k1_relat_1(esk7_0),esk6_0)|~r1_tarski(esk7_0,k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0))), inference(spm,[status(thm)],[c_0_143, c_0_38])).
cnf(c_0_152, plain, (X1=k1_xboole_0|v1_funct_2(X2,k1_relat_1(X2),X1)|~r1_tarski(X2,k2_zfmisc_1(k1_relat_1(X2),X1))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_153, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(X1,k2_zfmisc_1(X2,X4))|~r1_tarski(X4,X3)), inference(spm,[status(thm)],[c_0_114, c_0_90])).
cnf(c_0_154, negated_conjecture, (v1_funct_2(esk7_0,esk6_0,k1_relat_1(esk6_0))|~r1_tarski(esk6_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_147, c_0_148])).
cnf(c_0_155, negated_conjecture, (k1_relat_1(esk6_0)=esk6_0|~r1_tarski(esk6_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_141]), c_0_56])).
cnf(c_0_156, negated_conjecture, (~v1_funct_2(esk7_0,esk6_0,esk6_0)|~r1_tarski(esk6_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_141]), c_0_120])).
cnf(c_0_157, negated_conjecture, (esk6_0=k1_xboole_0|~r1_tarski(esk7_0,k2_zfmisc_1(k1_relat_1(esk7_0),esk6_0))), inference(spm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_158, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),X2))|~r1_tarski(k2_relat_1(X1),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_153, c_0_112])).
cnf(c_0_159, negated_conjecture, (~r1_tarski(esk6_0,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_156])).
cnf(c_0_160, negated_conjecture, (esk6_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_106]), c_0_79]), c_0_80])])).
cnf(c_0_161, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_159, c_0_160]), c_0_60])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 162
# Proof object clause steps            : 120
# Proof object formula steps           : 42
# Proof object conjectures             : 47
# Proof object clause conjectures      : 44
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 74
# Proof object simplifying inferences  : 72
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 56
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 55
# Processed clauses                    : 3725
# ...of these trivial                  : 13
# ...subsumed                          : 2986
# ...remaining for further processing  : 726
# Other redundant clauses eliminated   : 22
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 39
# Backward-rewritten                   : 127
# Generated clauses                    : 13117
# ...of the previous two non-trivial   : 10422
# Contextual simplify-reflections      : 117
# Paramodulations                      : 13097
# Factorizations                       : 0
# Equation resolutions                 : 22
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 491
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 462
# Current number of unprocessed clauses: 6746
# ...number of literals in the above   : 25662
# Current number of archived formulas  : 0
# Current number of archived clauses   : 220
# Clause-clause subsumption calls (NU) : 74631
# Rec. Clause-clause subsumption calls : 49476
# Non-unit clause-clause subsumptions  : 2633
# Unit Clause-clause subsumption calls : 595
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 14
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 161182

# -------------------------------------------------
# User time                : 0.133 s
# System time              : 0.000 s
# Total time               : 0.133 s
# Maximum resident set size: 28648 pages
