# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t60_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(m1_setfam_1(X2,X1)<=>k5_setfam_1(X1,X2)=X1)), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', t60_setfam_1)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', dt_u1_pre_topc)).
fof(d12_setfam_1, axiom, ![X1, X2]:(m1_setfam_1(X2,X1)<=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', d12_setfam_1)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', l49_zfmisc_1)).
fof(redefinition_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2)), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', redefinition_k5_setfam_1)).
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', t24_yellow_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', d1_pre_topc)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', t14_yellow_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/9de1279ad58540418008bd285fb4a937.E_input', d1_xboole_0)).
fof(c_0_9, plain, ![X15, X16]:((~m1_setfam_1(X16,X15)|k5_setfam_1(X15,X16)=X15|~m1_subset_1(X16,k1_zfmisc_1(k1_zfmisc_1(X15))))&(k5_setfam_1(X15,X16)!=X15|m1_setfam_1(X16,X15)|~m1_subset_1(X16,k1_zfmisc_1(k1_zfmisc_1(X15))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_setfam_1])])])).
fof(c_0_10, plain, ![X14]:(~l1_pre_topc(X14)|m1_subset_1(u1_pre_topc(X14),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X14))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
fof(c_0_11, plain, ![X22, X23]:((~m1_setfam_1(X23,X22)|r1_tarski(X22,k3_tarski(X23)))&(~r1_tarski(X22,k3_tarski(X23))|m1_setfam_1(X23,X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_setfam_1])])).
fof(c_0_12, plain, ![X12, X13]:(~r2_hidden(X12,X13)|r1_tarski(X12,k3_tarski(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
fof(c_0_13, plain, ![X24, X25]:(~m1_subset_1(X25,k1_zfmisc_1(k1_zfmisc_1(X24)))|k5_setfam_1(X24,X25)=k3_tarski(X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
cnf(c_0_14, plain, (k5_setfam_1(X2,X1)=X2|~m1_setfam_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (m1_setfam_1(X2,X1)|~r1_tarski(X1,k3_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
cnf(c_0_19, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (k5_setfam_1(u1_struct_0(X1),u1_pre_topc(X1))=u1_struct_0(X1)|~m1_setfam_1(u1_pre_topc(X1),u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21, plain, (m1_setfam_1(X1,X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_22, plain, ![X5, X6, X7, X8]:(((r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|~v2_pre_topc(X5)|~l1_pre_topc(X5))&((~m1_subset_1(X6,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|(~r1_tarski(X6,u1_pre_topc(X5))|r2_hidden(k5_setfam_1(u1_struct_0(X5),X6),u1_pre_topc(X5)))|~v2_pre_topc(X5)|~l1_pre_topc(X5))&(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5)))|(~r2_hidden(X7,u1_pre_topc(X5))|~r2_hidden(X8,u1_pre_topc(X5))|r2_hidden(k9_subset_1(u1_struct_0(X5),X7,X8),u1_pre_topc(X5))))|~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((m1_subset_1(esk3_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&((m1_subset_1(esk4_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(((r2_hidden(esk3_1(X5),u1_pre_topc(X5))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(r2_hidden(esk4_1(X5),u1_pre_topc(X5))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))&(~r2_hidden(k9_subset_1(u1_struct_0(X5),esk3_1(X5),esk4_1(X5)),u1_pre_topc(X5))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))))&(((m1_subset_1(esk3_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&((m1_subset_1(esk4_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(((r2_hidden(esk3_1(X5),u1_pre_topc(X5))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(r2_hidden(esk4_1(X5),u1_pre_topc(X5))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))&(~r2_hidden(k9_subset_1(u1_struct_0(X5),esk3_1(X5),esk4_1(X5)),u1_pre_topc(X5))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))))&((m1_subset_1(esk3_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&((m1_subset_1(esk4_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(((r2_hidden(esk3_1(X5),u1_pre_topc(X5))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(r2_hidden(esk4_1(X5),u1_pre_topc(X5))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))&(~r2_hidden(k9_subset_1(u1_struct_0(X5),esk3_1(X5),esk4_1(X5)),u1_pre_topc(X5))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_23, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
fof(c_0_24, plain, ![X17]:(v1_xboole_0(X17)|(~r2_hidden(k3_tarski(X17),X17)|k4_yellow_0(k2_yellow_1(X17))=k3_tarski(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
fof(c_0_25, plain, ![X18, X19, X20]:((~v1_xboole_0(X18)|~r2_hidden(X19,X18))&(r2_hidden(esk5_1(X20),X20)|v1_xboole_0(X20))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_26, plain, (k5_setfam_1(u1_struct_0(X1),u1_pre_topc(X1))=k3_tarski(u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_19, c_0_15])).
cnf(c_0_27, plain, (k5_setfam_1(u1_struct_0(X1),u1_pre_topc(X1))=u1_struct_0(X1)|~r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k3_tarski(u1_pre_topc(X1))=u1_struct_0(X1)|~r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_35, plain, (k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(csr,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_30])])).
cnf(c_0_37, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_34])]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 20
# Proof object formula steps           : 19
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 60
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 60
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 16
# ...of the previous two non-trivial   : 12
# Contextual simplify-reflections      : 1
# Paramodulations                      : 16
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 13
# ...number of literals in the above   : 67
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 877
# Rec. Clause-clause subsumption calls : 84
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3126

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.004 s
# Total time               : 0.011 s
# Maximum resident set size: 28676 pages
