# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t149_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', t149_zfmisc_1)).
fof(l3_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2))), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', l3_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', t71_enumset1)).
fof(t75_xboole_1, axiom, ![X1, X2]:~((~(r1_xboole_0(X1,X2))&r1_xboole_0(k3_xboole_0(X1,X2),X2))), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', t75_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', d7_xboole_0)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', l27_zfmisc_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', l51_zfmisc_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', t3_xboole_0)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', t70_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/9de90998caa94b2bbb89806b7bc555f6.E_input', symmetry_r1_xboole_0)).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2))), inference(assume_negation,[status(cth)],[t149_zfmisc_1])).
fof(c_0_13, plain, ![X9, X10]:((~r1_tarski(X9,k1_tarski(X10))|(X9=k1_xboole_0|X9=k1_tarski(X10)))&((X9!=k1_xboole_0|r1_tarski(X9,k1_tarski(X10)))&(X9!=k1_tarski(X10)|r1_tarski(X9,k1_tarski(X10))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_zfmisc_1])])])).
fof(c_0_14, plain, ![X11]:k2_tarski(X11,X11)=k1_tarski(X11), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_15, plain, ![X12, X13]:k1_enumset1(X12,X12,X13)=k2_tarski(X12,X13), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X14, X15, X16]:k2_enumset1(X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_17, negated_conjecture, ((r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))&(r2_hidden(esk4_0,esk3_0)&r1_xboole_0(esk3_0,esk2_0)))&~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_18, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|~r1_tarski(X1,k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, ![X17, X18]:(r1_xboole_0(X17,X18)|~r1_xboole_0(k3_xboole_0(X17,X18),X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t75_xboole_1])])])).
cnf(c_0_24, plain, (X1=k1_xboole_0|X1=k2_enumset1(X2,X2,X2,X2)|~r1_tarski(X1,k2_enumset1(X2,X2,X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_19]), c_0_20]), c_0_20]), c_0_21]), c_0_21])).
cnf(c_0_25, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_19]), c_0_20]), c_0_21])).
fof(c_0_26, plain, ![X19, X20]:((~r1_xboole_0(X19,X20)|k3_xboole_0(X19,X20)=k1_xboole_0)&(k3_xboole_0(X19,X20)!=k1_xboole_0|r1_xboole_0(X19,X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_27, plain, ![X21, X22]:(r2_hidden(X21,X22)|r1_xboole_0(k1_tarski(X21),X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
cnf(c_0_28, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k3_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (k3_xboole_0(esk1_0,esk2_0)=k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0)|k3_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_32, plain, ![X23, X24]:k3_tarski(k2_tarski(X23,X24))=k2_xboole_0(X23,X24), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_33, plain, ![X25, X26, X28, X29, X30]:(((r2_hidden(esk5_2(X25,X26),X25)|r1_xboole_0(X25,X26))&(r2_hidden(esk5_2(X25,X26),X26)|r1_xboole_0(X25,X26)))&(~r2_hidden(X30,X28)|~r2_hidden(X30,X29)|~r1_xboole_0(X28,X29))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_34, negated_conjecture, (k3_xboole_0(esk1_0,esk2_0)=k1_xboole_0|~r1_xboole_0(k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_35, plain, (r2_hidden(X1,X2)|r1_xboole_0(k2_enumset1(X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_19]), c_0_20]), c_0_21])).
fof(c_0_36, plain, ![X31, X32, X33]:((r1_xboole_0(X31,k2_xboole_0(X32,X33))|(~r1_xboole_0(X31,X32)|~r1_xboole_0(X31,X33)))&((r1_xboole_0(X31,X32)|~r1_xboole_0(X31,k2_xboole_0(X32,X33)))&(r1_xboole_0(X31,X33)|~r1_xboole_0(X31,k2_xboole_0(X32,X33))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
cnf(c_0_37, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, negated_conjecture, (r1_xboole_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_40, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_41, negated_conjecture, (k3_xboole_0(esk1_0,esk2_0)=k1_xboole_0|r2_hidden(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
fof(c_0_42, plain, ![X34, X35]:(~r1_xboole_0(X34,X35)|r1_xboole_0(X35,X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_43, plain, (r1_xboole_0(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_37, c_0_20])).
cnf(c_0_45, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, negated_conjecture, (r1_xboole_0(esk1_0,esk2_0)|r2_hidden(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_48, negated_conjecture, (~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_49, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, plain, (r1_xboole_0(X1,k3_tarski(k2_enumset1(X2,X2,X2,X3)))|~r1_xboole_0(X1,X3)|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44]), c_0_21])).
cnf(c_0_51, negated_conjecture, (r1_xboole_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_52, negated_conjecture, (~r1_xboole_0(k3_tarski(k2_enumset1(esk1_0,esk1_0,esk1_0,esk3_0)),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_44]), c_0_21])).
cnf(c_0_53, plain, (r1_xboole_0(k3_tarski(k2_enumset1(X1,X1,X1,X2)),X3)|~r1_xboole_0(X3,X2)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_49, c_0_39])).
cnf(c_0_55, negated_conjecture, (r1_xboole_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_51])).
cnf(c_0_56, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_55])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 32
# Proof object formula steps           : 25
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 18
# Processed clauses                    : 105
# ...of these trivial                  : 2
# ...subsumed                          : 23
# ...remaining for further processing  : 80
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 4
# Generated clauses                    : 131
# ...of the previous two non-trivial   : 109
# Contextual simplify-reflections      : 2
# Paramodulations                      : 128
# Factorizations                       : 1
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 44
# Current number of unprocessed clauses: 34
# ...number of literals in the above   : 100
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 470
# Rec. Clause-clause subsumption calls : 412
# Non-unit clause-clause subsumptions  : 26
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3219

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28648 pages
