# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d16_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k4_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X5,X4)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/9e3eac3f6b8a4aeab254ace418f8e58e.E_input', d16_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/9e3eac3f6b8a4aeab254ace418f8e58e.E_input', t4_subset)).
fof(fraenkel_a_2_1_waybel_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X5,X4)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/9e3eac3f6b8a4aeab254ace418f8e58e.E_input', fraenkel_a_2_1_waybel_0)).
fof(t15_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2))), file('data/MPTP/proofs/9e3eac3f6b8a4aeab254ace418f8e58e.E_input', t15_waybel_0)).
fof(dt_k4_waybel_0, axiom, ![X1, X2]:((l1_orders_2(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k4_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/9e3eac3f6b8a4aeab254ace418f8e58e.E_input', dt_k4_waybel_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/9e3eac3f6b8a4aeab254ace418f8e58e.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/9e3eac3f6b8a4aeab254ace418f8e58e.E_input', d10_xboole_0)).
fof(c_0_7, plain, ![X6, X7, X8, X9, X11, X13]:((((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r2_hidden(X9,X8)|~m1_subset_1(X9,u1_struct_0(X6))|X8!=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((r1_orders_2(X6,esk1_4(X6,X7,X8,X9),X9)|~r2_hidden(X9,X8)|~m1_subset_1(X9,u1_struct_0(X6))|X8!=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&(r2_hidden(esk1_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8)|~m1_subset_1(X9,u1_struct_0(X6))|X8!=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))))&(~m1_subset_1(X11,u1_struct_0(X6))|(~r1_orders_2(X6,X11,X9)|~r2_hidden(X11,X7))|r2_hidden(X9,X8)|~m1_subset_1(X9,u1_struct_0(X6))|X8!=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6)))&((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X6))|X8=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((~r2_hidden(esk2_3(X6,X7,X8),X8)|(~m1_subset_1(X13,u1_struct_0(X6))|(~r1_orders_2(X6,X13,esk2_3(X6,X7,X8))|~r2_hidden(X13,X7)))|X8=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((m1_subset_1(esk3_3(X6,X7,X8),u1_struct_0(X6))|r2_hidden(esk2_3(X6,X7,X8),X8)|X8=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((r1_orders_2(X6,esk3_3(X6,X7,X8),esk2_3(X6,X7,X8))|r2_hidden(esk2_3(X6,X7,X8),X8)|X8=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&(r2_hidden(esk3_3(X6,X7,X8),X7)|r2_hidden(esk2_3(X6,X7,X8),X8)|X8=k4_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_waybel_0])])])])])).
fof(c_0_8, plain, ![X15, X16, X17]:(~r2_hidden(X15,X16)|~m1_subset_1(X16,k1_zfmisc_1(X17))|m1_subset_1(X15,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_9, plain, ![X18, X19, X20, X23, X24]:(((m1_subset_1(esk4_3(X18,X19,X20),u1_struct_0(X19))|~r2_hidden(X18,a_2_1_waybel_0(X19,X20))|(v2_struct_0(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))))&((X18=esk4_3(X18,X19,X20)|~r2_hidden(X18,a_2_1_waybel_0(X19,X20))|(v2_struct_0(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))))&((m1_subset_1(esk5_3(X18,X19,X20),u1_struct_0(X19))|~r2_hidden(X18,a_2_1_waybel_0(X19,X20))|(v2_struct_0(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))))&((r1_orders_2(X19,esk5_3(X18,X19,X20),esk4_3(X18,X19,X20))|~r2_hidden(X18,a_2_1_waybel_0(X19,X20))|(v2_struct_0(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))))&(r2_hidden(esk5_3(X18,X19,X20),X20)|~r2_hidden(X18,a_2_1_waybel_0(X19,X20))|(v2_struct_0(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))))))))&(~m1_subset_1(X23,u1_struct_0(X19))|(X18!=X23|(~m1_subset_1(X24,u1_struct_0(X19))|(~r1_orders_2(X19,X24,X23)|~r2_hidden(X24,X20))))|r2_hidden(X18,a_2_1_waybel_0(X19,X20))|(v2_struct_0(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_waybel_0])])])])])])).
fof(c_0_10, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k4_waybel_0(X1,X2)=a_2_1_waybel_0(X1,X2)))), inference(assume_negation,[status(cth)],[t15_waybel_0])).
fof(c_0_11, plain, ![X27, X28]:(~l1_orders_2(X27)|~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X27)))|m1_subset_1(k4_waybel_0(X27,X28),k1_zfmisc_1(u1_struct_0(X27)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_waybel_0])])).
cnf(c_0_12, plain, (r1_orders_2(X1,esk1_4(X1,X2,X3,X4),X4)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r2_hidden(esk1_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k4_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, plain, (r2_hidden(X3,a_2_1_waybel_0(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r1_orders_2(X2,X4,X1)|~r2_hidden(X4,X5)|~l1_orders_2(X2)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk6_0)&l1_orders_2(esk6_0))&(m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(esk6_0)))&k4_waybel_0(esk6_0,esk7_0)!=a_2_1_waybel_0(esk6_0,esk7_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])).
cnf(c_0_17, plain, (m1_subset_1(k4_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (r1_orders_2(X1,esk1_4(X1,X2,X3,X4),X4)|X3!=k4_waybel_0(X1,X2)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, plain, (r2_hidden(esk1_4(X1,X2,X3,X4),X2)|X3!=k4_waybel_0(X1,X2)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_14, c_0_13])).
cnf(c_0_20, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X1,X3)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k4_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21, plain, (v2_struct_0(X1)|r2_hidden(X2,a_2_1_waybel_0(X1,X3))|~r1_orders_2(X1,X4,X2)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_15, c_0_13])])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk7_0,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (l1_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k4_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_13, c_0_17])).
cnf(c_0_26, plain, (r1_orders_2(X1,esk1_4(X1,X2,k4_waybel_0(X1,X3),X4),X4)|k4_waybel_0(X1,X3)!=k4_waybel_0(X1,X2)|~r2_hidden(X4,k4_waybel_0(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_27, plain, (r2_hidden(esk1_4(X1,X2,k4_waybel_0(X1,X3),X4),X2)|k4_waybel_0(X1,X3)!=k4_waybel_0(X1,X2)|~r2_hidden(X4,k4_waybel_0(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_19, c_0_17])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|X2!=k4_waybel_0(X3,X4)|~r1_orders_2(X3,X5,X1)|~r2_hidden(X5,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_20, c_0_13])).
cnf(c_0_29, plain, (X1=esk4_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_30, plain, ![X29, X30, X31, X32, X33]:((~r1_tarski(X29,X30)|(~r2_hidden(X31,X29)|r2_hidden(X31,X30)))&((r2_hidden(esk8_2(X32,X33),X32)|r1_tarski(X32,X33))&(~r2_hidden(esk8_2(X32,X33),X33)|r1_tarski(X32,X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_31, negated_conjecture, (r2_hidden(X1,a_2_1_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,X2,X1)|~r2_hidden(X2,esk7_0)|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk6_0))|~r2_hidden(X1,k4_waybel_0(esk6_0,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_22]), c_0_23])])).
cnf(c_0_33, negated_conjecture, (r1_orders_2(esk6_0,esk1_4(esk6_0,esk7_0,k4_waybel_0(esk6_0,X1),X2),X2)|k4_waybel_0(esk6_0,X1)!=k4_waybel_0(esk6_0,esk7_0)|~r2_hidden(X2,k4_waybel_0(esk6_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_22]), c_0_23])])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk1_4(esk6_0,esk7_0,k4_waybel_0(esk6_0,X1),X2),esk7_0)|k4_waybel_0(esk6_0,X1)!=k4_waybel_0(esk6_0,esk7_0)|~r2_hidden(X2,k4_waybel_0(esk6_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk6_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_22]), c_0_23])])).
cnf(c_0_35, plain, (r2_hidden(X1,k4_waybel_0(X2,X3))|k4_waybel_0(X2,X3)!=k4_waybel_0(X2,X4)|~r1_orders_2(X2,X5,X1)|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_28, c_0_17])).
cnf(c_0_36, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37, negated_conjecture, (esk4_3(X1,esk6_0,esk7_0)=X1|~r2_hidden(X1,a_2_1_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_38, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (r1_orders_2(X1,esk5_3(X2,X1,X3),esk4_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_waybel_0(X1,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_40, negated_conjecture, (r2_hidden(X1,a_2_1_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,X2,X1)|~r2_hidden(X1,k4_waybel_0(esk6_0,esk7_0))|~r2_hidden(X2,esk7_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk6_0,esk1_4(esk6_0,esk7_0,k4_waybel_0(esk6_0,esk7_0),X1),X1)|~r2_hidden(X1,k4_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_33, c_0_22])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk1_4(esk6_0,esk7_0,k4_waybel_0(esk6_0,esk7_0),X1),esk7_0)|~r2_hidden(X1,k4_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_34, c_0_22])).
cnf(c_0_43, negated_conjecture, (r2_hidden(X1,k4_waybel_0(esk6_0,X2))|k4_waybel_0(esk6_0,X2)!=k4_waybel_0(esk6_0,esk7_0)|~r1_orders_2(esk6_0,X3,X1)|~r2_hidden(X3,esk7_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk6_0)))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_22]), c_0_23])])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk4_3(X1,esk6_0,esk7_0),u1_struct_0(esk6_0))|~r2_hidden(X1,a_2_1_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_45, negated_conjecture, (esk4_3(esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0)=esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1)|r1_tarski(a_2_1_waybel_0(esk6_0,esk7_0),X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(esk6_0,esk5_3(X1,esk6_0,esk7_0),esk4_3(X1,esk6_0,esk7_0))|~r2_hidden(X1,a_2_1_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_47, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_48, plain, (r1_tarski(X1,X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_49, negated_conjecture, (r2_hidden(X1,a_2_1_waybel_0(esk6_0,esk7_0))|~r2_hidden(X1,k4_waybel_0(esk6_0,esk7_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_50, negated_conjecture, (r2_hidden(X1,k4_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,X2,X1)|~r2_hidden(X2,esk7_0)|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(spm,[status(thm)],[c_0_43, c_0_22])).
cnf(c_0_51, negated_conjecture, (r1_tarski(a_2_1_waybel_0(esk6_0,esk7_0),X1)|m1_subset_1(esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1),u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_38])).
cnf(c_0_52, negated_conjecture, (r1_tarski(a_2_1_waybel_0(esk6_0,esk7_0),X1)|r1_orders_2(esk6_0,esk5_3(esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0),esk4_3(esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_46, c_0_38])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk5_3(X1,esk6_0,esk7_0),esk7_0)|~r2_hidden(X1,a_2_1_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_22]), c_0_23])]), c_0_24])).
fof(c_0_54, plain, ![X35, X36]:(((r1_tarski(X35,X36)|X35!=X36)&(r1_tarski(X36,X35)|X35!=X36))&(~r1_tarski(X35,X36)|~r1_tarski(X36,X35)|X35=X36)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_55, negated_conjecture, (r1_tarski(X1,a_2_1_waybel_0(esk6_0,esk7_0))|~r2_hidden(esk8_2(X1,a_2_1_waybel_0(esk6_0,esk7_0)),k4_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_56, negated_conjecture, (r1_tarski(a_2_1_waybel_0(esk6_0,esk7_0),X1)|r2_hidden(esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1),k4_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,X2,esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1))|~r2_hidden(X2,esk7_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_57, negated_conjecture, (r1_tarski(a_2_1_waybel_0(esk6_0,esk7_0),X1)|r1_orders_2(esk6_0,esk5_3(esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0),esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1))), inference(spm,[status(thm)],[c_0_52, c_0_45])).
cnf(c_0_58, negated_conjecture, (r1_tarski(a_2_1_waybel_0(esk6_0,esk7_0),X1)|r2_hidden(esk5_3(esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1),esk6_0,esk7_0),esk7_0)), inference(spm,[status(thm)],[c_0_53, c_0_38])).
cnf(c_0_59, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_60, negated_conjecture, (r1_tarski(k4_waybel_0(esk6_0,esk7_0),a_2_1_waybel_0(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_55, c_0_38])).
cnf(c_0_61, negated_conjecture, (k4_waybel_0(esk6_0,esk7_0)!=a_2_1_waybel_0(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_62, negated_conjecture, (r1_tarski(a_2_1_waybel_0(esk6_0,esk7_0),X1)|r2_hidden(esk8_2(a_2_1_waybel_0(esk6_0,esk7_0),X1),k4_waybel_0(esk6_0,esk7_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])).
cnf(c_0_63, negated_conjecture, (~r1_tarski(a_2_1_waybel_0(esk6_0,esk7_0),k4_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
cnf(c_0_64, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_62]), c_0_63]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 50
# Proof object formula steps           : 15
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 311
# ...of these trivial                  : 3
# ...subsumed                          : 78
# ...remaining for further processing  : 230
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 0
# Generated clauses                    : 833
# ...of the previous two non-trivial   : 788
# Contextual simplify-reflections      : 14
# Paramodulations                      : 830
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 195
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 188
# Current number of unprocessed clauses: 521
# ...number of literals in the above   : 2774
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 8809
# Rec. Clause-clause subsumption calls : 2962
# Non-unit clause-clause subsumptions  : 98
# Unit Clause-clause subsumption calls : 76
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 28786

# -------------------------------------------------
# User time                : 0.023 s
# System time              : 0.003 s
# Total time               : 0.026 s
# Maximum resident set size: 28876 pages
