# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_SE_CS_SP_PS_S4d
# and selection function SelectCQIPrecWNTNp.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t32_lattices, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v11_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((k4_lattices(X1,X2,X3)=k4_lattices(X1,X2,X4)&k3_lattices(X1,X2,X3)=k3_lattices(X1,X2,X4))=>X3=X4))))), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', t32_lattices)).
fof(dt_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', dt_k3_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', cc1_lattices)).
fof(redefinition_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k3_lattices(X1,X2,X3)=k1_lattices(X1,X2,X3)), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', redefinition_k3_lattices)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', dt_l3_lattices)).
fof(commutativity_k3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', commutativity_k3_lattices)).
fof(t17_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k1_lattices(X1,X2,X2)=X2)), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', t17_lattices)).
fof(d9_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>(v9_lattices(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2)))), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', d9_lattices)).
fof(commutativity_k4_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&l1_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', commutativity_k4_lattices)).
fof(redefinition_k4_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&l1_lattices(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k4_lattices(X1,X2,X3)=k2_lattices(X1,X2,X3)), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', redefinition_k4_lattices)).
fof(t31_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v11_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>k3_lattices(X1,X2,k4_lattices(X1,X3,X4))=k4_lattices(X1,k3_lattices(X1,X2,X3),k3_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/9e5197c9186c4513abd33f33d14ae62a.E_input', t31_lattices)).
fof(c_0_11, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v11_lattices(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((k4_lattices(X1,X2,X3)=k4_lattices(X1,X2,X4)&k3_lattices(X1,X2,X3)=k3_lattices(X1,X2,X4))=>X3=X4)))))), inference(assume_negation,[status(cth)],[t32_lattices])).
fof(c_0_12, plain, ![X9, X10, X11]:(v2_struct_0(X9)|(~v4_lattices(X9)|~l2_lattices(X9))|(~m1_subset_1(X10,u1_struct_0(X9))|~m1_subset_1(X11,u1_struct_0(X9)))|m1_subset_1(k3_lattices(X9,X10,X11),u1_struct_0(X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_lattices])])])).
fof(c_0_13, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v11_lattices(esk1_0)&l3_lattices(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&((k4_lattices(esk1_0,esk2_0,esk3_0)=k4_lattices(esk1_0,esk2_0,esk4_0)&k3_lattices(esk1_0,esk2_0,esk3_0)=k3_lattices(esk1_0,esk2_0,esk4_0))&esk3_0!=esk4_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
cnf(c_0_14, plain, (v2_struct_0(X1)|m1_subset_1(k3_lattices(X1,X2,X3),u1_struct_0(X1))|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, ![X12]:((~v2_struct_0(X12)|(v2_struct_0(X12)|~v10_lattices(X12))|~l3_lattices(X12))&((v4_lattices(X12)|(v2_struct_0(X12)|~v10_lattices(X12))|~l3_lattices(X12))&((v5_lattices(X12)|(v2_struct_0(X12)|~v10_lattices(X12))|~l3_lattices(X12))&((v6_lattices(X12)|(v2_struct_0(X12)|~v10_lattices(X12))|~l3_lattices(X12))&((v7_lattices(X12)|(v2_struct_0(X12)|~v10_lattices(X12))|~l3_lattices(X12))&((v8_lattices(X12)|(v2_struct_0(X12)|~v10_lattices(X12))|~l3_lattices(X12))&(v9_lattices(X12)|(v2_struct_0(X12)|~v10_lattices(X12))|~l3_lattices(X12)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_18, plain, ![X13, X14, X15]:(v2_struct_0(X13)|(~v4_lattices(X13)|~l2_lattices(X13))|(~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13)))|k3_lattices(X13,X14,X15)=k1_lattices(X13,X14,X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_lattices])])])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(k3_lattices(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_20, plain, (v4_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_23, plain, ![X16]:((l1_lattices(X16)|~l3_lattices(X16))&(l2_lattices(X16)|~l3_lattices(X16))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
cnf(c_0_24, plain, (v2_struct_0(X1)|k3_lattices(X1,X2,X3)=k1_lattices(X1,X2,X3)|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_26, plain, ![X17, X18, X19]:(v2_struct_0(X17)|(~v4_lattices(X17)|~l2_lattices(X17))|(~m1_subset_1(X18,u1_struct_0(X17))|~m1_subset_1(X19,u1_struct_0(X17)))|k3_lattices(X17,X18,X19)=k3_lattices(X17,X19,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k3_lattices])])])).
fof(c_0_27, plain, ![X20, X21]:(v2_struct_0(X20)|(~v6_lattices(X20)|(~v8_lattices(X20)|(~v9_lattices(X20)|~l3_lattices(X20))))|(~m1_subset_1(X21,u1_struct_0(X20))|k1_lattices(X20,X21,X21)=X21)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_lattices])])])])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(k3_lattices(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|~l2_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_29, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, ![X22, X23, X24]:((~v9_lattices(X22)|(~m1_subset_1(X23,u1_struct_0(X22))|(~m1_subset_1(X24,u1_struct_0(X22))|k2_lattices(X22,X23,k1_lattices(X22,X23,X24))=X23))|(v2_struct_0(X22)|~l3_lattices(X22)))&((m1_subset_1(esk5_1(X22),u1_struct_0(X22))|v9_lattices(X22)|(v2_struct_0(X22)|~l3_lattices(X22)))&((m1_subset_1(esk6_1(X22),u1_struct_0(X22))|v9_lattices(X22)|(v2_struct_0(X22)|~l3_lattices(X22)))&(k2_lattices(X22,esk5_1(X22),k1_lattices(X22,esk5_1(X22),esk6_1(X22)))!=esk5_1(X22)|v9_lattices(X22)|(v2_struct_0(X22)|~l3_lattices(X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d9_lattices])])])])])])).
cnf(c_0_31, negated_conjecture, (k1_lattices(esk1_0,X1,esk2_0)=k3_lattices(esk1_0,X1,esk2_0)|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_16])).
cnf(c_0_32, plain, (v2_struct_0(X1)|k3_lattices(X1,X2,X3)=k3_lattices(X1,X3,X2)|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_33, plain, ![X27, X28, X29]:(v2_struct_0(X27)|(~v6_lattices(X27)|~l1_lattices(X27))|(~m1_subset_1(X28,u1_struct_0(X27))|~m1_subset_1(X29,u1_struct_0(X27)))|k4_lattices(X27,X28,X29)=k4_lattices(X27,X29,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k4_lattices])])])).
cnf(c_0_34, plain, (v2_struct_0(X1)|k1_lattices(X1,X2,X2)=X2|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_35, plain, ![X30, X31, X32]:(v2_struct_0(X30)|(~v6_lattices(X30)|~l1_lattices(X30))|(~m1_subset_1(X31,u1_struct_0(X30))|~m1_subset_1(X32,u1_struct_0(X30)))|k4_lattices(X30,X31,X32)=k2_lattices(X30,X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k4_lattices])])])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(k3_lattices(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_21])])).
cnf(c_0_37, negated_conjecture, (k3_lattices(esk1_0,esk2_0,esk3_0)=k3_lattices(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_38, plain, (k2_lattices(X1,X2,k1_lattices(X1,X2,X3))=X2|v2_struct_0(X1)|~v9_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, negated_conjecture, (k1_lattices(esk1_0,X1,esk2_0)=k3_lattices(esk1_0,X1,esk2_0)|~l2_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_20]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_40, negated_conjecture, (k3_lattices(esk1_0,X1,esk4_0)=k3_lattices(esk1_0,esk4_0,X1)|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_15]), c_0_16])).
cnf(c_0_41, plain, (v2_struct_0(X1)|k4_lattices(X1,X2,X3)=k4_lattices(X1,X3,X2)|~v6_lattices(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_43, negated_conjecture, (k1_lattices(esk1_0,esk4_0,esk4_0)=esk4_0|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_15]), c_0_21])]), c_0_16])).
cnf(c_0_44, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_45, plain, (v2_struct_0(X1)|k4_lattices(X1,X2,X3)=k2_lattices(X1,X2,X3)|~v6_lattices(X1)|~l1_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(k3_lattices(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_25]), c_0_37])).
cnf(c_0_47, negated_conjecture, (k2_lattices(esk1_0,X1,k1_lattices(esk1_0,X1,esk2_0))=X1|~v9_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_25]), c_0_21])]), c_0_16])).
cnf(c_0_48, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_49, negated_conjecture, (k1_lattices(esk1_0,X1,esk2_0)=k3_lattices(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_29]), c_0_21])])).
cnf(c_0_50, negated_conjecture, (k3_lattices(esk1_0,X1,esk4_0)=k3_lattices(esk1_0,esk4_0,X1)|~l2_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_20]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_51, negated_conjecture, (k4_lattices(esk1_0,X1,esk4_0)=k4_lattices(esk1_0,esk4_0,X1)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_15]), c_0_16])).
cnf(c_0_52, plain, (l1_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_53, negated_conjecture, (k3_lattices(esk1_0,X1,esk3_0)=k3_lattices(esk1_0,esk3_0,X1)|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_42]), c_0_16])).
fof(c_0_54, plain, ![X33, X34, X35, X36]:(v2_struct_0(X33)|(~v10_lattices(X33)|(~v11_lattices(X33)|~l3_lattices(X33)))|(~m1_subset_1(X34,u1_struct_0(X33))|(~m1_subset_1(X35,u1_struct_0(X33))|(~m1_subset_1(X36,u1_struct_0(X33))|k3_lattices(X33,X34,k4_lattices(X33,X35,X36))=k4_lattices(X33,k3_lattices(X33,X34,X35),k3_lattices(X33,X34,X36)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_lattices])])])])).
cnf(c_0_55, negated_conjecture, (k1_lattices(esk1_0,X1,esk4_0)=k3_lattices(esk1_0,X1,esk4_0)|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_15]), c_0_16])).
cnf(c_0_56, negated_conjecture, (k1_lattices(esk1_0,esk4_0,esk4_0)=esk4_0|~v9_lattices(esk1_0)|~v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_57, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_58, negated_conjecture, (k2_lattices(esk1_0,X1,k3_lattices(esk1_0,esk2_0,esk3_0))=k4_lattices(esk1_0,X1,k3_lattices(esk1_0,esk2_0,esk3_0))|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_16])).
cnf(c_0_59, negated_conjecture, (k2_lattices(esk1_0,X1,k1_lattices(esk1_0,X1,esk2_0))=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_60, negated_conjecture, (k1_lattices(esk1_0,esk4_0,esk2_0)=k3_lattices(esk1_0,esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_49, c_0_15])).
cnf(c_0_61, negated_conjecture, (k3_lattices(esk1_0,X1,esk4_0)=k3_lattices(esk1_0,esk4_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_29]), c_0_21])])).
cnf(c_0_62, negated_conjecture, (k4_lattices(esk1_0,X1,esk4_0)=k4_lattices(esk1_0,esk4_0,X1)|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_21])])).
cnf(c_0_63, negated_conjecture, (k1_lattices(esk1_0,esk3_0,esk3_0)=esk3_0|~v9_lattices(esk1_0)|~v8_lattices(esk1_0)|~v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_42]), c_0_21])]), c_0_16])).
cnf(c_0_64, negated_conjecture, (k3_lattices(esk1_0,X1,esk3_0)=k3_lattices(esk1_0,esk3_0,X1)|~l2_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_20]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_65, plain, (v2_struct_0(X1)|k3_lattices(X1,X2,k4_lattices(X1,X3,X4))=k4_lattices(X1,k3_lattices(X1,X2,X3),k3_lattices(X1,X2,X4))|~v10_lattices(X1)|~v11_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, negated_conjecture, (v11_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_67, negated_conjecture, (k1_lattices(esk1_0,X1,esk4_0)=k3_lattices(esk1_0,X1,esk4_0)|~l2_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_20]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_68, negated_conjecture, (k1_lattices(esk1_0,esk4_0,esk4_0)=esk4_0|~v9_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_69, negated_conjecture, (k2_lattices(esk1_0,X1,k3_lattices(esk1_0,esk2_0,esk3_0))=k4_lattices(esk1_0,X1,k3_lattices(esk1_0,esk2_0,esk3_0))|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_52]), c_0_21])])).
cnf(c_0_70, negated_conjecture, (k2_lattices(esk1_0,esk4_0,k3_lattices(esk1_0,esk4_0,esk2_0))=esk4_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_15]), c_0_60])).
cnf(c_0_71, negated_conjecture, (k3_lattices(esk1_0,esk4_0,esk2_0)=k3_lattices(esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_25]), c_0_37])).
cnf(c_0_72, negated_conjecture, (k4_lattices(esk1_0,X1,esk4_0)=k4_lattices(esk1_0,esk4_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_57]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_73, negated_conjecture, (k1_lattices(esk1_0,X1,esk3_0)=k3_lattices(esk1_0,X1,esk3_0)|~l2_lattices(esk1_0)|~v4_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_42]), c_0_16])).
cnf(c_0_74, negated_conjecture, (k1_lattices(esk1_0,esk3_0,esk3_0)=esk3_0|~v9_lattices(esk1_0)|~v6_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_44]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_75, negated_conjecture, (k4_lattices(esk1_0,X1,esk3_0)=k4_lattices(esk1_0,esk3_0,X1)|~l1_lattices(esk1_0)|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_16])).
cnf(c_0_76, negated_conjecture, (k1_lattices(esk1_0,esk3_0,esk2_0)=k3_lattices(esk1_0,esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_49, c_0_42])).
cnf(c_0_77, negated_conjecture, (k3_lattices(esk1_0,X1,esk3_0)=k3_lattices(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_29]), c_0_21])])).
cnf(c_0_78, negated_conjecture, (k4_lattices(esk1_0,k3_lattices(esk1_0,X1,X2),k3_lattices(esk1_0,X1,esk4_0))=k3_lattices(esk1_0,X1,k4_lattices(esk1_0,X2,esk4_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_15]), c_0_21]), c_0_66]), c_0_22])]), c_0_16])).
cnf(c_0_79, negated_conjecture, (k4_lattices(esk1_0,esk2_0,esk3_0)=k4_lattices(esk1_0,esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_80, negated_conjecture, (k1_lattices(esk1_0,X1,esk4_0)=k3_lattices(esk1_0,X1,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_29]), c_0_21])])).
cnf(c_0_81, negated_conjecture, (k1_lattices(esk1_0,esk4_0,esk4_0)=esk4_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_48]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_82, negated_conjecture, (k2_lattices(esk1_0,X1,k3_lattices(esk1_0,esk2_0,esk3_0))=k4_lattices(esk1_0,X1,k3_lattices(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_57]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_83, negated_conjecture, (k2_lattices(esk1_0,esk4_0,k3_lattices(esk1_0,esk2_0,esk3_0))=esk4_0), inference(rw,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_84, negated_conjecture, (k4_lattices(esk1_0,esk4_0,k3_lattices(esk1_0,esk2_0,esk3_0))=k4_lattices(esk1_0,k3_lattices(esk1_0,esk2_0,esk3_0),esk4_0)), inference(spm,[status(thm)],[c_0_72, c_0_46])).
cnf(c_0_85, negated_conjecture, (k1_lattices(esk1_0,X1,esk3_0)=k3_lattices(esk1_0,X1,esk3_0)|~l2_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_20]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_86, negated_conjecture, (k1_lattices(esk1_0,esk3_0,esk3_0)=esk3_0|~v9_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_57]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_87, negated_conjecture, (k4_lattices(esk1_0,X1,esk3_0)=k4_lattices(esk1_0,esk3_0,X1)|~v6_lattices(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_52]), c_0_21])])).
cnf(c_0_88, negated_conjecture, (k2_lattices(esk1_0,esk3_0,k3_lattices(esk1_0,esk3_0,esk2_0))=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_42]), c_0_76])).
cnf(c_0_89, negated_conjecture, (k3_lattices(esk1_0,esk3_0,esk2_0)=k3_lattices(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_77, c_0_25])).
cnf(c_0_90, negated_conjecture, (k4_lattices(esk1_0,k3_lattices(esk1_0,X1,X2),k3_lattices(esk1_0,X1,esk3_0))=k3_lattices(esk1_0,X1,k4_lattices(esk1_0,X2,esk3_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_42]), c_0_21]), c_0_66]), c_0_22])]), c_0_16])).
cnf(c_0_91, negated_conjecture, (k4_lattices(esk1_0,k3_lattices(esk1_0,X1,esk2_0),k3_lattices(esk1_0,X1,esk4_0))=k3_lattices(esk1_0,X1,k4_lattices(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_25]), c_0_79])).
cnf(c_0_92, negated_conjecture, (k3_lattices(esk1_0,esk4_0,esk4_0)=esk4_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_15]), c_0_81])).
cnf(c_0_93, negated_conjecture, (k4_lattices(esk1_0,k3_lattices(esk1_0,esk2_0,esk3_0),esk4_0)=esk4_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_15]), c_0_83]), c_0_84])).
cnf(c_0_94, negated_conjecture, (k1_lattices(esk1_0,X1,esk3_0)=k3_lattices(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_29]), c_0_21])])).
cnf(c_0_95, negated_conjecture, (k1_lattices(esk1_0,esk3_0,esk3_0)=esk3_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_48]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_96, negated_conjecture, (k4_lattices(esk1_0,X1,esk3_0)=k4_lattices(esk1_0,esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_57]), c_0_21]), c_0_22])]), c_0_16])).
cnf(c_0_97, negated_conjecture, (k2_lattices(esk1_0,esk3_0,k3_lattices(esk1_0,esk2_0,esk3_0))=esk3_0), inference(rw,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_98, negated_conjecture, (k4_lattices(esk1_0,k3_lattices(esk1_0,X1,esk2_0),k3_lattices(esk1_0,X1,esk3_0))=k3_lattices(esk1_0,X1,k4_lattices(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_90, c_0_25])).
cnf(c_0_99, negated_conjecture, (k3_lattices(esk1_0,esk4_0,esk3_0)=k3_lattices(esk1_0,esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_61, c_0_42])).
cnf(c_0_100, negated_conjecture, (k4_lattices(esk1_0,k3_lattices(esk1_0,esk2_0,esk3_0),k3_lattices(esk1_0,esk3_0,esk4_0))=k3_lattices(esk1_0,esk3_0,k4_lattices(esk1_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_42]), c_0_89])).
cnf(c_0_101, negated_conjecture, (k3_lattices(esk1_0,esk4_0,k4_lattices(esk1_0,esk2_0,esk3_0))=esk4_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_15]), c_0_71]), c_0_92]), c_0_93])).
cnf(c_0_102, negated_conjecture, (k3_lattices(esk1_0,esk3_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_42]), c_0_95])).
cnf(c_0_103, negated_conjecture, (k4_lattices(esk1_0,k3_lattices(esk1_0,esk2_0,esk3_0),esk3_0)=k4_lattices(esk1_0,esk3_0,k3_lattices(esk1_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_96, c_0_46])).
cnf(c_0_104, negated_conjecture, (k4_lattices(esk1_0,esk3_0,k3_lattices(esk1_0,esk2_0,esk3_0))=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_42]), c_0_97])).
cnf(c_0_105, negated_conjecture, (esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_106, negated_conjecture, (esk4_0=k3_lattices(esk1_0,esk3_0,k4_lattices(esk1_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_15]), c_0_71]), c_0_99]), c_0_100]), c_0_101])).
cnf(c_0_107, negated_conjecture, (k3_lattices(esk1_0,esk3_0,k4_lattices(esk1_0,esk2_0,esk3_0))=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_42]), c_0_89]), c_0_102]), c_0_103]), c_0_104])).
cnf(c_0_108, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_106]), c_0_107])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 109
# Proof object clause steps            : 86
# Proof object formula steps           : 23
# Proof object conjectures             : 75
# Proof object clause conjectures      : 72
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 59
# Proof object simplifying inferences  : 137
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 817
# ...of these trivial                  : 21
# ...subsumed                          : 18
# ...remaining for further processing  : 777
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 128
# Backward-rewritten                   : 411
# Generated clauses                    : 17034
# ...of the previous two non-trivial   : 17240
# Contextual simplify-reflections      : 10
# Paramodulations                      : 17034
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 209
#    Positive orientable unit clauses  : 70
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 138
# Current number of unprocessed clauses: 16173
# ...number of literals in the above   : 30837
# Current number of archived formulas  : 0
# Current number of archived clauses   : 568
# Clause-clause subsumption calls (NU) : 3509
# Rec. Clause-clause subsumption calls : 1740
# Non-unit clause-clause subsumptions  : 156
# Unit Clause-clause subsumption calls : 44
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11750
# BW rewrite match successes           : 53
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 591276

# -------------------------------------------------
# User time                : 0.135 s
# System time              : 0.013 s
# Total time               : 0.148 s
# Maximum resident set size: 28680 pages
