# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/9e66154314cb485799885c44c6d50569.E_input', d5_funct_1)).
fof(t21_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(k5_relat_1(X3,X2)))<=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X1),k1_relat_1(X2)))))), file('data/MPTP/proofs/9e66154314cb485799885c44c6d50569.E_input', t21_funct_1)).
fof(t27_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(k1_relat_1(k5_relat_1(X2,X1))=k1_relat_1(X2)=>r1_tarski(k2_relat_1(X2),k1_relat_1(X1))))), file('data/MPTP/proofs/9e66154314cb485799885c44c6d50569.E_input', t27_funct_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/9e66154314cb485799885c44c6d50569.E_input', d3_tarski)).
fof(c_0_4, plain, ![X10, X11, X12, X14, X15, X16, X18]:((((r2_hidden(esk3_3(X10,X11,X12),k1_relat_1(X10))|~r2_hidden(X12,X11)|X11!=k2_relat_1(X10)|(~v1_relat_1(X10)|~v1_funct_1(X10)))&(X12=k1_funct_1(X10,esk3_3(X10,X11,X12))|~r2_hidden(X12,X11)|X11!=k2_relat_1(X10)|(~v1_relat_1(X10)|~v1_funct_1(X10))))&(~r2_hidden(X15,k1_relat_1(X10))|X14!=k1_funct_1(X10,X15)|r2_hidden(X14,X11)|X11!=k2_relat_1(X10)|(~v1_relat_1(X10)|~v1_funct_1(X10))))&((~r2_hidden(esk4_2(X10,X16),X16)|(~r2_hidden(X18,k1_relat_1(X10))|esk4_2(X10,X16)!=k1_funct_1(X10,X18))|X16=k2_relat_1(X10)|(~v1_relat_1(X10)|~v1_funct_1(X10)))&((r2_hidden(esk5_2(X10,X16),k1_relat_1(X10))|r2_hidden(esk4_2(X10,X16),X16)|X16=k2_relat_1(X10)|(~v1_relat_1(X10)|~v1_funct_1(X10)))&(esk4_2(X10,X16)=k1_funct_1(X10,esk5_2(X10,X16))|r2_hidden(esk4_2(X10,X16),X16)|X16=k2_relat_1(X10)|(~v1_relat_1(X10)|~v1_funct_1(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
fof(c_0_5, plain, ![X7, X8, X9]:(((r2_hidden(X7,k1_relat_1(X9))|~r2_hidden(X7,k1_relat_1(k5_relat_1(X9,X8)))|(~v1_relat_1(X9)|~v1_funct_1(X9))|(~v1_relat_1(X8)|~v1_funct_1(X8)))&(r2_hidden(k1_funct_1(X9,X7),k1_relat_1(X8))|~r2_hidden(X7,k1_relat_1(k5_relat_1(X9,X8)))|(~v1_relat_1(X9)|~v1_funct_1(X9))|(~v1_relat_1(X8)|~v1_funct_1(X8))))&(~r2_hidden(X7,k1_relat_1(X9))|~r2_hidden(k1_funct_1(X9,X7),k1_relat_1(X8))|r2_hidden(X7,k1_relat_1(k5_relat_1(X9,X8)))|(~v1_relat_1(X9)|~v1_funct_1(X9))|(~v1_relat_1(X8)|~v1_funct_1(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_funct_1])])])])).
cnf(c_0_6, plain, (X1=k1_funct_1(X2,esk3_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(k1_relat_1(k5_relat_1(X2,X1))=k1_relat_1(X2)=>r1_tarski(k2_relat_1(X2),k1_relat_1(X1)))))), inference(assume_negation,[status(cth)],[t27_funct_1])).
cnf(c_0_8, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relat_1(X3))|~r2_hidden(X2,k1_relat_1(k5_relat_1(X1,X3)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (k1_funct_1(X1,esk3_3(X1,k2_relat_1(X1),X2))=X2|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, ((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&(k1_relat_1(k5_relat_1(esk2_0,esk1_0))=k1_relat_1(esk2_0)&~r1_tarski(k2_relat_1(esk2_0),k1_relat_1(esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_11, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(esk3_3(X3,k2_relat_1(X3),X1),k1_relat_1(k5_relat_1(X3,X2)))|~r2_hidden(X1,k2_relat_1(X3))|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, negated_conjecture, (k1_relat_1(k5_relat_1(esk2_0,esk1_0))=k1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (r2_hidden(esk3_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_18, plain, ![X20, X21, X22, X23, X24]:((~r1_tarski(X20,X21)|(~r2_hidden(X22,X20)|r2_hidden(X22,X21)))&((r2_hidden(esk6_2(X23,X24),X23)|r1_tarski(X23,X24))&(~r2_hidden(esk6_2(X23,X24),X24)|r1_tarski(X23,X24)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_19, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk1_0))|~r2_hidden(esk3_3(esk2_0,k2_relat_1(esk2_0),X1),k1_relat_1(esk2_0))|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_20, plain, (r2_hidden(esk3_3(X1,k2_relat_1(X1),X2),k1_relat_1(X1))|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk1_0))|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_14]), c_0_16])])).
cnf(c_0_23, negated_conjecture, (r1_tarski(X1,k1_relat_1(esk1_0))|~r2_hidden(esk6_2(X1,k1_relat_1(esk1_0)),k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_24, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, negated_conjecture, (~r1_tarski(k2_relat_1(esk2_0),k1_relat_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_26, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 27
# Proof object clause steps            : 18
# Proof object formula steps           : 9
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 35
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 33
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 46
# ...of the previous two non-trivial   : 40
# Contextual simplify-reflections      : 0
# Paramodulations                      : 43
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 23
# ...number of literals in the above   : 149
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 209
# Rec. Clause-clause subsumption calls : 52
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2591

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28744 pages
