# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', t106_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', t100_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', t48_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', t28_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', commutativity_k3_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', t36_xboole_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', d1_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', fc1_subset_1)).
fof(t31_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,X3)<=>r1_tarski(k3_subset_1(X1,X3),k3_subset_1(X1,X2))))), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', t31_subset_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', d5_subset_1)).
fof(t86_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X1,X3))=>r1_tarski(X1,k4_xboole_0(X2,X3))), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', t86_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/9ed8acef75a746429bb659cf8278dcfb.E_input', symmetry_r1_xboole_0)).
fof(c_0_13, plain, ![X7, X8, X9]:((r1_tarski(X7,X8)|~r1_tarski(X7,k4_xboole_0(X8,X9)))&(r1_xboole_0(X7,X9)|~r1_tarski(X7,k4_xboole_0(X8,X9)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
fof(c_0_14, plain, ![X13, X14]:k4_xboole_0(X13,X14)=k5_xboole_0(X13,k3_xboole_0(X13,X14)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_15, plain, ![X35, X36]:k4_xboole_0(X35,k4_xboole_0(X35,X36))=k3_xboole_0(X35,X36), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_16, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_19, plain, ![X31, X32]:(~r1_tarski(X31,X32)|k3_xboole_0(X31,X32)=X31), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_20, plain, ![X22, X23]:k3_xboole_0(X22,X23)=k3_xboole_0(X23,X22), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_21, plain, ![X11, X12]:r1_tarski(k4_xboole_0(X11,X12),X11), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_17]), c_0_17])).
cnf(c_0_24, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_27, plain, ![X24, X25, X26, X27, X28, X29]:(((~r2_hidden(X26,X25)|r1_tarski(X26,X24)|X25!=k1_zfmisc_1(X24))&(~r1_tarski(X27,X24)|r2_hidden(X27,X25)|X25!=k1_zfmisc_1(X24)))&((~r2_hidden(esk4_2(X28,X29),X29)|~r1_tarski(esk4_2(X28,X29),X28)|X29=k1_zfmisc_1(X28))&(r2_hidden(esk4_2(X28,X29),X29)|r1_tarski(esk4_2(X28,X29),X28)|X29=k1_zfmisc_1(X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29, plain, (k3_xboole_0(X1,X2)=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_26, c_0_17])).
cnf(c_0_31, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_32, plain, ![X33, X34]:(((~m1_subset_1(X34,X33)|r2_hidden(X34,X33)|v1_xboole_0(X33))&(~r2_hidden(X34,X33)|m1_subset_1(X34,X33)|v1_xboole_0(X33)))&((~m1_subset_1(X34,X33)|v1_xboole_0(X34)|~v1_xboole_0(X33))&(~v1_xboole_0(X34)|m1_subset_1(X34,X33)|~v1_xboole_0(X33)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_33, plain, ![X10]:~v1_xboole_0(k1_zfmisc_1(X10)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_34, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(r1_tarski(X2,X3)<=>r1_tarski(k3_subset_1(X1,X3),k3_subset_1(X1,X2)))))), inference(assume_negation,[status(cth)],[t31_subset_1])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_30, c_0_23])).
cnf(c_0_37, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_40, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&((~r1_tarski(esk2_0,esk3_0)|~r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0)))&(r1_tarski(esk2_0,esk3_0)|r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])).
cnf(c_0_41, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_44, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_tarski(X4,X3)|~r1_tarski(X1,X4)), inference(spm,[status(thm)],[c_0_35, c_0_41])).
cnf(c_0_45, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_46, negated_conjecture, (r1_tarski(k3_xboole_0(X1,X2),esk1_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_47, plain, (r1_tarski(X1,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_36, c_0_24])).
fof(c_0_48, plain, ![X15, X16]:(~m1_subset_1(X16,k1_zfmisc_1(X15))|k3_subset_1(X15,X16)=k4_xboole_0(X15,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_49, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_52, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X2,esk3_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_46, c_0_29])).
cnf(c_0_53, negated_conjecture, (r1_tarski(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_47, c_0_45])).
cnf(c_0_54, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2)))=k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_23, c_0_23])).
cnf(c_0_56, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_57, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_49])).
fof(c_0_58, plain, ![X17, X18, X19]:(~r1_tarski(X17,X18)|~r1_xboole_0(X17,X19)|r1_tarski(X17,k4_xboole_0(X18,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_xboole_1])])).
cnf(c_0_59, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_50])).
cnf(c_0_60, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k5_xboole_0(X3,k3_xboole_0(X3,X2)))), inference(rw,[status(thm)],[c_0_51, c_0_17])).
cnf(c_0_61, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_62, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_54, c_0_17])).
cnf(c_0_63, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,k5_xboole_0(X1,X2))|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_55, c_0_29])).
cnf(c_0_64, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_39])).
cnf(c_0_65, plain, (r1_tarski(X1,k4_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_66, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,X1)))=X1|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_67, negated_conjecture, (r1_tarski(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_47, c_0_59])).
cnf(c_0_68, plain, (r1_tarski(k5_xboole_0(X1,X1),X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_30, c_0_24])).
cnf(c_0_69, plain, (r1_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r1_tarski(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_60, c_0_23])).
cnf(c_0_70, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,X2)))=X2|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_23, c_0_29])).
cnf(c_0_71, negated_conjecture, (r1_tarski(k3_xboole_0(esk3_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_61, c_0_36])).
cnf(c_0_72, plain, (k3_xboole_0(X1,k5_xboole_0(X1,X2))=k3_subset_1(X1,X2)|~r1_tarski(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])).
cnf(c_0_73, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_29])).
cnf(c_0_74, plain, (r1_tarski(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)), inference(rw,[status(thm)],[c_0_65, c_0_17])).
cnf(c_0_75, plain, (k5_xboole_0(X1,X1)=k3_xboole_0(X1,X2)|~r1_tarski(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_76, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,k5_xboole_0(esk2_0,esk2_0)))=esk2_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_77, negated_conjecture, (r1_tarski(k5_xboole_0(esk2_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_68, c_0_59])).
cnf(c_0_78, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k3_xboole_0(X3,k5_xboole_0(X3,X2)))|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_79, negated_conjecture, (r1_tarski(k3_subset_1(esk3_0,X1),esk1_0)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
fof(c_0_80, plain, ![X20, X21]:(~r1_xboole_0(X20,X21)|r1_xboole_0(X21,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_81, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k3_subset_1(X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_60, c_0_62])).
cnf(c_0_82, negated_conjecture, (r1_tarski(esk2_0,esk3_0)|r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_83, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X3|~r1_xboole_0(X3,X2)|~r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_84, plain, (r1_tarski(k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(k3_xboole_0(X1,X2),X3)),X1)), inference(spm,[status(thm)],[c_0_28, c_0_30])).
cnf(c_0_85, negated_conjecture, (k3_xboole_0(esk2_0,k5_xboole_0(esk2_0,esk2_0))=k5_xboole_0(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_67])])).
cnf(c_0_86, negated_conjecture, (k5_xboole_0(esk2_0,k5_xboole_0(esk2_0,esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_29]), c_0_77])])).
cnf(c_0_87, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k3_subset_1(X3,X2))|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_78, c_0_72])).
cnf(c_0_88, negated_conjecture, (r1_tarski(k3_subset_1(esk3_0,X1),k3_subset_1(esk3_0,X1))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_47, c_0_79])).
cnf(c_0_89, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k3_subset_1(X1,X2)))=k3_xboole_0(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_23, c_0_62])).
cnf(c_0_90, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_91, negated_conjecture, (r1_xboole_0(k3_subset_1(esk1_0,esk3_0),esk2_0)|r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_50])])).
cnf(c_0_92, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(k3_xboole_0(X1,X2),X3))=X1|~r1_xboole_0(X1,X3)|~r1_tarski(X1,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_93, negated_conjecture, (k3_xboole_0(esk2_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_85]), c_0_85]), c_0_86]), c_0_86])).
cnf(c_0_94, negated_conjecture, (r1_xboole_0(k3_subset_1(esk3_0,X1),X1)|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_95, plain, (r1_tarski(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X1,k3_subset_1(X2,X3))|~r1_tarski(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_74, c_0_89])).
cnf(c_0_96, negated_conjecture, (r1_xboole_0(esk2_0,k3_subset_1(esk1_0,esk3_0))|r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_97, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_28, c_0_25])).
cnf(c_0_98, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_47, c_0_30])).
cnf(c_0_99, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,X1))=esk2_0|~r1_xboole_0(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_67])])).
cnf(c_0_100, negated_conjecture, (r1_xboole_0(X1,k3_subset_1(esk3_0,X1))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_90, c_0_94])).
cnf(c_0_101, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_59]), c_0_43])]), c_0_97])).
cnf(c_0_102, plain, (r1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_60, c_0_98])).
cnf(c_0_103, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,k3_subset_1(esk3_0,esk2_0)))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])])).
cnf(c_0_104, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_36, c_0_25])).
cnf(c_0_105, plain, (k3_xboole_0(X1,X2)=X3|~r1_xboole_0(X3,k3_subset_1(X1,X2))|~r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_tarski(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_83, c_0_89])).
cnf(c_0_106, negated_conjecture, (r1_xboole_0(esk2_0,k3_subset_1(esk3_0,esk2_0))), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_107, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)|~r1_tarski(k3_subset_1(esk1_0,esk3_0),k3_subset_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_108, plain, (r1_tarski(X1,k3_subset_1(X2,X3))|~r1_xboole_0(X1,X3)|~r1_tarski(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_74, c_0_62])).
cnf(c_0_109, plain, (r1_xboole_0(k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_60, c_0_104])).
cnf(c_0_110, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=esk2_0|~m1_subset_1(esk2_0,k1_zfmisc_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_25]), c_0_25]), c_0_36]), c_0_101])])).
cnf(c_0_111, negated_conjecture, (~r1_xboole_0(k3_subset_1(esk1_0,esk3_0),esk2_0)|~r1_tarski(k3_subset_1(esk1_0,esk3_0),esk1_0)|~r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_50])])).
cnf(c_0_112, plain, (k3_subset_1(X1,X2)=k5_xboole_0(X1,X2)|~r1_tarski(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_29]), c_0_64])).
cnf(c_0_113, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X3,X2))|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_109, c_0_70])).
cnf(c_0_114, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_64]), c_0_101])])).
cnf(c_0_115, negated_conjecture, (~r1_xboole_0(k5_xboole_0(esk1_0,esk3_0),esk2_0)|~r1_tarski(k5_xboole_0(esk1_0,esk3_0),esk1_0)|~r1_tarski(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_45])])).
cnf(c_0_116, negated_conjecture, (r1_xboole_0(esk2_0,k5_xboole_0(X1,esk3_0))|~r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_117, negated_conjecture, (~r1_xboole_0(k5_xboole_0(esk1_0,esk3_0),esk2_0)|~r1_tarski(k5_xboole_0(esk1_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_115, c_0_101])])).
cnf(c_0_118, negated_conjecture, (r1_xboole_0(k5_xboole_0(X1,esk3_0),esk2_0)|~r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_90, c_0_116])).
cnf(c_0_119, negated_conjecture, (~r1_tarski(k5_xboole_0(esk1_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_45])])).
cnf(c_0_120, plain, (r1_tarski(k5_xboole_0(X1,X2),X1)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_30, c_0_29])).
cnf(c_0_121, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_45])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 122
# Proof object clause steps            : 95
# Proof object formula steps           : 27
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 67
# Proof object simplifying inferences  : 47
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 26435
# ...of these trivial                  : 600
# ...subsumed                          : 22511
# ...remaining for further processing  : 3324
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 513
# Backward-rewritten                   : 695
# Generated clauses                    : 328837
# ...of the previous two non-trivial   : 287149
# Contextual simplify-reflections      : 156
# Paramodulations                      : 328817
# Factorizations                       : 18
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2092
#    Positive orientable unit clauses  : 230
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 43
#    Non-unit-clauses                  : 1818
# Current number of unprocessed clauses: 252218
# ...number of literals in the above   : 807026
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1231
# Clause-clause subsumption calls (NU) : 718691
# Rec. Clause-clause subsumption calls : 424717
# Non-unit clause-clause subsumptions  : 10812
# Unit Clause-clause subsumption calls : 13343
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 965
# BW rewrite match successes           : 103
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6990584

# -------------------------------------------------
# User time                : 3.623 s
# System time              : 0.056 s
# Total time               : 3.679 s
# Maximum resident set size: 173628 pages
