# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t85_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(X2=k8_relat_1(X1,X3)<=>(![X4]:(r2_hidden(X4,k1_relat_1(X2))<=>(r2_hidden(X4,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X4),X1)))&![X4]:(r2_hidden(X4,k1_relat_1(X2))=>k1_funct_1(X2,X4)=k1_funct_1(X3,X4)))))), file('data/MPTP/proofs/9f7080d461ca4cacb04f7cbfa73290af.E_input', t85_funct_1)).
fof(d12_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k8_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X5,X1)&r2_hidden(k4_tarski(X4,X5),X2)))))), file('data/MPTP/proofs/9f7080d461ca4cacb04f7cbfa73290af.E_input', d12_relat_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/9f7080d461ca4cacb04f7cbfa73290af.E_input', t8_funct_1)).
fof(c_0_3, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(X2=k8_relat_1(X1,X3)<=>(![X4]:(r2_hidden(X4,k1_relat_1(X2))<=>(r2_hidden(X4,k1_relat_1(X3))&r2_hidden(k1_funct_1(X3,X4),X1)))&![X4]:(r2_hidden(X4,k1_relat_1(X2))=>k1_funct_1(X2,X4)=k1_funct_1(X3,X4))))))), inference(assume_negation,[status(cth)],[t85_funct_1])).
fof(c_0_4, plain, ![X16, X17, X18, X19, X20, X21, X22]:((((r2_hidden(X20,X16)|~r2_hidden(k4_tarski(X19,X20),X18)|X18!=k8_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(X19,X20),X17)|~r2_hidden(k4_tarski(X19,X20),X18)|X18!=k8_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17)))&(~r2_hidden(X22,X16)|~r2_hidden(k4_tarski(X21,X22),X17)|r2_hidden(k4_tarski(X21,X22),X18)|X18!=k8_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17)))&((~r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk7_3(X16,X17,X18)),X18)|(~r2_hidden(esk7_3(X16,X17,X18),X16)|~r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk7_3(X16,X17,X18)),X17))|X18=k8_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17))&((r2_hidden(esk7_3(X16,X17,X18),X16)|r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk7_3(X16,X17,X18)),X18)|X18=k8_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk7_3(X16,X17,X18)),X17)|r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk7_3(X16,X17,X18)),X18)|X18=k8_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_relat_1])])])])])])).
fof(c_0_5, negated_conjecture, ![X11, X12]:((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((((r2_hidden(esk5_0,k1_relat_1(esk2_0))|(~r2_hidden(esk4_0,k1_relat_1(esk2_0))|(~r2_hidden(esk4_0,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)))|esk2_0!=k8_relat_1(esk1_0,esk3_0))&(k1_funct_1(esk2_0,esk5_0)!=k1_funct_1(esk3_0,esk5_0)|(~r2_hidden(esk4_0,k1_relat_1(esk2_0))|(~r2_hidden(esk4_0,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)))|esk2_0!=k8_relat_1(esk1_0,esk3_0)))&(((r2_hidden(esk5_0,k1_relat_1(esk2_0))|(r2_hidden(esk4_0,k1_relat_1(esk3_0))|r2_hidden(esk4_0,k1_relat_1(esk2_0)))|esk2_0!=k8_relat_1(esk1_0,esk3_0))&(k1_funct_1(esk2_0,esk5_0)!=k1_funct_1(esk3_0,esk5_0)|(r2_hidden(esk4_0,k1_relat_1(esk3_0))|r2_hidden(esk4_0,k1_relat_1(esk2_0)))|esk2_0!=k8_relat_1(esk1_0,esk3_0)))&((r2_hidden(esk5_0,k1_relat_1(esk2_0))|(r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|r2_hidden(esk4_0,k1_relat_1(esk2_0)))|esk2_0!=k8_relat_1(esk1_0,esk3_0))&(k1_funct_1(esk2_0,esk5_0)!=k1_funct_1(esk3_0,esk5_0)|(r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|r2_hidden(esk4_0,k1_relat_1(esk2_0)))|esk2_0!=k8_relat_1(esk1_0,esk3_0)))))&((((r2_hidden(X11,k1_relat_1(esk3_0))|~r2_hidden(X11,k1_relat_1(esk2_0))|esk2_0=k8_relat_1(esk1_0,esk3_0))&(r2_hidden(k1_funct_1(esk3_0,X11),esk1_0)|~r2_hidden(X11,k1_relat_1(esk2_0))|esk2_0=k8_relat_1(esk1_0,esk3_0)))&(~r2_hidden(X11,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,X11),esk1_0)|r2_hidden(X11,k1_relat_1(esk2_0))|esk2_0=k8_relat_1(esk1_0,esk3_0)))&(~r2_hidden(X12,k1_relat_1(esk2_0))|k1_funct_1(esk2_0,X12)=k1_funct_1(esk3_0,X12)|esk2_0=k8_relat_1(esk1_0,esk3_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_6, plain, (r2_hidden(k4_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),X2)|r2_hidden(k4_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),X3)|X3=k8_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, ![X13, X14, X15]:(((r2_hidden(X13,k1_relat_1(X15))|~r2_hidden(k4_tarski(X13,X14),X15)|(~v1_relat_1(X15)|~v1_funct_1(X15)))&(X14=k1_funct_1(X15,X13)|~r2_hidden(k4_tarski(X13,X14),X15)|(~v1_relat_1(X15)|~v1_funct_1(X15))))&(~r2_hidden(X13,k1_relat_1(X15))|X14!=k1_funct_1(X15,X13)|r2_hidden(k4_tarski(X13,X14),X15)|(~v1_relat_1(X15)|~v1_funct_1(X15)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_9, negated_conjecture, (k8_relat_1(X1,X2)=esk2_0|r2_hidden(k4_tarski(esk6_3(X1,X2,esk2_0),esk7_3(X1,X2,esk2_0)),esk2_0)|r2_hidden(k4_tarski(esk6_3(X1,X2,esk2_0),esk7_3(X1,X2,esk2_0)),X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,X1),esk1_0)|esk2_0=k8_relat_1(esk1_0,esk3_0)|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, negated_conjecture, (k1_funct_1(esk2_0,X1)=k1_funct_1(esk3_0,X1)|esk2_0=k8_relat_1(esk1_0,esk3_0)|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_12, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, negated_conjecture, (k8_relat_1(X1,esk2_0)=esk2_0|r2_hidden(k4_tarski(esk6_3(X1,esk2_0,esk2_0),esk7_3(X1,esk2_0,esk2_0)),esk2_0)), inference(spm,[status(thm)],[c_0_9, c_0_7])).
cnf(c_0_14, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17, plain, (X3=k8_relat_1(X1,X2)|~r2_hidden(k4_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),X3)|~r2_hidden(esk7_3(X1,X2,X3),X1)|~r2_hidden(k4_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),X2)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_18, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|r2_hidden(k1_funct_1(esk2_0,X1),esk1_0)|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_19, negated_conjecture, (k1_funct_1(esk2_0,esk6_3(X1,esk2_0,esk2_0))=esk7_3(X1,esk2_0,esk2_0)|k8_relat_1(X1,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_7])])).
cnf(c_0_20, negated_conjecture, (k8_relat_1(X1,esk2_0)=esk2_0|r2_hidden(esk6_3(X1,esk2_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_13]), c_0_14]), c_0_7])])).
cnf(c_0_21, negated_conjecture, (k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(k4_tarski(esk6_3(X1,esk3_0,esk2_0),esk7_3(X1,esk3_0,esk2_0)),esk3_0)|r2_hidden(k4_tarski(esk6_3(X1,esk3_0,esk2_0),esk7_3(X1,esk3_0,esk2_0)),esk2_0)), inference(spm,[status(thm)],[c_0_9, c_0_16])).
cnf(c_0_22, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),X4)|X4!=k8_relat_1(X2,X5)|~v1_relat_1(X4)|~v1_relat_1(X5)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_24, negated_conjecture, (k8_relat_1(X1,esk2_0)=esk2_0|~r2_hidden(esk7_3(X1,esk2_0,esk2_0),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_13]), c_0_7])]), c_0_13])).
cnf(c_0_25, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(X1,esk2_0)=esk2_0|r2_hidden(esk7_3(X1,esk2_0,esk2_0),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_26, negated_conjecture, (k1_funct_1(esk3_0,esk6_3(X1,esk3_0,esk2_0))=esk7_3(X1,esk3_0,esk2_0)|k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(k4_tarski(esk6_3(X1,esk3_0,esk2_0),esk7_3(X1,esk3_0,esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_21]), c_0_22]), c_0_16])])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k8_relat_1(X2,X4))|~v1_relat_1(k8_relat_1(X2,X4))|~v1_relat_1(X4)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_28, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(esk1_0,esk2_0)=esk2_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_30, negated_conjecture, (k1_funct_1(esk3_0,esk6_3(X1,esk3_0,esk2_0))=esk7_3(X1,esk3_0,esk2_0)|k1_funct_1(esk2_0,esk6_3(X1,esk3_0,esk2_0))=esk7_3(X1,esk3_0,esk2_0)|k8_relat_1(X1,esk3_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_26]), c_0_14]), c_0_7])])).
cnf(c_0_31, negated_conjecture, (k8_relat_1(esk1_0,esk2_0)=esk2_0|r2_hidden(X1,esk1_0)|~r2_hidden(k4_tarski(X2,X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_7]), c_0_16])])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_33, negated_conjecture, (k1_funct_1(esk2_0,esk6_3(X1,esk3_0,esk2_0))=esk7_3(X1,esk3_0,esk2_0)|k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(X1,esk3_0)=esk2_0|~r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_11, c_0_30])).
cnf(c_0_34, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk3_0))|esk2_0=k8_relat_1(esk1_0,esk3_0)|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_35, plain, (r2_hidden(esk7_3(X1,X2,X3),X1)|r2_hidden(k4_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),X3)|X3=k8_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_36, negated_conjecture, (k8_relat_1(esk1_0,esk2_0)=esk2_0|k8_relat_1(X1,esk2_0)=esk2_0|r2_hidden(esk7_3(X1,esk2_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_31, c_0_13])).
cnf(c_0_37, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(k4_tarski(esk6_3(X1,esk3_0,esk2_0),esk7_3(X1,esk3_0,esk2_0)),esk2_0)|~r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_14]), c_0_7])])).
cnf(c_0_38, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|r2_hidden(k4_tarski(X1,k1_funct_1(esk2_0,X1)),esk3_0)|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_11]), c_0_22]), c_0_16])]), c_0_34])).
cnf(c_0_39, negated_conjecture, (k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(k4_tarski(esk6_3(X1,esk3_0,esk2_0),esk7_3(X1,esk3_0,esk2_0)),esk2_0)|r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_21]), c_0_22]), c_0_16])])).
cnf(c_0_40, plain, (k8_relat_1(X1,X2)=k8_relat_1(X3,X4)|r2_hidden(esk7_3(X3,X4,k8_relat_1(X1,X2)),X3)|r2_hidden(esk7_3(X3,X4,k8_relat_1(X1,X2)),X1)|~v1_relat_1(k8_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_27, c_0_35])).
cnf(c_0_41, negated_conjecture, (k8_relat_1(esk1_0,esk2_0)=esk2_0), inference(spm,[status(thm)],[c_0_24, c_0_36])).
cnf(c_0_42, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X1,X2),X4)|X4!=k8_relat_1(X5,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_43, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(X1,esk3_0)=esk2_0|~r2_hidden(k4_tarski(esk6_3(X1,esk3_0,esk2_0),esk7_3(X1,esk3_0,esk2_0)),esk3_0)|~r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk2_0))|~r2_hidden(esk7_3(X1,esk3_0,esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_37]), c_0_7]), c_0_16])])).
cnf(c_0_44, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(k4_tarski(esk6_3(X1,esk3_0,esk2_0),esk7_3(X1,esk3_0,esk2_0)),esk3_0)|~r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_38, c_0_33])).
cnf(c_0_45, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk2_0))|esk2_0=k8_relat_1(esk1_0,esk3_0)|~r2_hidden(X1,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,X1),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_46, negated_conjecture, (k1_funct_1(esk3_0,esk6_3(X1,esk3_0,esk2_0))=esk7_3(X1,esk3_0,esk2_0)|k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_26]), c_0_14]), c_0_7])])).
cnf(c_0_47, negated_conjecture, (k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk3_0))|r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_39]), c_0_14]), c_0_7])])).
cnf(c_0_48, negated_conjecture, (k8_relat_1(X1,X2)=esk2_0|r2_hidden(esk7_3(X1,X2,esk2_0),esk1_0)|r2_hidden(esk7_3(X1,X2,esk2_0),X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_7])])).
cnf(c_0_49, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X1,X2),k8_relat_1(X4,X3))|~v1_relat_1(k8_relat_1(X4,X3))|~v1_relat_1(X3)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_50, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(X1,esk3_0)=esk2_0|~r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk2_0))|~r2_hidden(esk7_3(X1,esk3_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(esk6_3(X1,esk3_0,esk2_0),k1_relat_1(esk2_0))|~r2_hidden(esk7_3(X1,esk3_0,esk2_0),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
cnf(c_0_52, negated_conjecture, (k8_relat_1(X1,esk3_0)=esk2_0|r2_hidden(esk7_3(X1,esk3_0,esk2_0),esk1_0)|r2_hidden(esk7_3(X1,esk3_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_48, c_0_16])).
cnf(c_0_53, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(k8_relat_1(X2,X3),X1)),X3)|~r2_hidden(X1,k1_relat_1(k8_relat_1(X2,X3)))|~v1_funct_1(k8_relat_1(X2,X3))|~v1_relat_1(k8_relat_1(X2,X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_49, c_0_32])).
cnf(c_0_54, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|k8_relat_1(X1,esk3_0)=esk2_0|~r2_hidden(esk7_3(X1,esk3_0,esk2_0),esk1_0)|~r2_hidden(esk7_3(X1,esk3_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_55, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0|r2_hidden(esk7_3(esk1_0,esk3_0,esk2_0),esk1_0)), inference(ef,[status(thm)],[c_0_52])).
cnf(c_0_56, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,k1_relat_1(k8_relat_1(X3,X2)))|~v1_funct_1(k8_relat_1(X3,X2))|~v1_funct_1(X2)|~v1_relat_1(k8_relat_1(X3,X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_15, c_0_53])).
cnf(c_0_57, negated_conjecture, (k8_relat_1(esk1_0,esk3_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_55])).
cnf(c_0_58, negated_conjecture, (k1_funct_1(esk2_0,esk5_0)!=k1_funct_1(esk3_0,esk5_0)|~r2_hidden(esk4_0,k1_relat_1(esk2_0))|~r2_hidden(esk4_0,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|esk2_0!=k8_relat_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_59, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk3_0))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_14]), c_0_22]), c_0_7]), c_0_16])])).
cnf(c_0_60, plain, (k1_funct_1(k8_relat_1(X1,X2),X3)=k1_funct_1(X2,X3)|~r2_hidden(X3,k1_relat_1(k8_relat_1(X1,X2)))|~v1_funct_1(k8_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_relat_1(k8_relat_1(X1,X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_12, c_0_53])).
cnf(c_0_61, plain, (r2_hidden(k1_funct_1(k8_relat_1(X1,X2),X3),X1)|~r2_hidden(X3,k1_relat_1(k8_relat_1(X1,X2)))|~v1_funct_1(k8_relat_1(X1,X2))|~v1_relat_1(k8_relat_1(X1,X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_27, c_0_32])).
cnf(c_0_62, negated_conjecture, (k1_funct_1(esk3_0,esk5_0)!=k1_funct_1(esk2_0,esk5_0)|~r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|~r2_hidden(esk4_0,k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_57])]), c_0_59])).
cnf(c_0_63, negated_conjecture, (k1_funct_1(esk3_0,X1)=k1_funct_1(esk2_0,X1)|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_57]), c_0_14]), c_0_22]), c_0_7]), c_0_16])])).
cnf(c_0_64, negated_conjecture, (r2_hidden(k1_funct_1(esk2_0,X1),esk1_0)|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_41]), c_0_14]), c_0_7])])).
cnf(c_0_65, plain, (r2_hidden(k4_tarski(X3,X1),X5)|~r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),X4)|X5!=k8_relat_1(X2,X4)|~v1_relat_1(X5)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_66, negated_conjecture, (k1_funct_1(esk3_0,esk5_0)!=k1_funct_1(esk2_0,esk5_0)|~r2_hidden(esk4_0,k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])).
cnf(c_0_67, plain, (r2_hidden(k4_tarski(X1,X2),k8_relat_1(X3,X4))|~r2_hidden(k4_tarski(X1,X2),X4)|~r2_hidden(X2,X3)|~v1_relat_1(k8_relat_1(X3,X4))|~v1_relat_1(X4)), inference(er,[status(thm)],[c_0_65])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk5_0,k1_relat_1(esk2_0))|~r2_hidden(esk4_0,k1_relat_1(esk2_0))|~r2_hidden(esk4_0,k1_relat_1(esk3_0))|~r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|esk2_0!=k8_relat_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_69, negated_conjecture, (~r2_hidden(esk4_0,k1_relat_1(esk2_0))|~r2_hidden(esk5_0,k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_66, c_0_63])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk4_0,k1_relat_1(esk3_0))|r2_hidden(esk4_0,k1_relat_1(esk2_0))|k1_funct_1(esk2_0,esk5_0)!=k1_funct_1(esk3_0,esk5_0)|esk2_0!=k8_relat_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_71, plain, (r2_hidden(X1,k1_relat_1(k8_relat_1(X2,X3)))|~r2_hidden(k4_tarski(X1,X4),X3)|~r2_hidden(X4,X2)|~v1_funct_1(k8_relat_1(X2,X3))|~v1_relat_1(k8_relat_1(X2,X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_15, c_0_67])).
cnf(c_0_72, negated_conjecture, (~r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|~r2_hidden(esk4_0,k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_57])]), c_0_59]), c_0_69])).
cnf(c_0_73, negated_conjecture, (r2_hidden(esk4_0,k1_relat_1(esk3_0))|r2_hidden(esk4_0,k1_relat_1(esk2_0))|k1_funct_1(esk3_0,esk5_0)!=k1_funct_1(esk2_0,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_57])])).
cnf(c_0_74, negated_conjecture, (r2_hidden(esk5_0,k1_relat_1(esk2_0))|r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|r2_hidden(esk4_0,k1_relat_1(esk2_0))|esk2_0!=k8_relat_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_75, plain, (r2_hidden(X1,k1_relat_1(k8_relat_1(X2,X3)))|~r2_hidden(k1_funct_1(X3,X1),X2)|~r2_hidden(X1,k1_relat_1(X3))|~v1_funct_1(k8_relat_1(X2,X3))|~v1_funct_1(X3)|~v1_relat_1(k8_relat_1(X2,X3))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_71, c_0_32])).
cnf(c_0_76, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|r2_hidden(esk4_0,k1_relat_1(esk2_0))|k1_funct_1(esk2_0,esk5_0)!=k1_funct_1(esk3_0,esk5_0)|esk2_0!=k8_relat_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_77, negated_conjecture, (~r2_hidden(esk4_0,k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_63]), c_0_64])).
cnf(c_0_78, negated_conjecture, (r2_hidden(esk5_0,k1_relat_1(esk2_0))|r2_hidden(esk4_0,k1_relat_1(esk3_0))|r2_hidden(esk4_0,k1_relat_1(esk2_0))|esk2_0!=k8_relat_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_79, negated_conjecture, (r2_hidden(esk4_0,k1_relat_1(esk3_0))|~r2_hidden(esk5_0,k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_63]), c_0_59])).
cnf(c_0_80, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|r2_hidden(esk5_0,k1_relat_1(esk2_0))|r2_hidden(esk4_0,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_57])])).
cnf(c_0_81, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk2_0))|~r2_hidden(k1_funct_1(esk3_0,X1),esk1_0)|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_57]), c_0_14]), c_0_22]), c_0_7]), c_0_16])])).
cnf(c_0_82, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|k1_funct_1(esk3_0,esk5_0)!=k1_funct_1(esk2_0,esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_57])]), c_0_77])).
cnf(c_0_83, negated_conjecture, (r2_hidden(esk4_0,k1_relat_1(esk3_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_57])]), c_0_77]), c_0_79])).
cnf(c_0_84, negated_conjecture, (r2_hidden(k1_funct_1(esk3_0,esk4_0),esk1_0)|r2_hidden(esk5_0,k1_relat_1(esk2_0))), inference(sr,[status(thm)],[c_0_80, c_0_77])).
cnf(c_0_85, negated_conjecture, (k1_funct_1(esk3_0,esk5_0)!=k1_funct_1(esk2_0,esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])]), c_0_77])).
cnf(c_0_86, negated_conjecture, (r2_hidden(esk5_0,k1_relat_1(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_84]), c_0_83])]), c_0_77])).
cnf(c_0_87, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_63]), c_0_86])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 88
# Proof object clause steps            : 81
# Proof object formula steps           : 7
# Proof object conjectures             : 64
# Proof object clause conjectures      : 61
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 47
# Proof object simplifying inferences  : 94
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 398
# ...of these trivial                  : 0
# ...subsumed                          : 127
# ...remaining for further processing  : 270
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 28
# Backward-rewritten                   : 80
# Generated clauses                    : 925
# ...of the previous two non-trivial   : 849
# Contextual simplify-reflections      : 35
# Paramodulations                      : 918
# Factorizations                       : 2
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 134
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 124
# Current number of unprocessed clauses: 401
# ...number of literals in the above   : 3151
# Current number of archived formulas  : 0
# Current number of archived clauses   : 132
# Clause-clause subsumption calls (NU) : 10769
# Rec. Clause-clause subsumption calls : 854
# Non-unit clause-clause subsumptions  : 185
# Unit Clause-clause subsumption calls : 118
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 36643

# -------------------------------------------------
# User time                : 0.032 s
# System time              : 0.003 s
# Total time               : 0.035 s
# Maximum resident set size: 28568 pages
