# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t47_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(X3=X4=>k2_pre_topc(X2,X4)=k9_subset_1(u1_struct_0(X2),k2_pre_topc(X1,X3),k2_struct_0(X2))))))), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t47_pre_topc)).
fof(t39_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t39_pre_topc)).
fof(t31_xboole_1, axiom, ![X1, X2, X3]:r1_tarski(k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)),k2_xboole_0(X2,X3)), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t31_xboole_1)).
fof(t23_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t23_xboole_1)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t22_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t2_boole)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t52_pre_topc)).
fof(t34_tops_2, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_pre_topc(X3,X1)=>(v4_pre_topc(X2,X1)=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))=>(X4=X2=>v4_pre_topc(X4,X3))))))), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t34_tops_2)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', redefinition_k9_subset_1)).
fof(commutativity_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k9_subset_1(X1,X3,X2)), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', commutativity_k9_subset_1)).
fof(t43_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(v4_pre_topc(X3,X2)<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(v4_pre_topc(X4,X1)&k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2))=X3)))))), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t43_pre_topc)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', commutativity_k3_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', d10_xboole_0)).
fof(t48_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_tarski(X2,k2_pre_topc(X1,X2)))), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', t48_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/9fc5eef171014c87823f332c6dc8fb47.E_input', dt_m1_pre_topc)).
fof(c_0_15, plain, ![X16, X17, X18, X19]:(~l1_pre_topc(X16)|(~m1_pre_topc(X17,X16)|(~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X16)))|(~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X17)))|(X18!=X19|k2_pre_topc(X17,X19)=k9_subset_1(u1_struct_0(X17),k2_pre_topc(X16,X18),k2_struct_0(X17))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t47_pre_topc])])])).
fof(c_0_16, plain, ![X20, X21, X22]:(~l1_pre_topc(X20)|(~m1_pre_topc(X21,X20)|(~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))|m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t39_pre_topc])])])).
fof(c_0_17, plain, ![X11, X12, X13]:r1_tarski(k2_xboole_0(k3_xboole_0(X11,X12),k3_xboole_0(X11,X13)),k2_xboole_0(X12,X13)), inference(variable_rename,[status(thm)],[t31_xboole_1])).
fof(c_0_18, plain, ![X30, X31, X32]:k3_xboole_0(X30,k2_xboole_0(X31,X32))=k2_xboole_0(k3_xboole_0(X30,X31),k3_xboole_0(X30,X32)), inference(variable_rename,[status(thm)],[t23_xboole_1])).
fof(c_0_19, plain, ![X35, X36]:k2_xboole_0(X35,k3_xboole_0(X35,X36))=X35, inference(variable_rename,[status(thm)],[t22_xboole_1])).
fof(c_0_20, plain, ![X37]:k3_xboole_0(X37,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_21, plain, (k2_pre_topc(X2,X4)=k9_subset_1(u1_struct_0(X2),k2_pre_topc(X1,X3),k2_struct_0(X2))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, plain, ![X41, X42]:((~v4_pre_topc(X42,X41)|k2_pre_topc(X41,X42)=X42|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_pre_topc(X41))&(~v2_pre_topc(X41)|k2_pre_topc(X41,X42)!=X42|v4_pre_topc(X42,X41)|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_pre_topc(X41))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
fof(c_0_24, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_pre_topc(X3,X1)=>(v4_pre_topc(X2,X1)=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))=>(X4=X2=>v4_pre_topc(X4,X3)))))))), inference(assume_negation,[status(cth)],[t34_tops_2])).
cnf(c_0_25, plain, (r1_tarski(k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)),k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, ![X27, X28, X29]:(~m1_subset_1(X29,k1_zfmisc_1(X27))|k9_subset_1(X27,X28,X29)=k3_xboole_0(X28,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
fof(c_0_30, plain, ![X38, X39, X40]:(~m1_subset_1(X40,k1_zfmisc_1(X38))|k9_subset_1(X38,X39,X40)=k9_subset_1(X38,X40,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k9_subset_1])])).
cnf(c_0_31, plain, (k9_subset_1(u1_struct_0(X1),k2_pre_topc(X2,X3),k2_struct_0(X1))=k2_pre_topc(X1,X3)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_21]), c_0_22])).
cnf(c_0_32, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_33, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_pre_topc(esk3_0,esk1_0)&(v4_pre_topc(esk2_0,esk1_0)&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))&(esk4_0=esk2_0&~v4_pre_topc(esk4_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
fof(c_0_34, plain, ![X45, X46, X47, X49]:(((m1_subset_1(esk5_3(X45,X46,X47),k1_zfmisc_1(u1_struct_0(X45)))|~v4_pre_topc(X47,X46)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46)))|~m1_pre_topc(X46,X45)|~l1_pre_topc(X45))&((v4_pre_topc(esk5_3(X45,X46,X47),X45)|~v4_pre_topc(X47,X46)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46)))|~m1_pre_topc(X46,X45)|~l1_pre_topc(X45))&(k9_subset_1(u1_struct_0(X46),esk5_3(X45,X46,X47),k2_struct_0(X46))=X47|~v4_pre_topc(X47,X46)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46)))|~m1_pre_topc(X46,X45)|~l1_pre_topc(X45))))&(~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X45)))|(~v4_pre_topc(X49,X45)|k9_subset_1(u1_struct_0(X46),X49,k2_struct_0(X46))!=X47)|v4_pre_topc(X47,X46)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X46)))|~m1_pre_topc(X46,X45)|~l1_pre_topc(X45))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_pre_topc])])])])])).
cnf(c_0_35, plain, (r1_tarski(k3_xboole_0(X1,k2_xboole_0(X2,X3)),k2_xboole_0(X2,X3))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_36, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_37, plain, ![X33, X34]:k3_xboole_0(X33,X34)=k3_xboole_0(X34,X33), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_38, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (k9_subset_1(X2,X3,X1)=k9_subset_1(X2,X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (k9_subset_1(u1_struct_0(X1),X2,k2_struct_0(X1))=k2_pre_topc(X1,X2)|~v4_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_22])).
cnf(c_0_41, negated_conjecture, (v4_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (v4_pre_topc(X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v4_pre_topc(X1,X2)|k9_subset_1(u1_struct_0(X3),X1,k2_struct_0(X3))!=X4|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_44, plain, ![X43, X44]:(((r1_tarski(X43,X44)|X43!=X44)&(r1_tarski(X44,X43)|X43!=X44))&(~r1_tarski(X43,X44)|~r1_tarski(X44,X43)|X43=X44)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_45, plain, ![X14, X15]:(~l1_pre_topc(X14)|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))|r1_tarski(X15,k2_pre_topc(X14,X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t48_pre_topc])])])).
cnf(c_0_46, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_47, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (k9_subset_1(X1,X2,X3)=k3_xboole_0(X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_49, negated_conjecture, (k9_subset_1(u1_struct_0(X1),esk2_0,k2_struct_0(X1))=k2_pre_topc(X1,esk2_0)|~m1_pre_topc(X1,esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42])])).
cnf(c_0_50, plain, (v4_pre_topc(k9_subset_1(u1_struct_0(X1),X2,k2_struct_0(X1)),X1)|~v4_pre_topc(X2,X3)|~m1_pre_topc(X1,X3)|~m1_subset_1(k9_subset_1(u1_struct_0(X1),X2,k2_struct_0(X1)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_51, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_53, negated_conjecture, (esk4_0=esk2_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_54, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (r1_tarski(X2,k2_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_56, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_57, negated_conjecture, (k3_xboole_0(esk2_0,k2_struct_0(X1))=k2_pre_topc(X1,esk2_0)|~m1_pre_topc(X1,esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_47])).
fof(c_0_58, plain, ![X50, X51]:(~l1_pre_topc(X50)|(~m1_pre_topc(X51,X50)|l1_pre_topc(X51))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_59, negated_conjecture, (v4_pre_topc(k9_subset_1(u1_struct_0(esk3_0),X1,k2_struct_0(esk3_0)),esk3_0)|~v4_pre_topc(X1,esk1_0)|~m1_subset_1(k9_subset_1(u1_struct_0(esk3_0),X1,k2_struct_0(esk3_0)),k1_zfmisc_1(u1_struct_0(esk3_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_42])])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(rw,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_62, plain, (k2_pre_topc(X1,X2)=X2|~r1_tarski(k2_pre_topc(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_63, negated_conjecture, (r1_tarski(k2_pre_topc(X1,esk2_0),esk2_0)|~m1_pre_topc(X1,esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_64, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65, negated_conjecture, (~v4_pre_topc(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_66, negated_conjecture, (v4_pre_topc(k2_pre_topc(esk3_0,esk2_0),esk3_0)|~m1_subset_1(k2_pre_topc(esk3_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_49]), c_0_41]), c_0_60]), c_0_51]), c_0_61])])).
cnf(c_0_67, negated_conjecture, (k2_pre_topc(X1,esk2_0)=esk2_0|~m1_pre_topc(X1,esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_51]), c_0_42])])).
cnf(c_0_69, negated_conjecture, (~v4_pre_topc(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_65, c_0_53])).
cnf(c_0_70, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_61]), c_0_51]), c_0_68])]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 71
# Proof object clause steps            : 40
# Proof object formula steps           : 31
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 374
# ...of these trivial                  : 12
# ...subsumed                          : 204
# ...remaining for further processing  : 158
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 1
# Generated clauses                    : 766
# ...of the previous two non-trivial   : 682
# Contextual simplify-reflections      : 22
# Paramodulations                      : 762
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 123
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 95
# Current number of unprocessed clauses: 354
# ...number of literals in the above   : 2070
# Current number of archived formulas  : 0
# Current number of archived clauses   : 31
# Clause-clause subsumption calls (NU) : 3461
# Rec. Clause-clause subsumption calls : 1702
# Non-unit clause-clause subsumptions  : 216
# Unit Clause-clause subsumption calls : 7
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 28
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15319

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.000 s
# Total time               : 0.019 s
# Maximum resident set size: 28684 pages
