# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.011 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', t39_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', t24_orders_2)).
fof(t18_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', t18_yellow_0)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', d7_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', dt_k1_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', d10_yellow_0)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', dt_k2_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', t7_yellow_0)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', t69_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', d1_tarski)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', redefinition_k6_domain_1)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', d9_lattice3)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', dt_l1_orders_2)).
fof(t30_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3))=>(r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4)))))&((r2_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X3,X4)=>r1_orders_2(X1,X2,X4))))=>(X2=k1_yellow_0(X1,X3)&r1_yellow_0(X1,X3)))))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', t30_yellow_0)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/9fd922d5eb4a4e42993ff36ffff19018.E_input', t19_yellow_0)).
fof(c_0_17, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_18, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5))))))&((r1_orders_2(X1,X2,X4)&(r1_orders_2(X1,X3,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X2,X5)&r1_orders_2(X1,X3,X5))=>r1_orders_2(X1,X4,X5)))))=>(X4=k10_lattice3(X1,X2,X3)&r1_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_17])).
fof(c_0_19, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2)))), inference(assume_negation,[status(cth)],[t39_yellow_0])).
fof(c_0_20, plain, ![X67, X68, X69, X70, X71]:(((r1_orders_2(X69,X68,X70)|(X70!=k10_lattice3(X69,X68,X67)|~r1_yellow_0(X69,k2_tarski(X68,X67)))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67))&((r1_orders_2(X69,X67,X70)|(X70!=k10_lattice3(X69,X68,X67)|~r1_yellow_0(X69,k2_tarski(X68,X67)))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67))&(~m1_subset_1(X71,u1_struct_0(X69))|(~r1_orders_2(X69,X68,X71)|~r1_orders_2(X69,X67,X71)|r1_orders_2(X69,X70,X71))|(X70!=k10_lattice3(X69,X68,X67)|~r1_yellow_0(X69,k2_tarski(X68,X67)))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67))))&(((X70=k10_lattice3(X69,X68,X67)|(m1_subset_1(esk12_4(X67,X68,X69,X70),u1_struct_0(X69))|~r1_orders_2(X69,X67,X70)|~r1_orders_2(X69,X68,X70))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67))&(r1_yellow_0(X69,k2_tarski(X68,X67))|(m1_subset_1(esk12_4(X67,X68,X69,X70),u1_struct_0(X69))|~r1_orders_2(X69,X67,X70)|~r1_orders_2(X69,X68,X70))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67)))&((((X70=k10_lattice3(X69,X68,X67)|(r1_orders_2(X69,X68,esk12_4(X67,X68,X69,X70))|~r1_orders_2(X69,X67,X70)|~r1_orders_2(X69,X68,X70))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67))&(r1_yellow_0(X69,k2_tarski(X68,X67))|(r1_orders_2(X69,X68,esk12_4(X67,X68,X69,X70))|~r1_orders_2(X69,X67,X70)|~r1_orders_2(X69,X68,X70))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67)))&((X70=k10_lattice3(X69,X68,X67)|(r1_orders_2(X69,X67,esk12_4(X67,X68,X69,X70))|~r1_orders_2(X69,X67,X70)|~r1_orders_2(X69,X68,X70))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67))&(r1_yellow_0(X69,k2_tarski(X68,X67))|(r1_orders_2(X69,X67,esk12_4(X67,X68,X69,X70))|~r1_orders_2(X69,X67,X70)|~r1_orders_2(X69,X68,X70))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67))))&((X70=k10_lattice3(X69,X68,X67)|(~r1_orders_2(X69,X70,esk12_4(X67,X68,X69,X70))|~r1_orders_2(X69,X67,X70)|~r1_orders_2(X69,X68,X70))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67))&(r1_yellow_0(X69,k2_tarski(X68,X67))|(~r1_orders_2(X69,X70,esk12_4(X67,X68,X69,X70))|~r1_orders_2(X69,X67,X70)|~r1_orders_2(X69,X68,X70))|~m1_subset_1(X70,u1_struct_0(X69))|~epred1_3(X69,X68,X67)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])).
fof(c_0_21, plain, ![X8, X9]:(v2_struct_0(X8)|(~v3_orders_2(X8)|~l1_orders_2(X8))|(~m1_subset_1(X9,u1_struct_0(X8))|r1_orders_2(X8,X9,X9))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_22, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_19])])])])).
cnf(c_0_23, plain, (X1=k10_lattice3(X2,X3,X4)|~r1_orders_2(X2,X1,esk12_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X4,esk12_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (X1=k10_lattice3(X2,X3,X4)|r1_orders_2(X2,X3,esk12_4(X4,X3,X2,X1))|~r1_orders_2(X2,X4,X1)|~r1_orders_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (k10_lattice3(X1,X2,X3)=X3|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])]), c_0_28])).
cnf(c_0_32, plain, (k10_lattice3(X1,X2,X3)=X2|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X2)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_23, c_0_29])).
fof(c_0_33, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t18_yellow_0, c_0_17])).
cnf(c_0_34, negated_conjecture, (k10_lattice3(esk1_0,X1,X2)=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, negated_conjecture, (k10_lattice3(esk1_0,X1,X2)=X1|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_32, c_0_31])).
fof(c_0_36, plain, ![X10, X11, X12]:(~v5_orders_2(X10)|~l1_orders_2(X10)|(~m1_subset_1(X11,u1_struct_0(X10))|(~m1_subset_1(X12,u1_struct_0(X10))|epred1_3(X10,X11,X12)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).
cnf(c_0_37, negated_conjecture, (X1=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_39, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_40, plain, ![X13, X14, X16, X17, X19, X20, X23]:(((m1_subset_1(esk3_2(X13,X14),u1_struct_0(X13))|~r1_yellow_0(X13,X14)|~l1_orders_2(X13))&((r2_lattice3(X13,X14,esk3_2(X13,X14))|~r1_yellow_0(X13,X14)|~l1_orders_2(X13))&((~m1_subset_1(X16,u1_struct_0(X13))|(~r2_lattice3(X13,X14,X16)|r1_orders_2(X13,esk3_2(X13,X14),X16))|~r1_yellow_0(X13,X14)|~l1_orders_2(X13))&((m1_subset_1(esk4_3(X13,X14,X17),u1_struct_0(X13))|~r2_lattice3(X13,X14,X17)|X17=esk3_2(X13,X14)|~m1_subset_1(X17,u1_struct_0(X13))|~r1_yellow_0(X13,X14)|~l1_orders_2(X13))&((r2_lattice3(X13,X14,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X14,X17)|X17=esk3_2(X13,X14)|~m1_subset_1(X17,u1_struct_0(X13))|~r1_yellow_0(X13,X14)|~l1_orders_2(X13))&(~r1_orders_2(X13,X17,esk4_3(X13,X14,X17))|~r2_lattice3(X13,X14,X17)|X17=esk3_2(X13,X14)|~m1_subset_1(X17,u1_struct_0(X13))|~r1_yellow_0(X13,X14)|~l1_orders_2(X13)))))))&(((m1_subset_1(esk6_3(X13,X19,X20),u1_struct_0(X13))|m1_subset_1(esk5_3(X13,X19,X20),u1_struct_0(X13))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))&(((r2_lattice3(X13,X19,esk6_3(X13,X19,X20))|m1_subset_1(esk5_3(X13,X19,X20),u1_struct_0(X13))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))&(~m1_subset_1(X23,u1_struct_0(X13))|(~r2_lattice3(X13,X19,X23)|r1_orders_2(X13,esk6_3(X13,X19,X20),X23))|m1_subset_1(esk5_3(X13,X19,X20),u1_struct_0(X13))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13)))&(esk6_3(X13,X19,X20)!=X20|m1_subset_1(esk5_3(X13,X19,X20),u1_struct_0(X13))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))))&(((m1_subset_1(esk6_3(X13,X19,X20),u1_struct_0(X13))|r2_lattice3(X13,X19,esk5_3(X13,X19,X20))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))&(((r2_lattice3(X13,X19,esk6_3(X13,X19,X20))|r2_lattice3(X13,X19,esk5_3(X13,X19,X20))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))&(~m1_subset_1(X23,u1_struct_0(X13))|(~r2_lattice3(X13,X19,X23)|r1_orders_2(X13,esk6_3(X13,X19,X20),X23))|r2_lattice3(X13,X19,esk5_3(X13,X19,X20))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13)))&(esk6_3(X13,X19,X20)!=X20|r2_lattice3(X13,X19,esk5_3(X13,X19,X20))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))))&((m1_subset_1(esk6_3(X13,X19,X20),u1_struct_0(X13))|~r1_orders_2(X13,X20,esk5_3(X13,X19,X20))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))&(((r2_lattice3(X13,X19,esk6_3(X13,X19,X20))|~r1_orders_2(X13,X20,esk5_3(X13,X19,X20))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))&(~m1_subset_1(X23,u1_struct_0(X13))|(~r2_lattice3(X13,X19,X23)|r1_orders_2(X13,esk6_3(X13,X19,X20),X23))|~r1_orders_2(X13,X20,esk5_3(X13,X19,X20))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13)))&(esk6_3(X13,X19,X20)!=X20|~r1_orders_2(X13,X20,esk5_3(X13,X19,X20))|~r2_lattice3(X13,X19,X20)|~m1_subset_1(X20,u1_struct_0(X13))|r1_yellow_0(X13,X19)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
fof(c_0_41, plain, ![X24, X25, X26, X27]:(((r2_lattice3(X24,X25,X26)|X26!=k1_yellow_0(X24,X25)|~r1_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&(~m1_subset_1(X27,u1_struct_0(X24))|(~r2_lattice3(X24,X25,X27)|r1_orders_2(X24,X26,X27))|X26!=k1_yellow_0(X24,X25)|~r1_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24)))&((m1_subset_1(esk7_3(X24,X25,X26),u1_struct_0(X24))|~r2_lattice3(X24,X25,X26)|X26=k1_yellow_0(X24,X25)|~r1_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&((r2_lattice3(X24,X25,esk7_3(X24,X25,X26))|~r2_lattice3(X24,X25,X26)|X26=k1_yellow_0(X24,X25)|~r1_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))&(~r1_orders_2(X24,X26,esk7_3(X24,X25,X26))|~r2_lattice3(X24,X25,X26)|X26=k1_yellow_0(X24,X25)|~r1_yellow_0(X24,X25)|~m1_subset_1(X26,u1_struct_0(X24))|~l1_orders_2(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_42, plain, ![X29, X30]:(~l1_orders_2(X29)|m1_subset_1(k1_yellow_0(X29,X30),u1_struct_0(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_43, negated_conjecture, (X1=X2|~r1_orders_2(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_27]), c_0_39])])).
cnf(c_0_44, plain, (r1_orders_2(X2,esk3_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_46, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_47, plain, ![X49, X50, X51, X52]:(((r1_lattice3(X49,X50,X51)|X51!=k2_yellow_0(X49,X50)|~r2_yellow_0(X49,X50)|~m1_subset_1(X51,u1_struct_0(X49))|~l1_orders_2(X49))&(~m1_subset_1(X52,u1_struct_0(X49))|(~r1_lattice3(X49,X50,X52)|r1_orders_2(X49,X52,X51))|X51!=k2_yellow_0(X49,X50)|~r2_yellow_0(X49,X50)|~m1_subset_1(X51,u1_struct_0(X49))|~l1_orders_2(X49)))&((m1_subset_1(esk10_3(X49,X50,X51),u1_struct_0(X49))|~r1_lattice3(X49,X50,X51)|X51=k2_yellow_0(X49,X50)|~r2_yellow_0(X49,X50)|~m1_subset_1(X51,u1_struct_0(X49))|~l1_orders_2(X49))&((r1_lattice3(X49,X50,esk10_3(X49,X50,X51))|~r1_lattice3(X49,X50,X51)|X51=k2_yellow_0(X49,X50)|~r2_yellow_0(X49,X50)|~m1_subset_1(X51,u1_struct_0(X49))|~l1_orders_2(X49))&(~r1_orders_2(X49,esk10_3(X49,X50,X51),X51)|~r1_lattice3(X49,X50,X51)|X51=k2_yellow_0(X49,X50)|~r2_yellow_0(X49,X50)|~m1_subset_1(X51,u1_struct_0(X49))|~l1_orders_2(X49))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_48, plain, ![X54, X55]:(~l1_orders_2(X54)|m1_subset_1(k2_yellow_0(X54,X55),u1_struct_0(X54))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
fof(c_0_49, plain, ![X46, X47, X48]:((~r1_lattice3(X46,k1_tarski(X48),X47)|r1_orders_2(X46,X47,X48)|~m1_subset_1(X48,u1_struct_0(X46))|~m1_subset_1(X47,u1_struct_0(X46))|~l1_orders_2(X46))&((~r1_orders_2(X46,X47,X48)|r1_lattice3(X46,k1_tarski(X48),X47)|~m1_subset_1(X48,u1_struct_0(X46))|~m1_subset_1(X47,u1_struct_0(X46))|~l1_orders_2(X46))&((~r2_lattice3(X46,k1_tarski(X48),X47)|r1_orders_2(X46,X48,X47)|~m1_subset_1(X48,u1_struct_0(X46))|~m1_subset_1(X47,u1_struct_0(X46))|~l1_orders_2(X46))&(~r1_orders_2(X46,X48,X47)|r2_lattice3(X46,k1_tarski(X48),X47)|~m1_subset_1(X48,u1_struct_0(X46))|~m1_subset_1(X47,u1_struct_0(X46))|~l1_orders_2(X46))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
fof(c_0_50, plain, ![X43]:k2_tarski(X43,X43)=k1_tarski(X43), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_51, negated_conjecture, (esk3_2(esk1_0,X1)=X2|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,X2,esk3_2(esk1_0,X1))|~m1_subset_1(esk3_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_27])])).
cnf(c_0_52, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_45]), c_0_46])).
cnf(c_0_53, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, negated_conjecture, (esk3_2(esk1_0,X1)=k1_yellow_0(esk1_0,X2)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X2))|~r2_lattice3(esk1_0,X2,esk3_2(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk3_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_27])])).
cnf(c_0_59, plain, (r2_lattice3(X1,X2,esk3_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_60, plain, ![X36, X37, X38, X39, X40, X41]:(((~r2_hidden(X38,X37)|X38=X36|X37!=k1_tarski(X36))&(X39!=X36|r2_hidden(X39,X37)|X37!=k1_tarski(X36)))&((~r2_hidden(esk9_2(X40,X41),X41)|esk9_2(X40,X41)!=X40|X41=k1_tarski(X40))&(r2_hidden(esk9_2(X40,X41),X41)|esk9_2(X40,X41)=X40|X41=k1_tarski(X40)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_61, plain, ![X44, X45]:(v1_xboole_0(X44)|~m1_subset_1(X45,X44)|k6_domain_1(X44,X45)=k1_tarski(X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_62, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_53]), c_0_54])).
cnf(c_0_63, plain, (r1_orders_2(X1,X3,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_64, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_57]), c_0_54])).
cnf(c_0_65, negated_conjecture, (esk3_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk3_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_27])])).
cnf(c_0_66, plain, (m1_subset_1(esk3_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_67, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_68, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_69, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_70, negated_conjecture, (X1=k2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X2)|~r1_lattice3(esk1_0,X2,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),X1)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_62]), c_0_27])])).
cnf(c_0_71, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X2,X2)),X2)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_54])).
cnf(c_0_72, negated_conjecture, (esk3_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_27])])).
cnf(c_0_73, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_67]), c_0_46])).
cnf(c_0_74, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_75, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_68, c_0_56])).
cnf(c_0_76, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_69, c_0_56])).
cnf(c_0_77, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_27])])).
cnf(c_0_78, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_79, plain, (r1_orders_2(X1,X2,X3)|~r2_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_80, negated_conjecture, (esk3_2(esk1_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_27])])).
fof(c_0_81, plain, ![X31, X32, X33, X34]:((~r2_lattice3(X31,X32,X33)|(~m1_subset_1(X34,u1_struct_0(X31))|(~r2_hidden(X34,X32)|r1_orders_2(X31,X34,X33)))|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&((m1_subset_1(esk8_3(X31,X32,X33),u1_struct_0(X31))|r2_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&((r2_hidden(esk8_3(X31,X32,X33),X32)|r2_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))&(~r1_orders_2(X31,esk8_3(X31,X32,X33),X33)|r2_lattice3(X31,X32,X33)|~m1_subset_1(X33,u1_struct_0(X31))|~l1_orders_2(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_82, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_74, c_0_56])).
cnf(c_0_83, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_75])])).
cnf(c_0_84, plain, (k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))=k6_domain_1(u1_struct_0(X1),k2_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_76, c_0_54])).
cnf(c_0_85, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_54]), c_0_27])])).
cnf(c_0_86, plain, (r1_lattice3(X1,k2_tarski(X3,X3),X2)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_78, c_0_56])).
cnf(c_0_87, plain, (r1_orders_2(X1,X2,X3)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_lattice3(X1,k2_tarski(X2,X2),X3)), inference(rw,[status(thm)],[c_0_79, c_0_56])).
cnf(c_0_88, negated_conjecture, (r2_lattice3(esk1_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_80]), c_0_27])])).
cnf(c_0_89, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_90, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_82])).
cnf(c_0_91, plain, (r2_hidden(esk8_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_92, plain, ![X56]:(v2_struct_0(X56)|~l1_struct_0(X56)|~v1_xboole_0(u1_struct_0(X56))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_93, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_94, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_95, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k2_yellow_0(X1,X2),k6_domain_1(u1_struct_0(X1),k2_yellow_0(X1,X2)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_96, negated_conjecture, (k2_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_27])]), c_0_31])).
cnf(c_0_97, negated_conjecture, (k1_yellow_0(esk1_0,X1)=X2|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,X2,k1_yellow_0(esk1_0,X1))|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_52]), c_0_27])])).
cnf(c_0_98, negated_conjecture, (r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,k2_tarski(X1,X1)))|~r1_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(k1_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_27])])).
cnf(c_0_99, plain, (k2_tarski(esk8_3(X1,X2,X3),esk8_3(X1,X2,X3))=k6_domain_1(u1_struct_0(X1),esk8_3(X1,X2,X3))|v1_xboole_0(u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_76, c_0_89])).
cnf(c_0_100, plain, (esk8_3(X1,k2_tarski(X2,X2),X3)=X2|r2_lattice3(X1,k2_tarski(X2,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_101, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_102, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_76, c_0_93])).
cnf(c_0_103, negated_conjecture, (r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,X2))|~r2_hidden(X1,X2)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(k1_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_88]), c_0_27])])).
cnf(c_0_104, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(X1,k6_domain_1(u1_struct_0(esk1_0),X1))|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_27])])).
cnf(c_0_105, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r1_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(k1_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_106, plain, (r2_lattice3(X1,k1_tarski(X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_107, plain, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(X2),X1)|v1_xboole_0(u1_struct_0(X2))|r2_lattice3(X2,k2_tarski(X1,X1),X3)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_108, negated_conjecture, (k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_109, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_28])).
cnf(c_0_110, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_orders_2(esk1_0,X1,k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1)))|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))|~m1_subset_1(k1_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_111, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r2_lattice3(esk1_0,k2_tarski(X1,X1),X1)|~r1_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_46]), c_0_27])])).
cnf(c_0_112, plain, (r2_lattice3(X1,k2_tarski(X2,X2),X3)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(rw,[status(thm)],[c_0_106, c_0_56])).
cnf(c_0_113, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|v1_xboole_0(u1_struct_0(esk1_0))|r2_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_93]), c_0_27])])).
fof(c_0_114, plain, ![X57]:(~l1_orders_2(X57)|l1_struct_0(X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_115, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|~l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_116, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk2_0,k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~m1_subset_1(k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_109]), c_0_93])])).
cnf(c_0_117, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(X1,X1))=X1|~r1_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_27])]), c_0_31])).
cnf(c_0_118, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk8_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_119, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|r2_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)|~l1_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_113]), c_0_28])).
cnf(c_0_120, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_121, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_96]), c_0_93])])).
cnf(c_0_122, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_93])])).
cnf(c_0_123, plain, (r2_lattice3(X1,k2_tarski(X2,X2),X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_118, c_0_100])).
cnf(c_0_124, negated_conjecture, (k2_tarski(X1,X1)=k6_domain_1(u1_struct_0(esk1_0),X1)|r2_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_27])])).
cnf(c_0_125, negated_conjecture, (k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))!=esk2_0|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_120]), c_0_27])])).
cnf(c_0_126, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)|~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_122]), c_0_28])).
fof(c_0_127, plain, ![X58, X59, X60, X61, X62]:(((r2_lattice3(X58,X60,X59)|(X59!=k1_yellow_0(X58,X60)|~r1_yellow_0(X58,X60))|~m1_subset_1(X59,u1_struct_0(X58))|(~v5_orders_2(X58)|~l1_orders_2(X58)))&(~m1_subset_1(X61,u1_struct_0(X58))|(~r2_lattice3(X58,X60,X61)|r1_orders_2(X58,X59,X61))|(X59!=k1_yellow_0(X58,X60)|~r1_yellow_0(X58,X60))|~m1_subset_1(X59,u1_struct_0(X58))|(~v5_orders_2(X58)|~l1_orders_2(X58))))&(((X59=k1_yellow_0(X58,X62)|(m1_subset_1(esk11_3(X58,X59,X62),u1_struct_0(X58))|~r2_lattice3(X58,X62,X59))|~m1_subset_1(X59,u1_struct_0(X58))|(~v5_orders_2(X58)|~l1_orders_2(X58)))&(r1_yellow_0(X58,X62)|(m1_subset_1(esk11_3(X58,X59,X62),u1_struct_0(X58))|~r2_lattice3(X58,X62,X59))|~m1_subset_1(X59,u1_struct_0(X58))|(~v5_orders_2(X58)|~l1_orders_2(X58))))&(((X59=k1_yellow_0(X58,X62)|(r2_lattice3(X58,X62,esk11_3(X58,X59,X62))|~r2_lattice3(X58,X62,X59))|~m1_subset_1(X59,u1_struct_0(X58))|(~v5_orders_2(X58)|~l1_orders_2(X58)))&(r1_yellow_0(X58,X62)|(r2_lattice3(X58,X62,esk11_3(X58,X59,X62))|~r2_lattice3(X58,X62,X59))|~m1_subset_1(X59,u1_struct_0(X58))|(~v5_orders_2(X58)|~l1_orders_2(X58))))&((X59=k1_yellow_0(X58,X62)|(~r1_orders_2(X58,X59,esk11_3(X58,X59,X62))|~r2_lattice3(X58,X62,X59))|~m1_subset_1(X59,u1_struct_0(X58))|(~v5_orders_2(X58)|~l1_orders_2(X58)))&(r1_yellow_0(X58,X62)|(~r1_orders_2(X58,X59,esk11_3(X58,X59,X62))|~r2_lattice3(X58,X62,X59))|~m1_subset_1(X59,u1_struct_0(X58))|(~v5_orders_2(X58)|~l1_orders_2(X58))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_yellow_0])])])])])])).
cnf(c_0_128, negated_conjecture, (r2_lattice3(X1,k6_domain_1(u1_struct_0(esk1_0),X2),X3)|r2_lattice3(esk1_0,k2_tarski(X2,X2),esk2_0)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_123, c_0_124])).
fof(c_0_129, plain, ![X3, X2, X1]:(epred2_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
cnf(c_0_130, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk2_0)|~m1_subset_1(k1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_116]), c_0_93])]), c_0_125])).
cnf(c_0_131, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_120]), c_0_27])])).
cnf(c_0_132, plain, (r1_yellow_0(X1,X2)|r2_lattice3(X1,X2,esk11_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_133, plain, (r1_yellow_0(X1,X2)|m1_subset_1(esk11_3(X1,X3,X2),u1_struct_0(X1))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_134, negated_conjecture, (r2_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1),esk2_0)|r2_lattice3(esk1_0,k2_tarski(X1,X1),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_93]), c_0_27])])).
fof(c_0_135, plain, ![X3, X2, X1]:(epred2_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_129])).
cnf(c_0_136, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_117]), c_0_93])]), c_0_131])).
cnf(c_0_137, plain, (r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk11_3(X1,X3,X2))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_138, plain, (r1_yellow_0(X1,X2)|r1_orders_2(X1,X3,esk11_3(X1,X4,X2))|~r2_hidden(X3,X2)|~r2_lattice3(X1,X2,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_132]), c_0_133])).
cnf(c_0_139, negated_conjecture, (r2_lattice3(esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk2_0)|r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_31]), c_0_93])])).
fof(c_0_140, plain, ![X73, X74, X75, X76, X77]:(((r1_orders_2(X75,X76,X74)|(X76!=k11_lattice3(X75,X74,X73)|~r2_yellow_0(X75,k2_tarski(X74,X73)))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73))&((r1_orders_2(X75,X76,X73)|(X76!=k11_lattice3(X75,X74,X73)|~r2_yellow_0(X75,k2_tarski(X74,X73)))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73))&(~m1_subset_1(X77,u1_struct_0(X75))|(~r1_orders_2(X75,X77,X74)|~r1_orders_2(X75,X77,X73)|r1_orders_2(X75,X77,X76))|(X76!=k11_lattice3(X75,X74,X73)|~r2_yellow_0(X75,k2_tarski(X74,X73)))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73))))&(((X76=k11_lattice3(X75,X74,X73)|(m1_subset_1(esk13_4(X73,X74,X75,X76),u1_struct_0(X75))|~r1_orders_2(X75,X76,X73)|~r1_orders_2(X75,X76,X74))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73))&(r2_yellow_0(X75,k2_tarski(X74,X73))|(m1_subset_1(esk13_4(X73,X74,X75,X76),u1_struct_0(X75))|~r1_orders_2(X75,X76,X73)|~r1_orders_2(X75,X76,X74))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73)))&((((X76=k11_lattice3(X75,X74,X73)|(r1_orders_2(X75,esk13_4(X73,X74,X75,X76),X74)|~r1_orders_2(X75,X76,X73)|~r1_orders_2(X75,X76,X74))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73))&(r2_yellow_0(X75,k2_tarski(X74,X73))|(r1_orders_2(X75,esk13_4(X73,X74,X75,X76),X74)|~r1_orders_2(X75,X76,X73)|~r1_orders_2(X75,X76,X74))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73)))&((X76=k11_lattice3(X75,X74,X73)|(r1_orders_2(X75,esk13_4(X73,X74,X75,X76),X73)|~r1_orders_2(X75,X76,X73)|~r1_orders_2(X75,X76,X74))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73))&(r2_yellow_0(X75,k2_tarski(X74,X73))|(r1_orders_2(X75,esk13_4(X73,X74,X75,X76),X73)|~r1_orders_2(X75,X76,X73)|~r1_orders_2(X75,X76,X74))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73))))&((X76=k11_lattice3(X75,X74,X73)|(~r1_orders_2(X75,esk13_4(X73,X74,X75,X76),X76)|~r1_orders_2(X75,X76,X73)|~r1_orders_2(X75,X76,X74))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73))&(r2_yellow_0(X75,k2_tarski(X74,X73))|(~r1_orders_2(X75,esk13_4(X73,X74,X75,X76),X76)|~r1_orders_2(X75,X76,X73)|~r1_orders_2(X75,X76,X74))|~m1_subset_1(X76,u1_struct_0(X75))|~epred2_3(X75,X74,X73)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_135])])])])])).
cnf(c_0_141, negated_conjecture, (~l1_struct_0(esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_136]), c_0_28])).
cnf(c_0_142, plain, (r1_yellow_0(X1,X2)|~r2_hidden(X3,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_143, negated_conjecture, (r2_lattice3(esk1_0,k2_tarski(esk2_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_139, c_0_124])).
cnf(c_0_144, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~r1_orders_2(X1,esk13_4(X3,X2,X1,X4),X4)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_145, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk13_4(X3,X2,X1,X4),X3)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred2_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_146, negated_conjecture, (~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_120]), c_0_27])])).
cnf(c_0_147, negated_conjecture, (r1_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_83]), c_0_93]), c_0_27]), c_0_39])])).
cnf(c_0_148, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~epred2_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_149, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_143]), c_0_93]), c_0_27])])).
fof(c_0_150, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred2_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_129])).
cnf(c_0_151, negated_conjecture, (~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_147])])).
cnf(c_0_152, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~epred2_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_93])])).
fof(c_0_153, plain, ![X64, X65, X66]:(~v5_orders_2(X64)|~l1_orders_2(X64)|(~m1_subset_1(X65,u1_struct_0(X64))|(~m1_subset_1(X66,u1_struct_0(X64))|epred2_3(X64,X65,X66)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_150])])])).
cnf(c_0_154, negated_conjecture, (~epred2_3(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_149])])).
cnf(c_0_155, plain, (epred2_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_156, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_93]), c_0_27]), c_0_39])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 157
# Proof object clause steps            : 114
# Proof object formula steps           : 43
# Proof object conjectures             : 58
# Proof object clause conjectures      : 55
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 40
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 60
# Proof object simplifying inferences  : 108
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 85
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 84
# Processed clauses                    : 2242
# ...of these trivial                  : 11
# ...subsumed                          : 1193
# ...remaining for further processing  : 1038
# Other redundant clauses eliminated   : 291
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 196
# Backward-rewritten                   : 72
# Generated clauses                    : 20961
# ...of the previous two non-trivial   : 18387
# Contextual simplify-reflections      : 143
# Paramodulations                      : 20632
# Factorizations                       : 39
# Equation resolutions                 : 291
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 674
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 661
# Current number of unprocessed clauses: 15988
# ...number of literals in the above   : 137868
# Current number of archived formulas  : 0
# Current number of archived clauses   : 351
# Clause-clause subsumption calls (NU) : 81135
# Rec. Clause-clause subsumption calls : 6402
# Non-unit clause-clause subsumptions  : 1455
# Unit Clause-clause subsumption calls : 1432
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 657245

# -------------------------------------------------
# User time                : 0.443 s
# System time              : 0.003 s
# Total time               : 0.446 s
# Maximum resident set size: 28688 pages
