# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t22_relset_1, conjecture, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X4,X5),X3))))<=>k1_relset_1(X2,X1,X3)=X2)), file('data/MPTP/proofs/a06fa9576f204913b7f7e3395387e330.E_input', t22_relset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/a06fa9576f204913b7f7e3395387e330.E_input', t3_subset)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/a06fa9576f204913b7f7e3395387e330.E_input', redefinition_k1_relset_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/a06fa9576f204913b7f7e3395387e330.E_input', d4_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/a06fa9576f204913b7f7e3395387e330.E_input', t28_xboole_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/a06fa9576f204913b7f7e3395387e330.E_input', d4_relat_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/a06fa9576f204913b7f7e3395387e330.E_input', l54_zfmisc_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~(r2_hidden(k4_tarski(X4,X5),X3))))<=>k1_relset_1(X2,X1,X3)=X2))), inference(assume_negation,[status(cth)],[t22_relset_1])).
fof(c_0_8, plain, ![X13, X14]:((~m1_subset_1(X13,k1_zfmisc_1(X14))|r1_tarski(X13,X14))&(~r1_tarski(X13,X14)|m1_subset_1(X13,k1_zfmisc_1(X14)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_9, negated_conjecture, ![X10, X11]:(m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))&(((r2_hidden(esk4_0,esk2_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)!=esk2_0)&(~r2_hidden(k4_tarski(esk4_0,X10),esk3_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)!=esk2_0))&(~r2_hidden(X11,esk2_0)|r2_hidden(k4_tarski(X11,esk5_1(X11)),esk3_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)=esk2_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])])])).
fof(c_0_10, plain, ![X15, X16, X17]:(~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(X15,X16)))|k1_relset_1(X15,X16,X17)=k1_relat_1(X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_11, plain, ![X18, X19, X20, X21, X22, X23, X24, X25]:((((r2_hidden(X21,X18)|~r2_hidden(X21,X20)|X20!=k3_xboole_0(X18,X19))&(r2_hidden(X21,X19)|~r2_hidden(X21,X20)|X20!=k3_xboole_0(X18,X19)))&(~r2_hidden(X22,X18)|~r2_hidden(X22,X19)|r2_hidden(X22,X20)|X20!=k3_xboole_0(X18,X19)))&((~r2_hidden(esk6_3(X23,X24,X25),X25)|(~r2_hidden(esk6_3(X23,X24,X25),X23)|~r2_hidden(esk6_3(X23,X24,X25),X24))|X25=k3_xboole_0(X23,X24))&((r2_hidden(esk6_3(X23,X24,X25),X23)|r2_hidden(esk6_3(X23,X24,X25),X25)|X25=k3_xboole_0(X23,X24))&(r2_hidden(esk6_3(X23,X24,X25),X24)|r2_hidden(esk6_3(X23,X24,X25),X25)|X25=k3_xboole_0(X23,X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_12, plain, ![X27, X28]:(~r1_tarski(X27,X28)|k3_xboole_0(X27,X28)=X27), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_13, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19, negated_conjecture, (~r2_hidden(k4_tarski(esk4_0,X1),esk3_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20, negated_conjecture, (k1_relset_1(esk2_0,esk1_0,esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_15, c_0_14])).
fof(c_0_21, plain, ![X29, X30, X31, X33, X34, X35, X36, X38]:(((~r2_hidden(X31,X30)|r2_hidden(k4_tarski(X31,esk7_3(X29,X30,X31)),X29)|X30!=k1_relat_1(X29))&(~r2_hidden(k4_tarski(X33,X34),X29)|r2_hidden(X33,X30)|X30!=k1_relat_1(X29)))&((~r2_hidden(esk8_2(X35,X36),X36)|~r2_hidden(k4_tarski(esk8_2(X35,X36),X38),X35)|X36=k1_relat_1(X35))&(r2_hidden(esk8_2(X35,X36),X36)|r2_hidden(k4_tarski(esk8_2(X35,X36),esk9_2(X35,X36)),X35)|X36=k1_relat_1(X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_22, plain, ![X40, X41, X42, X43]:(((r2_hidden(X40,X42)|~r2_hidden(k4_tarski(X40,X41),k2_zfmisc_1(X42,X43)))&(r2_hidden(X41,X43)|~r2_hidden(k4_tarski(X40,X41),k2_zfmisc_1(X42,X43))))&(~r2_hidden(X40,X42)|~r2_hidden(X41,X43)|r2_hidden(k4_tarski(X40,X41),k2_zfmisc_1(X42,X43)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (k3_xboole_0(esk3_0,k2_zfmisc_1(esk2_0,esk1_0))=esk3_0), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25, negated_conjecture, (k1_relat_1(esk3_0)!=esk2_0|~r2_hidden(k4_tarski(esk4_0,X1),esk3_0)), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, plain, (r2_hidden(k4_tarski(X1,esk7_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (r2_hidden(esk4_0,esk2_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk2_0,esk1_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k4_tarski(X1,esk5_1(X1)),esk3_0)|k1_relset_1(esk2_0,esk1_0,esk3_0)=esk2_0|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31, negated_conjecture, (k1_relat_1(esk3_0)!=esk2_0|X1!=k1_relat_1(esk3_0)|~r2_hidden(esk4_0,X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk4_0,esk2_0)|k1_relat_1(esk3_0)!=esk2_0), inference(rw,[status(thm)],[c_0_27, c_0_20])).
cnf(c_0_33, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(k4_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_35, plain, (X2=k1_relat_1(X1)|~r2_hidden(esk8_2(X1,X2),X2)|~r2_hidden(k4_tarski(esk8_2(X1,X2),X3),X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_36, negated_conjecture, (k1_relat_1(esk3_0)=esk2_0|r2_hidden(k4_tarski(X1,esk5_1(X1)),esk3_0)|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[c_0_30, c_0_20])).
cnf(c_0_37, negated_conjecture, (k1_relat_1(esk3_0)!=esk2_0), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, negated_conjecture, (X1=k1_relat_1(esk3_0)|r2_hidden(esk8_2(esk3_0,X1),esk2_0)|r2_hidden(esk8_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, negated_conjecture, (X1=k1_relat_1(esk3_0)|~r2_hidden(esk8_2(esk3_0,X1),esk2_0)|~r2_hidden(esk8_2(esk3_0,X1),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_40, negated_conjecture, (r2_hidden(esk8_2(esk3_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(ef,[status(thm)],[c_0_38]), c_0_37])).
cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_40])]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 27
# Proof object formula steps           : 15
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 181
# ...of these trivial                  : 4
# ...subsumed                          : 47
# ...remaining for further processing  : 130
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 1023
# ...of the previous two non-trivial   : 998
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1004
# Factorizations                       : 8
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 106
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 100
# Current number of unprocessed clauses: 831
# ...number of literals in the above   : 3214
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 2208
# Rec. Clause-clause subsumption calls : 1772
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17508

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.003 s
# Total time               : 0.018 s
# Maximum resident set size: 28668 pages
