# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t51_relset_1, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>![X6, X7]:(r2_hidden(k4_tarski(X6,X7),k4_relset_1(X1,X2,X2,X3,X4,X5))<=>?[X8]:(m1_subset_1(X8,X2)&(r2_hidden(k4_tarski(X6,X8),X4)&r2_hidden(k4_tarski(X8,X7),X5))))))))), file('data/MPTP/proofs/a154ac245c454959aa0b6f4d75d96bb4.E_input', t51_relset_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/a154ac245c454959aa0b6f4d75d96bb4.E_input', l3_subset_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/a154ac245c454959aa0b6f4d75d96bb4.E_input', l54_zfmisc_1)).
fof(redefinition_k4_relset_1, axiom, ![X1, X2, X3, X4, X5, X6]:((m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k4_relset_1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/a154ac245c454959aa0b6f4d75d96bb4.E_input', redefinition_k4_relset_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/a154ac245c454959aa0b6f4d75d96bb4.E_input', t1_subset)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/a154ac245c454959aa0b6f4d75d96bb4.E_input', d8_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/a154ac245c454959aa0b6f4d75d96bb4.E_input', dt_k5_relat_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/a154ac245c454959aa0b6f4d75d96bb4.E_input', cc1_relset_1)).
fof(c_0_8, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>![X6, X7]:(r2_hidden(k4_tarski(X6,X7),k4_relset_1(X1,X2,X2,X3,X4,X5))<=>?[X8]:(m1_subset_1(X8,X2)&(r2_hidden(k4_tarski(X6,X8),X4)&r2_hidden(k4_tarski(X8,X7),X5)))))))))), inference(assume_negation,[status(cth)],[t51_relset_1])).
fof(c_0_9, plain, ![X18, X19, X20]:(~m1_subset_1(X19,k1_zfmisc_1(X18))|(~r2_hidden(X20,X19)|r2_hidden(X20,X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_10, negated_conjecture, ![X16]:(~v1_xboole_0(esk1_0)&(~v1_xboole_0(esk2_0)&(~v1_xboole_0(esk3_0)&(m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))&(m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))&((~r2_hidden(k4_tarski(esk6_0,esk7_0),k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0))|(~m1_subset_1(X16,esk2_0)|(~r2_hidden(k4_tarski(esk6_0,X16),esk4_0)|~r2_hidden(k4_tarski(X16,esk7_0),esk5_0))))&((m1_subset_1(esk8_0,esk2_0)|r2_hidden(k4_tarski(esk6_0,esk7_0),k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0)))&((r2_hidden(k4_tarski(esk6_0,esk8_0),esk4_0)|r2_hidden(k4_tarski(esk6_0,esk7_0),k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0)))&(r2_hidden(k4_tarski(esk8_0,esk7_0),esk5_0)|r2_hidden(k4_tarski(esk6_0,esk7_0),k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
fof(c_0_11, plain, ![X21, X22, X23, X24]:(((r2_hidden(X21,X23)|~r2_hidden(k4_tarski(X21,X22),k2_zfmisc_1(X23,X24)))&(r2_hidden(X22,X24)|~r2_hidden(k4_tarski(X21,X22),k2_zfmisc_1(X23,X24))))&(~r2_hidden(X21,X23)|~r2_hidden(X22,X24)|r2_hidden(k4_tarski(X21,X22),k2_zfmisc_1(X23,X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_12, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X25, X26, X27, X28, X29, X30]:(~m1_subset_1(X29,k1_zfmisc_1(k2_zfmisc_1(X25,X26)))|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X27,X28)))|k4_relset_1(X25,X26,X27,X28,X29,X30)=k5_relat_1(X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_relset_1])])).
cnf(c_0_15, negated_conjecture, (~r2_hidden(k4_tarski(esk6_0,esk7_0),k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0))|~m1_subset_1(X1,esk2_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)|~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, (r2_hidden(k4_tarski(esk8_0,esk7_0),esk5_0)|r2_hidden(k4_tarski(esk6_0,esk7_0),k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, plain, ![X31, X32]:(~r2_hidden(X31,X32)|m1_subset_1(X31,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk2_0,esk3_0))|~r2_hidden(X1,esk5_0)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_20, plain, ![X33, X34, X35, X36, X37, X39, X40, X41, X44]:((((r2_hidden(k4_tarski(X36,esk9_5(X33,X34,X35,X36,X37)),X33)|~r2_hidden(k4_tarski(X36,X37),X35)|X35!=k5_relat_1(X33,X34)|~v1_relat_1(X35)|~v1_relat_1(X34)|~v1_relat_1(X33))&(r2_hidden(k4_tarski(esk9_5(X33,X34,X35,X36,X37),X37),X34)|~r2_hidden(k4_tarski(X36,X37),X35)|X35!=k5_relat_1(X33,X34)|~v1_relat_1(X35)|~v1_relat_1(X34)|~v1_relat_1(X33)))&(~r2_hidden(k4_tarski(X39,X41),X33)|~r2_hidden(k4_tarski(X41,X40),X34)|r2_hidden(k4_tarski(X39,X40),X35)|X35!=k5_relat_1(X33,X34)|~v1_relat_1(X35)|~v1_relat_1(X34)|~v1_relat_1(X33)))&((~r2_hidden(k4_tarski(esk10_3(X33,X34,X35),esk11_3(X33,X34,X35)),X35)|(~r2_hidden(k4_tarski(esk10_3(X33,X34,X35),X44),X33)|~r2_hidden(k4_tarski(X44,esk11_3(X33,X34,X35)),X34))|X35=k5_relat_1(X33,X34)|~v1_relat_1(X35)|~v1_relat_1(X34)|~v1_relat_1(X33))&((r2_hidden(k4_tarski(esk10_3(X33,X34,X35),esk12_3(X33,X34,X35)),X33)|r2_hidden(k4_tarski(esk10_3(X33,X34,X35),esk11_3(X33,X34,X35)),X35)|X35=k5_relat_1(X33,X34)|~v1_relat_1(X35)|~v1_relat_1(X34)|~v1_relat_1(X33))&(r2_hidden(k4_tarski(esk12_3(X33,X34,X35),esk11_3(X33,X34,X35)),X34)|r2_hidden(k4_tarski(esk10_3(X33,X34,X35),esk11_3(X33,X34,X35)),X35)|X35=k5_relat_1(X33,X34)|~v1_relat_1(X35)|~v1_relat_1(X34)|~v1_relat_1(X33))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
fof(c_0_21, plain, ![X46, X47]:(~v1_relat_1(X46)|~v1_relat_1(X47)|v1_relat_1(k5_relat_1(X46,X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
fof(c_0_22, plain, ![X48, X49, X50]:(~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(X48,X49)))|v1_relat_1(X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_23, plain, (k4_relset_1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, negated_conjecture, (r2_hidden(k4_tarski(esk8_0,esk7_0),esk5_0)|~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)|~m1_subset_1(X1,esk2_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_25, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(k4_tarski(X1,X2),esk5_0)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(esk9_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (k4_relset_1(X1,X2,esk2_0,esk3_0,X3,esk5_0)=k5_relat_1(X3,esk5_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_23, c_0_13])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_32, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk8_0),esk4_0)|r2_hidden(k4_tarski(esk6_0,esk7_0),k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_33, negated_conjecture, (r2_hidden(k4_tarski(esk8_0,esk7_0),esk5_0)|~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(esk9_5(X1,X2,k5_relat_1(X1,X2),X3,X4),X4),X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X1,X2))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_27]), c_0_28])).
cnf(c_0_35, negated_conjecture, (v1_relat_1(esk5_0)), inference(spm,[status(thm)],[c_0_29, c_0_13])).
cnf(c_0_36, plain, (r2_hidden(k4_tarski(X1,esk9_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_37, negated_conjecture, (k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0)=k5_relat_1(esk4_0,esk5_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk8_0),esk4_0)|~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)|~m1_subset_1(X1,esk2_0)), inference(spm,[status(thm)],[c_0_15, c_0_32])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(esk8_0,esk2_0)|r2_hidden(k4_tarski(esk6_0,esk7_0),k4_relset_1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40, plain, (r2_hidden(k4_tarski(X1,X4),X6)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X2,X4),X5)|X6!=k5_relat_1(X3,X5)|~v1_relat_1(X6)|~v1_relat_1(X5)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_41, negated_conjecture, (r2_hidden(k4_tarski(esk8_0,esk7_0),esk5_0)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(esk6_0,esk9_5(X1,esk5_0,k5_relat_1(X1,esk5_0),X2,esk7_0)),esk4_0)|~r2_hidden(k4_tarski(X2,esk7_0),k5_relat_1(X1,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_42, plain, (r2_hidden(k4_tarski(X1,esk9_5(X2,X3,k5_relat_1(X2,X3),X1,X4)),X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,X4),k5_relat_1(X2,X3))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_36]), c_0_28])).
cnf(c_0_43, negated_conjecture, (v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_29, c_0_31])).
cnf(c_0_44, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk7_0),k5_relat_1(esk4_0,esk5_0))|r2_hidden(k4_tarski(esk8_0,esk7_0),esk5_0)), inference(rw,[status(thm)],[c_0_16, c_0_37])).
cnf(c_0_45, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk8_0),esk4_0)|~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_25]), c_0_26])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk8_0,esk2_0)|~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)|~m1_subset_1(X1,esk2_0)), inference(spm,[status(thm)],[c_0_15, c_0_39])).
cnf(c_0_47, plain, (r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,X4))|~v1_relat_1(X4)|~v1_relat_1(X3)|~r2_hidden(k4_tarski(X5,X2),X4)|~r2_hidden(k4_tarski(X1,X5),X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_40]), c_0_28])).
cnf(c_0_48, negated_conjecture, (r2_hidden(k4_tarski(esk8_0,esk7_0),esk5_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_35])]), c_0_44])).
cnf(c_0_49, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk8_0),esk4_0)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(esk6_0,esk9_5(X1,esk5_0,k5_relat_1(X1,esk5_0),X2,esk7_0)),esk4_0)|~r2_hidden(k4_tarski(X2,esk7_0),k5_relat_1(X1,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_34]), c_0_35])])).
cnf(c_0_50, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk7_0),k5_relat_1(esk4_0,esk5_0))|r2_hidden(k4_tarski(esk6_0,esk8_0),esk4_0)), inference(rw,[status(thm)],[c_0_32, c_0_37])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk8_0,esk2_0)|~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_25]), c_0_26])).
cnf(c_0_52, negated_conjecture, (~r2_hidden(k4_tarski(esk6_0,esk7_0),k5_relat_1(esk4_0,esk5_0))|~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)|~m1_subset_1(X1,esk2_0)), inference(rw,[status(thm)],[c_0_15, c_0_37])).
cnf(c_0_53, negated_conjecture, (r2_hidden(k4_tarski(X1,esk7_0),k5_relat_1(X2,esk5_0))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,esk8_0),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_35])])).
cnf(c_0_54, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk8_0),esk4_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_42]), c_0_43]), c_0_35])]), c_0_50])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(esk8_0,esk2_0)|~v1_relat_1(X1)|~r2_hidden(k4_tarski(esk6_0,esk9_5(X1,esk5_0,k5_relat_1(X1,esk5_0),X2,esk7_0)),esk4_0)|~r2_hidden(k4_tarski(X2,esk7_0),k5_relat_1(X1,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_34]), c_0_35])])).
cnf(c_0_56, negated_conjecture, (r2_hidden(k4_tarski(esk6_0,esk7_0),k5_relat_1(esk4_0,esk5_0))|m1_subset_1(esk8_0,esk2_0)), inference(rw,[status(thm)],[c_0_39, c_0_37])).
cnf(c_0_57, negated_conjecture, (~r2_hidden(k4_tarski(X1,esk7_0),esk5_0)|~r2_hidden(k4_tarski(esk6_0,X1),esk4_0)|~m1_subset_1(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_43]), c_0_54])])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(esk8_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_42]), c_0_43]), c_0_35])]), c_0_56])).
cnf(c_0_59, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_48]), c_0_54])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 43
# Proof object formula steps           : 17
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 23
# Processed clauses                    : 107
# ...of these trivial                  : 0
# ...subsumed                          : 9
# ...remaining for further processing  : 98
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 16
# Generated clauses                    : 118
# ...of the previous two non-trivial   : 112
# Contextual simplify-reflections      : 11
# Paramodulations                      : 115
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 34
# Current number of unprocessed clauses: 51
# ...number of literals in the above   : 256
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 890
# Rec. Clause-clause subsumption calls : 421
# Non-unit clause-clause subsumptions  : 24
# Unit Clause-clause subsumption calls : 44
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4805

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28732 pages
