# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_yellow_1, conjecture, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(k13_lattice3(k3_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3)&k12_lattice3(k3_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3)))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', t17_yellow_1)).
fof(t4_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1)), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', t4_yellow_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', t12_setfam_1)).
fof(fc7_yellow_1, axiom, ![X1]:(~(v2_struct_0(k3_yellow_1(X1)))&(v1_orders_2(k3_yellow_1(X1))&(v3_orders_2(k3_yellow_1(X1))&(v4_orders_2(k3_yellow_1(X1))&v5_orders_2(k3_yellow_1(X1)))))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', fc7_yellow_1)).
fof(dt_k3_yellow_1, axiom, ![X1]:(v1_orders_2(k3_yellow_1(X1))&l1_orders_2(k3_yellow_1(X1))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', dt_k3_yellow_1)).
fof(t1_yellow_1, axiom, ![X1]:(u1_struct_0(k2_yellow_1(X1))=X1&u1_orders_2(k2_yellow_1(X1))=k1_yellow_1(X1)), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', t1_yellow_1)).
fof(t9_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r2_hidden(k3_xboole_0(X2,X3),X1)=>k11_lattice3(k2_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3))))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', t9_yellow_1)).
fof(l19_yellow_1, axiom, ![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(r2_hidden(k3_xboole_0(X2,X3),k9_setfam_1(X1))&r2_hidden(k2_xboole_0(X2,X3),k9_setfam_1(X1))))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', l19_yellow_1)).
fof(redefinition_k13_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v1_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', redefinition_k13_lattice3)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', d1_xboole_0)).
fof(redefinition_k12_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', redefinition_k12_lattice3)).
fof(t12_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(![X2, X3]:((r2_hidden(X2,X1)&r2_hidden(X3,X1))=>r2_hidden(k3_xboole_0(X2,X3),X1))=>v2_lattice3(k2_yellow_1(X1)))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', t12_yellow_1)).
fof(t8_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))=>(r2_hidden(k2_xboole_0(X2,X3),X1)=>k10_lattice3(k2_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', t8_yellow_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', d2_subset_1)).
fof(t11_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(![X2, X3]:((r2_hidden(X2,X1)&r2_hidden(X3,X1))=>r2_hidden(k2_xboole_0(X2,X3),X1))=>v1_lattice3(k2_yellow_1(X1)))), file('data/MPTP/proofs/a30bd11e8eff49e6946810e721a2a6f3.E_input', t11_yellow_1)).
fof(c_0_15, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,u1_struct_0(k3_yellow_1(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k3_yellow_1(X1)))=>(k13_lattice3(k3_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3)&k12_lattice3(k3_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3))))), inference(assume_negation,[status(cth)],[t17_yellow_1])).
fof(c_0_16, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(k3_yellow_1(esk1_0)))&(k13_lattice3(k3_yellow_1(esk1_0),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k12_lattice3(k3_yellow_1(esk1_0),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
fof(c_0_17, plain, ![X7]:k3_yellow_1(X7)=k2_yellow_1(k9_setfam_1(X7)), inference(variable_rename,[status(thm)],[t4_yellow_1])).
fof(c_0_18, plain, ![X8, X9]:k1_setfam_1(k2_tarski(X8,X9))=k3_xboole_0(X8,X9), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_19, plain, ![X11]:(~v2_struct_0(k3_yellow_1(X11))&(v1_orders_2(k3_yellow_1(X11))&(v3_orders_2(k3_yellow_1(X11))&(v4_orders_2(k3_yellow_1(X11))&v5_orders_2(k3_yellow_1(X11)))))), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc7_yellow_1])])).
fof(c_0_20, plain, ![X10]:(v1_orders_2(k3_yellow_1(X10))&l1_orders_2(k3_yellow_1(X10))), inference(variable_rename,[status(thm)],[dt_k3_yellow_1])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k3_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (k3_yellow_1(X1)=k2_yellow_1(k9_setfam_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, ![X12]:(u1_struct_0(k2_yellow_1(X12))=X12&u1_orders_2(k2_yellow_1(X12))=k1_yellow_1(X12)), inference(variable_rename,[status(thm)],[t1_yellow_1])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_25, plain, ![X13, X14, X15]:(v1_xboole_0(X13)|(~m1_subset_1(X14,u1_struct_0(k2_yellow_1(X13)))|(~m1_subset_1(X15,u1_struct_0(k2_yellow_1(X13)))|(~r2_hidden(k3_xboole_0(X14,X15),X13)|k11_lattice3(k2_yellow_1(X13),X14,X15)=k3_xboole_0(X14,X15))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_yellow_1])])])])).
fof(c_0_26, plain, ![X16, X17, X18]:((r2_hidden(k3_xboole_0(X17,X18),k9_setfam_1(X16))|~m1_subset_1(X18,u1_struct_0(k3_yellow_1(X16)))|~m1_subset_1(X17,u1_struct_0(k3_yellow_1(X16))))&(r2_hidden(k2_xboole_0(X17,X18),k9_setfam_1(X16))|~m1_subset_1(X18,u1_struct_0(k3_yellow_1(X16)))|~m1_subset_1(X17,u1_struct_0(k3_yellow_1(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l19_yellow_1])])])])).
cnf(c_0_27, negated_conjecture, (k13_lattice3(k3_yellow_1(esk1_0),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k12_lattice3(k3_yellow_1(esk1_0),esk2_0,esk3_0)!=k3_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_29, plain, ![X19, X20, X21]:(~v5_orders_2(X19)|(~v1_lattice3(X19)|~l1_orders_2(X19))|(~m1_subset_1(X20,u1_struct_0(X19))|~m1_subset_1(X21,u1_struct_0(X19)))|k13_lattice3(X19,X20,X21)=k10_lattice3(X19,X20,X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_lattice3])])).
cnf(c_0_30, plain, (v5_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (l1_orders_2(k3_yellow_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_32, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(k2_yellow_1(k9_setfam_1(esk1_0))))), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_33, plain, (u1_struct_0(k2_yellow_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(k2_yellow_1(k9_setfam_1(esk1_0))))), inference(rw,[status(thm)],[c_0_24, c_0_22])).
cnf(c_0_35, plain, (v1_xboole_0(X1)|k11_lattice3(k2_yellow_1(X1),X2,X3)=k3_xboole_0(X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r2_hidden(k3_xboole_0(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_36, plain, ![X22, X23, X24]:((~v1_xboole_0(X22)|~r2_hidden(X23,X22))&(r2_hidden(esk4_1(X24),X24)|v1_xboole_0(X24))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_37, plain, (r2_hidden(k3_xboole_0(X1,X2),k9_setfam_1(X3))|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(X3)))|~m1_subset_1(X1,u1_struct_0(k3_yellow_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, negated_conjecture, (k13_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|k12_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)),esk2_0,esk3_0)!=k1_setfam_1(k2_tarski(esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_22]), c_0_22]), c_0_28])).
cnf(c_0_39, plain, (k13_lattice3(X1,X2,X3)=k10_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v1_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (v5_orders_2(k2_yellow_1(k9_setfam_1(X1)))), inference(rw,[status(thm)],[c_0_30, c_0_22])).
cnf(c_0_41, plain, (l1_orders_2(k2_yellow_1(k9_setfam_1(X1)))), inference(rw,[status(thm)],[c_0_31, c_0_22])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk3_0,k9_setfam_1(esk1_0))), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk2_0,k9_setfam_1(esk1_0))), inference(rw,[status(thm)],[c_0_34, c_0_33])).
fof(c_0_44, plain, ![X26, X27, X28]:(~v5_orders_2(X26)|(~v2_lattice3(X26)|~l1_orders_2(X26))|(~m1_subset_1(X27,u1_struct_0(X26))|~m1_subset_1(X28,u1_struct_0(X26)))|k12_lattice3(X26,X27,X28)=k11_lattice3(X26,X27,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_45, plain, (k11_lattice3(k2_yellow_1(X1),X2,X3)=k1_setfam_1(k2_tarski(X2,X3))|v1_xboole_0(X1)|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~r2_hidden(k1_setfam_1(k2_tarski(X2,X3)),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_28]), c_0_28])).
cnf(c_0_46, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, plain, (r2_hidden(k1_setfam_1(k2_tarski(X1,X2)),k9_setfam_1(X3))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k9_setfam_1(X3))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k9_setfam_1(X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_22]), c_0_22]), c_0_28])).
fof(c_0_48, plain, ![X37]:(((r2_hidden(esk7_1(X37),X37)|v2_lattice3(k2_yellow_1(X37))|v1_xboole_0(X37))&(r2_hidden(esk8_1(X37),X37)|v2_lattice3(k2_yellow_1(X37))|v1_xboole_0(X37)))&(~r2_hidden(k3_xboole_0(esk7_1(X37),esk8_1(X37)),X37)|v2_lattice3(k2_yellow_1(X37))|v1_xboole_0(X37))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_yellow_1])])])])])).
cnf(c_0_49, plain, (r2_hidden(k2_xboole_0(X1,X2),k9_setfam_1(X3))|~m1_subset_1(X2,u1_struct_0(k3_yellow_1(X3)))|~m1_subset_1(X1,u1_struct_0(k3_yellow_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_50, negated_conjecture, (k12_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)),esk2_0,esk3_0)!=k1_setfam_1(k2_tarski(esk2_0,esk3_0))|k10_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|~v1_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41]), c_0_33]), c_0_42]), c_0_33]), c_0_43])])).
cnf(c_0_51, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52, plain, (k11_lattice3(k2_yellow_1(X1),X2,X3)=k1_setfam_1(k2_tarski(X2,X3))|~r2_hidden(k1_setfam_1(k2_tarski(X2,X3)),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_33]), c_0_33]), c_0_46])).
cnf(c_0_53, plain, (r2_hidden(k1_setfam_1(k2_tarski(X1,X2)),k9_setfam_1(X3))|~m1_subset_1(X2,k9_setfam_1(X3))|~m1_subset_1(X1,k9_setfam_1(X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_33]), c_0_33])).
fof(c_0_54, plain, ![X29, X30, X31]:(v1_xboole_0(X29)|(~m1_subset_1(X30,u1_struct_0(k2_yellow_1(X29)))|(~m1_subset_1(X31,u1_struct_0(k2_yellow_1(X29)))|(~r2_hidden(k2_xboole_0(X30,X31),X29)|k10_lattice3(k2_yellow_1(X29),X30,X31)=k2_xboole_0(X30,X31))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_yellow_1])])])])).
cnf(c_0_55, plain, (v2_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)|~r2_hidden(k3_xboole_0(esk7_1(X1),esk8_1(X1)),X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_56, plain, ![X35, X36]:(((~m1_subset_1(X36,X35)|r2_hidden(X36,X35)|v1_xboole_0(X35))&(~r2_hidden(X36,X35)|m1_subset_1(X36,X35)|v1_xboole_0(X35)))&((~m1_subset_1(X36,X35)|v1_xboole_0(X36)|~v1_xboole_0(X35))&(~v1_xboole_0(X36)|m1_subset_1(X36,X35)|~v1_xboole_0(X35)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_57, plain, (r2_hidden(k2_xboole_0(X1,X2),k9_setfam_1(X3))|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(k9_setfam_1(X3))))|~m1_subset_1(X1,u1_struct_0(k2_yellow_1(k9_setfam_1(X3))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_22]), c_0_22])).
cnf(c_0_58, negated_conjecture, (k11_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)),esk2_0,esk3_0)!=k1_setfam_1(k2_tarski(esk2_0,esk3_0))|k10_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|~v2_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))|~v1_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_40]), c_0_41]), c_0_33]), c_0_42]), c_0_33]), c_0_43])])).
cnf(c_0_59, plain, (k11_lattice3(k2_yellow_1(k9_setfam_1(X1)),X2,X3)=k1_setfam_1(k2_tarski(X2,X3))|~m1_subset_1(X3,k9_setfam_1(X1))|~m1_subset_1(X2,k9_setfam_1(X1))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_60, plain, (v1_xboole_0(X1)|k10_lattice3(k2_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3)|~m1_subset_1(X2,u1_struct_0(k2_yellow_1(X1)))|~m1_subset_1(X3,u1_struct_0(k2_yellow_1(X1)))|~r2_hidden(k2_xboole_0(X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, plain, (v1_xboole_0(X1)|v2_lattice3(k2_yellow_1(X1))|~r2_hidden(k1_setfam_1(k2_tarski(esk7_1(X1),esk8_1(X1))),X1)), inference(rw,[status(thm)],[c_0_55, c_0_28])).
cnf(c_0_62, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (r2_hidden(k2_xboole_0(X1,X2),k9_setfam_1(X3))|~m1_subset_1(X2,k9_setfam_1(X3))|~m1_subset_1(X1,k9_setfam_1(X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_33]), c_0_33])).
fof(c_0_64, plain, ![X32]:(((r2_hidden(esk5_1(X32),X32)|v1_lattice3(k2_yellow_1(X32))|v1_xboole_0(X32))&(r2_hidden(esk6_1(X32),X32)|v1_lattice3(k2_yellow_1(X32))|v1_xboole_0(X32)))&(~r2_hidden(k2_xboole_0(esk5_1(X32),esk6_1(X32)),X32)|v1_lattice3(k2_yellow_1(X32))|v1_xboole_0(X32))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t11_yellow_1])])])])])).
cnf(c_0_65, negated_conjecture, (k10_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)),esk2_0,esk3_0)!=k2_xboole_0(esk2_0,esk3_0)|~v2_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))|~v1_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_42]), c_0_43])])).
cnf(c_0_66, plain, (k10_lattice3(k2_yellow_1(X1),X2,X3)=k2_xboole_0(X2,X3)|~r2_hidden(k2_xboole_0(X2,X3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_33]), c_0_33]), c_0_46])).
cnf(c_0_67, plain, (v2_lattice3(k2_yellow_1(X1))|~r2_hidden(k1_setfam_1(k2_tarski(esk7_1(X1),esk8_1(X1))),X1)), inference(csr,[status(thm)],[c_0_61, c_0_46])).
cnf(c_0_68, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_62, c_0_46])).
cnf(c_0_69, plain, (r2_hidden(esk8_1(X1),X1)|v2_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_70, plain, (r2_hidden(esk7_1(X1),X1)|v2_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_71, plain, (~v1_xboole_0(k9_setfam_1(X1))|~m1_subset_1(X2,k9_setfam_1(X1))|~m1_subset_1(X3,k9_setfam_1(X1))), inference(spm,[status(thm)],[c_0_46, c_0_63])).
cnf(c_0_72, plain, (v1_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)|~r2_hidden(k2_xboole_0(esk5_1(X1),esk6_1(X1)),X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_73, negated_conjecture, (~v2_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))|~v1_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))|~r2_hidden(k2_xboole_0(esk2_0,esk3_0),k9_setfam_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_42]), c_0_43])])).
cnf(c_0_74, plain, (v2_lattice3(k2_yellow_1(k9_setfam_1(X1)))|~m1_subset_1(esk8_1(k9_setfam_1(X1)),k9_setfam_1(X1))|~m1_subset_1(esk7_1(k9_setfam_1(X1)),k9_setfam_1(X1))), inference(spm,[status(thm)],[c_0_67, c_0_53])).
cnf(c_0_75, plain, (v2_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)|m1_subset_1(esk8_1(X1),X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76, plain, (v2_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)|m1_subset_1(esk7_1(X1),X1)), inference(spm,[status(thm)],[c_0_68, c_0_70])).
cnf(c_0_77, negated_conjecture, (~v1_xboole_0(k9_setfam_1(esk1_0))|~m1_subset_1(X1,k9_setfam_1(esk1_0))), inference(spm,[status(thm)],[c_0_71, c_0_42])).
cnf(c_0_78, plain, (v1_lattice3(k2_yellow_1(X1))|~r2_hidden(k2_xboole_0(esk5_1(X1),esk6_1(X1)),X1)), inference(csr,[status(thm)],[c_0_72, c_0_46])).
cnf(c_0_79, plain, (r2_hidden(esk6_1(X1),X1)|v1_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_80, plain, (r2_hidden(esk5_1(X1),X1)|v1_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_81, negated_conjecture, (~v2_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))|~v1_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_63]), c_0_42]), c_0_43])])).
cnf(c_0_82, plain, (v2_lattice3(k2_yellow_1(k9_setfam_1(X1)))|v1_xboole_0(k9_setfam_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
cnf(c_0_83, negated_conjecture, (~v1_xboole_0(k9_setfam_1(esk1_0))), inference(spm,[status(thm)],[c_0_77, c_0_42])).
cnf(c_0_84, plain, (v1_lattice3(k2_yellow_1(k9_setfam_1(X1)))|~m1_subset_1(esk6_1(k9_setfam_1(X1)),k9_setfam_1(X1))|~m1_subset_1(esk5_1(k9_setfam_1(X1)),k9_setfam_1(X1))), inference(spm,[status(thm)],[c_0_78, c_0_63])).
cnf(c_0_85, plain, (v1_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)|m1_subset_1(esk6_1(X1),X1)), inference(spm,[status(thm)],[c_0_68, c_0_79])).
cnf(c_0_86, plain, (v1_lattice3(k2_yellow_1(X1))|v1_xboole_0(X1)|m1_subset_1(esk5_1(X1),X1)), inference(spm,[status(thm)],[c_0_68, c_0_80])).
cnf(c_0_87, negated_conjecture, (~v1_lattice3(k2_yellow_1(k9_setfam_1(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])).
cnf(c_0_88, plain, (v1_lattice3(k2_yellow_1(k9_setfam_1(X1)))|v1_xboole_0(k9_setfam_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_83]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 59
# Proof object formula steps           : 31
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 57
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 29
# Processed clauses                    : 103
# ...of these trivial                  : 1
# ...subsumed                          : 18
# ...remaining for further processing  : 84
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 80
# ...of the previous two non-trivial   : 48
# Contextual simplify-reflections      : 7
# Paramodulations                      : 80
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 42
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 7
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 259
# Rec. Clause-clause subsumption calls : 151
# Non-unit clause-clause subsumptions  : 21
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3752

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28720 pages
