# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B02_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t131_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k8_relat_1(X1,X3),k8_relat_1(X2,X3)))), file('data/MPTP/proofs/a35ac0685b3346e6a833814f31a850ce.E_input', t131_relat_1)).
fof(d10_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k6_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>(r2_hidden(X3,X1)&X3=X4)))), file('data/MPTP/proofs/a35ac0685b3346e6a833814f31a850ce.E_input', d10_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/a35ac0685b3346e6a833814f31a850ce.E_input', d3_tarski)).
fof(t73_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(![X3]:(r2_hidden(X3,X1)=>r2_hidden(k4_tarski(X3,X3),X2))=>r1_tarski(k6_relat_1(X1),X2))), file('data/MPTP/proofs/a35ac0685b3346e6a833814f31a850ce.E_input', t73_relat_1)).
fof(t50_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>![X4]:(v1_relat_1(X4)=>((r1_tarski(X1,X2)&r1_tarski(X3,X4))=>r1_tarski(k5_relat_1(X1,X3),k5_relat_1(X2,X4))))))), file('data/MPTP/proofs/a35ac0685b3346e6a833814f31a850ce.E_input', t50_relat_1)).
fof(t123_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k8_relat_1(X1,X2)=k5_relat_1(X2,k6_relat_1(X1))), file('data/MPTP/proofs/a35ac0685b3346e6a833814f31a850ce.E_input', t123_relat_1)).
fof(dt_k6_relat_1, axiom, ![X1]:v1_relat_1(k6_relat_1(X1)), file('data/MPTP/proofs/a35ac0685b3346e6a833814f31a850ce.E_input', dt_k6_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/a35ac0685b3346e6a833814f31a850ce.E_input', d10_xboole_0)).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k8_relat_1(X1,X3),k8_relat_1(X2,X3))))), inference(assume_negation,[status(cth)],[t131_relat_1])).
fof(c_0_9, plain, ![X13, X14, X15, X16, X17, X18]:((((r2_hidden(X15,X13)|~r2_hidden(k4_tarski(X15,X16),X14)|X14!=k6_relat_1(X13)|~v1_relat_1(X14))&(X15=X16|~r2_hidden(k4_tarski(X15,X16),X14)|X14!=k6_relat_1(X13)|~v1_relat_1(X14)))&(~r2_hidden(X17,X13)|X17!=X18|r2_hidden(k4_tarski(X17,X18),X14)|X14!=k6_relat_1(X13)|~v1_relat_1(X14)))&((~r2_hidden(k4_tarski(esk5_2(X13,X14),esk6_2(X13,X14)),X14)|(~r2_hidden(esk5_2(X13,X14),X13)|esk5_2(X13,X14)!=esk6_2(X13,X14))|X14=k6_relat_1(X13)|~v1_relat_1(X14))&((r2_hidden(esk5_2(X13,X14),X13)|r2_hidden(k4_tarski(esk5_2(X13,X14),esk6_2(X13,X14)),X14)|X14=k6_relat_1(X13)|~v1_relat_1(X14))&(esk5_2(X13,X14)=esk6_2(X13,X14)|r2_hidden(k4_tarski(esk5_2(X13,X14),esk6_2(X13,X14)),X14)|X14=k6_relat_1(X13)|~v1_relat_1(X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_relat_1])])])])])])).
fof(c_0_10, plain, ![X28, X29, X30, X31, X32]:((~r1_tarski(X28,X29)|(~r2_hidden(X30,X28)|r2_hidden(X30,X29)))&((r2_hidden(esk7_2(X31,X32),X31)|r1_tarski(X31,X32))&(~r2_hidden(esk7_2(X31,X32),X32)|r1_tarski(X31,X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_11, negated_conjecture, (v1_relat_1(esk3_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k8_relat_1(esk1_0,esk3_0),k8_relat_1(esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)|X1!=X3|X4!=k6_relat_1(X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_15, plain, ![X10, X11]:((r2_hidden(esk4_2(X10,X11),X10)|r1_tarski(k6_relat_1(X10),X11)|~v1_relat_1(X11))&(~r2_hidden(k4_tarski(esk4_2(X10,X11),esk4_2(X10,X11)),X11)|r1_tarski(k6_relat_1(X10),X11)|~v1_relat_1(X11))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_relat_1])])])])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(X1,X1),X2)|X2!=k6_relat_1(X3)|~r2_hidden(X1,X3)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_17, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_18, plain, (r1_tarski(k6_relat_1(X1),X2)|~r2_hidden(k4_tarski(esk4_2(X1,X2),esk4_2(X1,X2)),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (r2_hidden(k4_tarski(X1,X1),X2)|X2!=k6_relat_1(esk2_0)|~r2_hidden(X1,esk1_0)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_21, plain, ![X21, X22, X23, X24]:(~v1_relat_1(X21)|(~v1_relat_1(X22)|(~v1_relat_1(X23)|(~v1_relat_1(X24)|(~r1_tarski(X21,X22)|~r1_tarski(X23,X24)|r1_tarski(k5_relat_1(X21,X23),k5_relat_1(X22,X24))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t50_relat_1])])])).
fof(c_0_22, plain, ![X8, X9]:(~v1_relat_1(X9)|k8_relat_1(X8,X9)=k5_relat_1(X9,k6_relat_1(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t123_relat_1])])).
fof(c_0_23, plain, ![X25]:v1_relat_1(k6_relat_1(X25)), inference(variable_rename,[status(thm)],[dt_k6_relat_1])).
cnf(c_0_24, negated_conjecture, (r1_tarski(k6_relat_1(X1),X2)|X2!=k6_relat_1(esk2_0)|~r2_hidden(esk4_2(X1,X2),esk1_0)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(k6_relat_1(X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X4!=k6_relat_1(X2)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(esk7_2(X1,X2),esk7_2(X1,X2)),X3)|r1_tarski(X1,X2)|X3!=k6_relat_1(X1)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_16, c_0_20])).
cnf(c_0_28, plain, (r1_tarski(k5_relat_1(X1,X3),k5_relat_1(X2,X4))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X4)|~r1_tarski(X1,X2)|~r1_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (k8_relat_1(X2,X1)=k5_relat_1(X1,k6_relat_1(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, plain, ![X26, X27]:(((r1_tarski(X26,X27)|X26!=X27)&(r1_tarski(X27,X26)|X26!=X27))&(~r1_tarski(X26,X27)|~r1_tarski(X27,X26)|X26=X27)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_32, negated_conjecture, (r1_tarski(k6_relat_1(esk1_0),X1)|X1!=k6_relat_1(esk2_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33, plain, (r2_hidden(esk7_2(X1,X2),X3)|r1_tarski(X1,X2)|X4!=k6_relat_1(X3)|X4!=k6_relat_1(X1)|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (r1_tarski(k5_relat_1(X1,X2),k8_relat_1(X3,X4))|~r1_tarski(X2,k6_relat_1(X3))|~r1_tarski(X1,X4)|~v1_relat_1(X2)|~v1_relat_1(X4)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_35, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, negated_conjecture, (r2_hidden(X1,X2)|X2!=k6_relat_1(esk2_0)|~r2_hidden(X1,k6_relat_1(esk1_0))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_13, c_0_32])).
cnf(c_0_37, plain, (r2_hidden(esk7_2(X1,X2),X3)|r1_tarski(X1,X2)|k6_relat_1(X3)!=k6_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_33]), c_0_30])])).
cnf(c_0_38, negated_conjecture, (~r1_tarski(k8_relat_1(esk1_0,esk3_0),k8_relat_1(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_39, plain, (r1_tarski(k8_relat_1(X1,X2),k8_relat_1(X3,X4))|~r1_tarski(k6_relat_1(X1),k6_relat_1(X3))|~r1_tarski(X2,X4)|~v1_relat_1(X4)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_29]), c_0_30])])).
cnf(c_0_40, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_35])).
cnf(c_0_41, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk7_2(X1,X2),X3)|r1_tarski(X1,X2)|k6_relat_1(k6_relat_1(esk1_0))!=k6_relat_1(X1)|X3!=k6_relat_1(esk2_0)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, negated_conjecture, (~r1_tarski(k6_relat_1(esk1_0),k6_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41])])).
cnf(c_0_45, negated_conjecture, (r1_tarski(X1,X2)|k6_relat_1(k6_relat_1(esk1_0))!=k6_relat_1(X1)|X2!=k6_relat_1(esk2_0)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_46, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 30
# Proof object formula steps           : 17
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 1064
# ...of these trivial                  : 113
# ...subsumed                          : 443
# ...remaining for further processing  : 508
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 0
# Generated clauses                    : 7440
# ...of the previous two non-trivial   : 7058
# Contextual simplify-reflections      : 6
# Paramodulations                      : 7371
# Factorizations                       : 0
# Equation resolutions                 : 69
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 472
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 465
# Current number of unprocessed clauses: 6009
# ...number of literals in the above   : 52021
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 115687
# Rec. Clause-clause subsumption calls : 7771
# Non-unit clause-clause subsumptions  : 463
# Unit Clause-clause subsumption calls : 351
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 178600

# -------------------------------------------------
# User time                : 0.125 s
# System time              : 0.000 s
# Total time               : 0.125 s
# Maximum resident set size: 28616 pages
