# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t83_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:~((r2_hidden(X1,k4_zfmisc_1(X2,X3,X4,X5))&![X6, X7, X8, X9]:~((r2_hidden(X6,X2)&(r2_hidden(X7,X3)&(r2_hidden(X8,X4)&(r2_hidden(X9,X5)&X1=k4_mcart_1(X6,X7,X8,X9)))))))), file('data/MPTP/proofs/a4455877986a4164add0710a814ada94.E_input', t83_mcart_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/a4455877986a4164add0710a814ada94.E_input', d4_zfmisc_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/a4455877986a4164add0710a814ada94.E_input', d3_zfmisc_1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/a4455877986a4164add0710a814ada94.E_input', t10_mcart_1)).
fof(d4_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4), file('data/MPTP/proofs/a4455877986a4164add0710a814ada94.E_input', d4_mcart_1)).
fof(t31_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4), file('data/MPTP/proofs/a4455877986a4164add0710a814ada94.E_input', t31_mcart_1)).
fof(l139_zfmisc_1, axiom, ![X1, X2, X3]:~((r2_hidden(X1,k2_zfmisc_1(X2,X3))&![X4, X5]:~(k4_tarski(X4,X5)=X1))), file('data/MPTP/proofs/a4455877986a4164add0710a814ada94.E_input', l139_zfmisc_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/a4455877986a4164add0710a814ada94.E_input', t7_mcart_1)).
fof(c_0_8, negated_conjecture, ~(![X1, X2, X3, X4, X5]:~((r2_hidden(X1,k4_zfmisc_1(X2,X3,X4,X5))&![X6, X7, X8, X9]:~((r2_hidden(X6,X2)&(r2_hidden(X7,X3)&(r2_hidden(X8,X4)&(r2_hidden(X9,X5)&X1=k4_mcart_1(X6,X7,X8,X9))))))))), inference(assume_negation,[status(cth)],[t83_mcart_1])).
fof(c_0_9, plain, ![X19, X20, X21, X22]:k4_zfmisc_1(X19,X20,X21,X22)=k2_zfmisc_1(k3_zfmisc_1(X19,X20,X21),X22), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_10, plain, ![X41, X42, X43]:k3_zfmisc_1(X41,X42,X43)=k2_zfmisc_1(k2_zfmisc_1(X41,X42),X43), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_11, negated_conjecture, ![X15, X16, X17, X18]:(r2_hidden(esk1_0,k4_zfmisc_1(esk2_0,esk3_0,esk4_0,esk5_0))&(~r2_hidden(X15,esk2_0)|(~r2_hidden(X16,esk3_0)|(~r2_hidden(X17,esk4_0)|(~r2_hidden(X18,esk5_0)|esk1_0!=k4_mcart_1(X15,X16,X17,X18)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])])).
cnf(c_0_12, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_14, plain, ![X38, X39, X40]:((r2_hidden(k1_mcart_1(X38),X39)|~r2_hidden(X38,k2_zfmisc_1(X39,X40)))&(r2_hidden(k2_mcart_1(X38),X40)|~r2_hidden(X38,k2_zfmisc_1(X39,X40)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_15, negated_conjecture, (r2_hidden(esk1_0,k4_zfmisc_1(esk2_0,esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),esk5_0))), inference(rw,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_19, plain, ![X29, X30, X31, X32]:k4_mcart_1(X29,X30,X31,X32)=k4_tarski(k3_mcart_1(X29,X30,X31),X32), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_20, plain, ![X23, X24, X25, X26]:k4_mcart_1(X23,X24,X25,X26)=k4_tarski(k4_tarski(k4_tarski(X23,X24),X25),X26), inference(variable_rename,[status(thm)],[t31_mcart_1])).
fof(c_0_21, plain, ![X33, X34, X35]:(~r2_hidden(X33,k2_zfmisc_1(X34,X35))|k4_tarski(esk6_1(X33),esk7_1(X33))=X33), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l139_zfmisc_1])])])])])])).
cnf(c_0_22, negated_conjecture, (r2_hidden(k1_mcart_1(esk1_0),k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X3,esk4_0)|~r2_hidden(X4,esk5_0)|esk1_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X27, X28]:(k1_mcart_1(k4_tarski(X27,X28))=X27&k2_mcart_1(k4_tarski(X27,X28))=X28), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_27, plain, (k4_tarski(esk6_1(X1),esk7_1(X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk1_0)),k2_zfmisc_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_17, c_0_22])).
cnf(c_0_29, negated_conjecture, (esk1_0!=k4_tarski(k3_mcart_1(X1,X2,X3),X4)|~r2_hidden(X4,esk5_0)|~r2_hidden(X3,esk4_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, plain, (k4_tarski(k3_mcart_1(X1,X2,X3),X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_25, c_0_24])).
cnf(c_0_31, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, negated_conjecture, (k4_tarski(esk6_1(k1_mcart_1(k1_mcart_1(esk1_0))),esk7_1(k1_mcart_1(k1_mcart_1(esk1_0))))=k1_mcart_1(k1_mcart_1(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, negated_conjecture, (k4_tarski(esk6_1(k1_mcart_1(esk1_0)),esk7_1(k1_mcart_1(esk1_0)))=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_27, c_0_22])).
cnf(c_0_34, negated_conjecture, (k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)!=esk1_0|~r2_hidden(X4,esk5_0)|~r2_hidden(X3,esk4_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, negated_conjecture, (esk6_1(k1_mcart_1(k1_mcart_1(esk1_0)))=k1_mcart_1(k1_mcart_1(k1_mcart_1(esk1_0)))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_36, negated_conjecture, (esk6_1(k1_mcart_1(esk1_0))=k1_mcart_1(k1_mcart_1(esk1_0))), inference(spm,[status(thm)],[c_0_31, c_0_33])).
cnf(c_0_37, negated_conjecture, (k4_tarski(esk6_1(esk1_0),esk7_1(esk1_0))=esk1_0), inference(spm,[status(thm)],[c_0_27, c_0_18])).
cnf(c_0_38, negated_conjecture, (k4_tarski(k4_tarski(k1_mcart_1(k1_mcart_1(esk1_0)),X1),X2)!=esk1_0|~r2_hidden(esk7_1(k1_mcart_1(k1_mcart_1(esk1_0))),esk3_0)|~r2_hidden(esk6_1(k1_mcart_1(k1_mcart_1(esk1_0))),esk2_0)|~r2_hidden(X2,esk5_0)|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_34, c_0_32])).
cnf(c_0_39, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(k1_mcart_1(esk1_0))),esk2_0)), inference(spm,[status(thm)],[c_0_17, c_0_28])).
cnf(c_0_40, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_41, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(k1_mcart_1(esk1_0))),esk7_1(k1_mcart_1(k1_mcart_1(esk1_0))))=k1_mcart_1(k1_mcart_1(esk1_0))), inference(rw,[status(thm)],[c_0_32, c_0_35])).
cnf(c_0_42, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_43, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk1_0)),esk7_1(k1_mcart_1(esk1_0)))=k1_mcart_1(esk1_0)), inference(rw,[status(thm)],[c_0_33, c_0_36])).
cnf(c_0_44, negated_conjecture, (esk6_1(esk1_0)=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_31, c_0_37])).
cnf(c_0_45, negated_conjecture, (k4_tarski(k4_tarski(k1_mcart_1(k1_mcart_1(esk1_0)),X1),X2)!=esk1_0|~r2_hidden(esk7_1(k1_mcart_1(k1_mcart_1(esk1_0))),esk3_0)|~r2_hidden(X2,esk5_0)|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_35]), c_0_39])])).
cnf(c_0_46, negated_conjecture, (esk7_1(k1_mcart_1(k1_mcart_1(esk1_0)))=k2_mcart_1(k1_mcart_1(k1_mcart_1(esk1_0)))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(k1_mcart_1(esk1_0))),esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_28])).
cnf(c_0_48, negated_conjecture, (esk7_1(k1_mcart_1(esk1_0))=k2_mcart_1(k1_mcart_1(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_43])).
cnf(c_0_49, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),esk7_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_37, c_0_44])).
cnf(c_0_50, negated_conjecture, (k4_tarski(k4_tarski(k1_mcart_1(k1_mcart_1(esk1_0)),X1),X2)!=esk1_0|~r2_hidden(X2,esk5_0)|~r2_hidden(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_51, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk1_0)),k2_mcart_1(k1_mcart_1(esk1_0)))=k1_mcart_1(esk1_0)), inference(rw,[status(thm)],[c_0_43, c_0_48])).
cnf(c_0_52, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk1_0)),esk4_0)), inference(spm,[status(thm)],[c_0_42, c_0_22])).
cnf(c_0_53, negated_conjecture, (esk7_1(esk1_0)=k2_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_40, c_0_49])).
cnf(c_0_54, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),X1)!=esk1_0|~r2_hidden(X1,esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])])).
cnf(c_0_55, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_49, c_0_53])).
cnf(c_0_56, negated_conjecture, (r2_hidden(k2_mcart_1(esk1_0),esk5_0)), inference(spm,[status(thm)],[c_0_42, c_0_18])).
cnf(c_0_57, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 41
# Proof object formula steps           : 17
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 18
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 8
# Processed clauses                    : 62
# ...of these trivial                  : 0
# ...subsumed                          : 8
# ...remaining for further processing  : 54
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 18
# Generated clauses                    : 63
# ...of the previous two non-trivial   : 78
# Contextual simplify-reflections      : 0
# Paramodulations                      : 63
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 9
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 7
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 139
# Rec. Clause-clause subsumption calls : 44
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 15
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2941

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28632 pages
