# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t138_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4)))), file('data/MPTP/proofs/a46c1f857c814b2b80de85d79c7459b7.E_input', t138_zfmisc_1)).
fof(t55_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_xboole_0(k2_tarski(X1,X2),X3)&r2_hidden(X1,X3))), file('data/MPTP/proofs/a46c1f857c814b2b80de85d79c7459b7.E_input', t55_zfmisc_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/a46c1f857c814b2b80de85d79c7459b7.E_input', t65_xboole_1)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/a46c1f857c814b2b80de85d79c7459b7.E_input', d2_zfmisc_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/a46c1f857c814b2b80de85d79c7459b7.E_input', d3_tarski)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/a46c1f857c814b2b80de85d79c7459b7.E_input', t2_tarski)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/a46c1f857c814b2b80de85d79c7459b7.E_input', l54_zfmisc_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3, X4]:(r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=>(k2_zfmisc_1(X1,X2)=k1_xboole_0|(r1_tarski(X1,X3)&r1_tarski(X2,X4))))), inference(assume_negation,[status(cth)],[t138_zfmisc_1])).
fof(c_0_8, plain, ![X11, X12, X13]:(~r1_xboole_0(k2_tarski(X11,X12),X13)|~r2_hidden(X11,X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_zfmisc_1])])).
fof(c_0_9, plain, ![X14]:r1_xboole_0(X14,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_10, plain, ![X15, X16, X17, X18, X21, X22, X23, X24, X25, X26, X28, X29]:((((r2_hidden(esk5_4(X15,X16,X17,X18),X15)|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))&((r2_hidden(esk6_4(X15,X16,X17,X18),X16)|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))&(X18=k4_tarski(esk5_4(X15,X16,X17,X18),esk6_4(X15,X16,X17,X18))|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))))&(~r2_hidden(X22,X15)|(~r2_hidden(X23,X16)|X21!=k4_tarski(X22,X23))|r2_hidden(X21,X17)|X17!=k2_zfmisc_1(X15,X16)))&((~r2_hidden(esk7_3(X24,X25,X26),X26)|(~r2_hidden(X28,X24)|(~r2_hidden(X29,X25)|esk7_3(X24,X25,X26)!=k4_tarski(X28,X29)))|X26=k2_zfmisc_1(X24,X25))&((r2_hidden(esk8_3(X24,X25,X26),X24)|r2_hidden(esk7_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25))&((r2_hidden(esk9_3(X24,X25,X26),X25)|r2_hidden(esk7_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25))&(esk7_3(X24,X25,X26)=k4_tarski(esk8_3(X24,X25,X26),esk9_3(X24,X25,X26))|r2_hidden(esk7_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_11, plain, ![X32, X33, X34, X35, X36]:((~r1_tarski(X32,X33)|(~r2_hidden(X34,X32)|r2_hidden(X34,X33)))&((r2_hidden(esk10_2(X35,X36),X35)|r1_tarski(X35,X36))&(~r2_hidden(esk10_2(X35,X36),X36)|r1_tarski(X35,X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_12, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))&(k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0&(~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_13, plain, (~r1_xboole_0(k2_tarski(X1,X2),X3)|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (r2_hidden(esk5_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_16, plain, ![X38, X39]:((~r2_hidden(esk11_2(X38,X39),X38)|~r2_hidden(esk11_2(X38,X39),X39)|X38=X39)&(r2_hidden(esk11_2(X38,X39),X38)|r2_hidden(esk11_2(X38,X39),X39)|X38=X39)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
fof(c_0_17, plain, ![X41, X42, X43, X44]:(((r2_hidden(X41,X43)|~r2_hidden(k4_tarski(X41,X42),k2_zfmisc_1(X43,X44)))&(r2_hidden(X42,X44)|~r2_hidden(k4_tarski(X41,X42),k2_zfmisc_1(X43,X44))))&(~r2_hidden(X41,X43)|~r2_hidden(X42,X44)|r2_hidden(k4_tarski(X41,X42),k2_zfmisc_1(X43,X44)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_18, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_21, plain, (r2_hidden(esk5_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (r2_hidden(esk11_2(X1,X2),X1)|r2_hidden(esk11_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk3_0,esk4_0))|~r2_hidden(X1,k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (~r2_hidden(X1,k2_zfmisc_1(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, plain, (k1_xboole_0=X1|r2_hidden(esk11_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_20, c_0_22])).
cnf(c_0_27, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(k4_tarski(X2,X1),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_32, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_29, c_0_24])).
cnf(c_0_34, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_30]), c_0_20])).
cnf(c_0_35, plain, (r2_hidden(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X2)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|~r2_hidden(esk10_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_37, negated_conjecture, (k1_xboole_0=esk1_0|r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_32, c_0_26])).
cnf(c_0_38, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)|~r1_tarski(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_39, plain, (r2_hidden(esk10_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_40, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X2,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_41, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k2_zfmisc_1(X3,X1))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, negated_conjecture, (k1_xboole_0=esk1_0|r1_tarski(X1,esk4_0)|~r2_hidden(esk10_2(X1,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk10_2(esk2_0,esk4_0),esk2_0)|~r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_45, plain, (r2_hidden(esk8_3(X1,X2,X3),X1)|r2_hidden(esk7_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_46, negated_conjecture, (k1_xboole_0=esk2_0|r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_40, c_0_26])).
cnf(c_0_47, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_26])).
cnf(c_0_48, negated_conjecture, (k1_xboole_0=esk1_0|~r1_tarski(esk1_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_42]), c_0_43])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk8_3(esk1_0,esk2_0,X1),esk1_0)|X1!=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_34])).
cnf(c_0_50, negated_conjecture, (k1_xboole_0=esk2_0|r1_tarski(X1,esk3_0)|~r2_hidden(esk10_2(X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_36, c_0_46])).
cnf(c_0_51, negated_conjecture, (k1_xboole_0!=esk2_0), inference(spm,[status(thm)],[c_0_44, c_0_47])).
cnf(c_0_52, negated_conjecture, (k1_xboole_0=esk1_0|r2_hidden(esk10_2(esk1_0,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_39])).
cnf(c_0_53, negated_conjecture, (k1_xboole_0!=esk1_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_49])).
cnf(c_0_54, negated_conjecture, (k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_55, negated_conjecture, (X1!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54]), c_0_54])])).
cnf(c_0_56, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_54, c_0_55]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 57
# Proof object clause steps            : 42
# Proof object formula steps           : 15
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 156
# ...of these trivial                  : 2
# ...subsumed                          : 44
# ...remaining for further processing  : 110
# Other redundant clauses eliminated   : 20
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 40
# Generated clauses                    : 795
# ...of the previous two non-trivial   : 776
# Contextual simplify-reflections      : 3
# Paramodulations                      : 760
# Factorizations                       : 8
# Equation resolutions                 : 26
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 43
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 37
# Current number of unprocessed clauses: 654
# ...number of literals in the above   : 3074
# Current number of archived formulas  : 0
# Current number of archived clauses   : 67
# Clause-clause subsumption calls (NU) : 939
# Rec. Clause-clause subsumption calls : 499
# Non-unit clause-clause subsumptions  : 22
# Unit Clause-clause subsumption calls : 304
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 11404

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.004 s
# Total time               : 0.019 s
# Maximum resident set size: 28776 pages
