# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t54_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>(~(v2_struct_0(X1))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4)))))))))), file('data/MPTP/proofs/a5a3987519484920985222e95b72ad18.E_input', t54_pre_topc)).
fof(d13_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k2_pre_topc(X1,X2)<=>![X4]:(r2_hidden(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X5,X1)&(r2_hidden(X4,X5)&r1_xboole_0(X2,X5))))))))))), file('data/MPTP/proofs/a5a3987519484920985222e95b72ad18.E_input', d13_pre_topc)).
fof(dt_k2_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/a5a3987519484920985222e95b72ad18.E_input', dt_k2_pre_topc)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/a5a3987519484920985222e95b72ad18.E_input', t5_subset)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/a5a3987519484920985222e95b72ad18.E_input', fc2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/a5a3987519484920985222e95b72ad18.E_input', dt_l1_pre_topc)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/a5a3987519484920985222e95b72ad18.E_input', t2_subset)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('data/MPTP/proofs/a5a3987519484920985222e95b72ad18.E_input', fc1_struct_0)).
fof(c_0_8, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>(~(v2_struct_0(X1))&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((v3_pre_topc(X4,X1)&(r2_hidden(X3,X4)&r1_xboole_0(X2,X4))))))))))), inference(assume_negation,[status(cth)],[t54_pre_topc])).
fof(c_0_9, plain, ![X11, X12, X13, X14, X15, X19]:(((~r2_hidden(X14,X13)|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X11)))|(~v3_pre_topc(X15,X11)|(~r2_hidden(X14,X15)|~r1_xboole_0(X12,X15))))|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((m1_subset_1(esk5_4(X11,X12,X13,X14),k1_zfmisc_1(u1_struct_0(X11)))|r2_hidden(X14,X13)|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((v3_pre_topc(esk5_4(X11,X12,X13,X14),X11)|r2_hidden(X14,X13)|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((r2_hidden(X14,esk5_4(X11,X12,X13,X14))|r2_hidden(X14,X13)|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&(r1_xboole_0(X12,esk5_4(X11,X12,X13,X14))|r2_hidden(X14,X13)|~r2_hidden(X14,u1_struct_0(X11))|X13!=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))))))&((r2_hidden(esk6_3(X11,X12,X13),u1_struct_0(X11))|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&(((m1_subset_1(esk7_3(X11,X12,X13),k1_zfmisc_1(u1_struct_0(X11)))|~r2_hidden(esk6_3(X11,X12,X13),X13)|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((v3_pre_topc(esk7_3(X11,X12,X13),X11)|~r2_hidden(esk6_3(X11,X12,X13),X13)|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&((r2_hidden(esk6_3(X11,X12,X13),esk7_3(X11,X12,X13))|~r2_hidden(esk6_3(X11,X12,X13),X13)|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))&(r1_xboole_0(X12,esk7_3(X11,X12,X13))|~r2_hidden(esk6_3(X11,X12,X13),X13)|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11)))))&(r2_hidden(esk6_3(X11,X12,X13),X13)|(~m1_subset_1(X19,k1_zfmisc_1(u1_struct_0(X11)))|(~v3_pre_topc(X19,X11)|(~r2_hidden(esk6_3(X11,X12,X13),X19)|~r1_xboole_0(X12,X19))))|X13=k2_pre_topc(X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(X11)))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))|~l1_pre_topc(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_pre_topc])])])])])).
fof(c_0_10, plain, ![X20, X21]:(~l1_pre_topc(X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|m1_subset_1(k2_pre_topc(X20,X21),k1_zfmisc_1(u1_struct_0(X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
fof(c_0_11, plain, ![X22, X23, X24]:(~r2_hidden(X22,X23)|(~m1_subset_1(X23,k1_zfmisc_1(X24))|~v1_xboole_0(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_12, negated_conjecture, ![X10]:(l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(((m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_struct_0(esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((v3_pre_topc(esk4_0,esk1_0)|v2_struct_0(esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&((r2_hidden(esk3_0,esk4_0)|v2_struct_0(esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&(r1_xboole_0(esk2_0,esk4_0)|v2_struct_0(esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))))))&((~v2_struct_0(esk1_0)|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))&(~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(esk1_0)))|(~v3_pre_topc(X10,esk1_0)|(~r2_hidden(esk3_0,X10)|~r1_xboole_0(esk2_0,X10)))|r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])])).
cnf(c_0_13, plain, (r1_xboole_0(X1,esk5_4(X2,X1,X3,X4))|r2_hidden(X4,X3)|~r2_hidden(X4,u1_struct_0(X2))|X3!=k2_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X25]:(v2_struct_0(X25)|~l1_struct_0(X25)|~v1_xboole_0(u1_struct_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_16, plain, ![X26]:(~l1_pre_topc(X26)|l1_struct_0(X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_17, plain, ![X27, X28]:(~m1_subset_1(X27,X28)|(v1_xboole_0(X28)|r2_hidden(X27,X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_18, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_19, plain, ![X29]:(~v2_struct_0(X29)|~l1_struct_0(X29)|v1_xboole_0(u1_struct_0(X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_20, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk3_0,X1)|~r1_xboole_0(esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (r1_xboole_0(X1,esk5_4(X2,X1,k2_pre_topc(X2,X1),X3))|r2_hidden(X3,k2_pre_topc(X2,X1))|~r2_hidden(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_14])).
cnf(c_0_22, plain, (v3_pre_topc(esk5_4(X1,X2,X3,X4),X1)|r2_hidden(X4,X3)|~r2_hidden(X4,u1_struct_0(X1))|X3!=k2_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27, plain, (~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,k2_pre_topc(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_18, c_0_14])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_30, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_31, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(X1,k2_pre_topc(X2,esk2_0))|~v3_pre_topc(esk5_4(X2,esk2_0,k2_pre_topc(X2,esk2_0),X1),esk1_0)|~r2_hidden(esk3_0,esk5_4(X2,esk2_0,k2_pre_topc(X2,esk2_0),X1))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(esk5_4(X2,esk2_0,k2_pre_topc(X2,esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_33, plain, (v3_pre_topc(esk5_4(X1,X2,k2_pre_topc(X1,X2),X3),X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|~r2_hidden(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_22]), c_0_14])).
cnf(c_0_34, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|r2_hidden(X4,X3)|~r2_hidden(X4,u1_struct_0(X1))|X3!=k2_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_36, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_37, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))|~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30])])).
cnf(c_0_38, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_31, c_0_24])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(X1,k2_pre_topc(esk1_0,esk2_0))|~r2_hidden(esk3_0,esk5_4(esk1_0,esk2_0,k2_pre_topc(esk1_0,esk2_0),X1))|~r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk5_4(esk1_0,esk2_0,k2_pre_topc(esk1_0,esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_29]), c_0_30])])).
cnf(c_0_40, plain, (r2_hidden(X1,k2_pre_topc(X2,X3))|m1_subset_1(esk5_4(X2,X3,k2_pre_topc(X2,X3),X1),k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_34]), c_0_14])).
cnf(c_0_41, plain, (r2_hidden(X1,esk5_4(X2,X3,X4,X1))|r2_hidden(X1,X4)|~r2_hidden(X1,u1_struct_0(X2))|X4!=k2_pre_topc(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_42, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk3_0,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_30])])).
cnf(c_0_43, negated_conjecture, (~v2_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_30])])).
cnf(c_0_44, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))|~v3_pre_topc(X3,X4)|~r2_hidden(X1,X3)|~r1_xboole_0(X5,X3)|~r2_hidden(X1,u1_struct_0(X4))|X2!=k2_pre_topc(X4,X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|~l1_pre_topc(X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))|r2_hidden(X1,k2_pre_topc(esk1_0,esk2_0))|~r2_hidden(esk3_0,esk5_4(esk1_0,esk2_0,k2_pre_topc(esk1_0,esk2_0),X1))|~r2_hidden(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_29]), c_0_30])])).
cnf(c_0_46, plain, (r2_hidden(X1,esk5_4(X2,X3,k2_pre_topc(X2,X3),X1))|r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_41]), c_0_14])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, plain, (~r1_xboole_0(X1,X2)|~v3_pre_topc(X2,X3)|~r2_hidden(X4,k2_pre_topc(X3,X1))|~r2_hidden(X4,u1_struct_0(X3))|~r2_hidden(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_44]), c_0_14])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47]), c_0_29]), c_0_30])])).
cnf(c_0_50, negated_conjecture, (r1_xboole_0(esk2_0,esk4_0)|v2_struct_0(esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_51, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)|v2_struct_0(esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk3_0,esk4_0)|v2_struct_0(esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_struct_0(esk1_0)|~r2_hidden(esk3_0,k2_pre_topc(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_54, negated_conjecture, (~r1_xboole_0(esk2_0,X1)|~v3_pre_topc(X1,esk1_0)|~r2_hidden(esk3_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_47]), c_0_29]), c_0_30])])).
cnf(c_0_55, negated_conjecture, (r1_xboole_0(esk2_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_49])]), c_0_43])).
cnf(c_0_56, negated_conjecture, (v3_pre_topc(esk4_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_49])]), c_0_43])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_49])]), c_0_43])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_49])]), c_0_43])).
cnf(c_0_59, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_57]), c_0_58])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 43
# Proof object formula steps           : 17
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 107
# ...of these trivial                  : 0
# ...subsumed                          : 10
# ...remaining for further processing  : 97
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 11
# Generated clauses                    : 64
# ...of the previous two non-trivial   : 62
# Contextual simplify-reflections      : 7
# Paramodulations                      : 58
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 25
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 1642
# Rec. Clause-clause subsumption calls : 118
# Non-unit clause-clause subsumptions  : 15
# Unit Clause-clause subsumption calls : 62
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5098

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28744 pages
