# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,X2)=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/a68f30f55dcb4deabc7611d87eda5833.E_input', d1_tops_2)).
fof(t7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/a68f30f55dcb4deabc7611d87eda5833.E_input', t7_subset_1)).
fof(t26_tops_2, conjecture, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)=>v3_pre_topc(k5_setfam_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/a68f30f55dcb4deabc7611d87eda5833.E_input', t26_tops_2)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/a68f30f55dcb4deabc7611d87eda5833.E_input', d5_pre_topc)).
fof(dt_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>m1_subset_1(k5_setfam_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/a68f30f55dcb4deabc7611d87eda5833.E_input', dt_k5_setfam_1)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/a68f30f55dcb4deabc7611d87eda5833.E_input', dt_u1_pre_topc)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/a68f30f55dcb4deabc7611d87eda5833.E_input', d1_pre_topc)).
fof(c_0_7, plain, ![X7, X8, X9]:((~v1_tops_2(X8,X7)|(~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X7)))|(~r2_hidden(X9,X8)|v3_pre_topc(X9,X7)))|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X7))))|~l1_pre_topc(X7))&((m1_subset_1(esk3_2(X7,X8),k1_zfmisc_1(u1_struct_0(X7)))|v1_tops_2(X8,X7)|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X7))))|~l1_pre_topc(X7))&((r2_hidden(esk3_2(X7,X8),X8)|v1_tops_2(X8,X7)|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X7))))|~l1_pre_topc(X7))&(~v3_pre_topc(esk3_2(X7,X8),X7)|v1_tops_2(X8,X7)|~m1_subset_1(X8,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X7))))|~l1_pre_topc(X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tops_2])])])])])).
fof(c_0_8, plain, ![X11, X12, X13]:((m1_subset_1(esk4_3(X11,X12,X13),X11)|r1_tarski(X12,X13)|~m1_subset_1(X13,k1_zfmisc_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(X11)))&((r2_hidden(esk4_3(X11,X12,X13),X12)|r1_tarski(X12,X13)|~m1_subset_1(X13,k1_zfmisc_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(X11)))&(~r2_hidden(esk4_3(X11,X12,X13),X13)|r1_tarski(X12,X13)|~m1_subset_1(X13,k1_zfmisc_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_subset_1])])])])])).
fof(c_0_9, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(v1_tops_2(X2,X1)=>v3_pre_topc(k5_setfam_1(u1_struct_0(X1),X2),X1))))), inference(assume_negation,[status(cth)],[t26_tops_2])).
fof(c_0_10, plain, ![X15, X16]:((~v3_pre_topc(X16,X15)|r2_hidden(X16,u1_pre_topc(X15))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))|~l1_pre_topc(X15))&(~r2_hidden(X16,u1_pre_topc(X15))|v3_pre_topc(X16,X15)|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))|~l1_pre_topc(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
cnf(c_0_11, plain, (v3_pre_topc(X3,X2)|~v1_tops_2(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_13, negated_conjecture, ((v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0))&(m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))&(v1_tops_2(esk2_0,esk1_0)&~v3_pre_topc(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
cnf(c_0_14, plain, (r2_hidden(X1,u1_pre_topc(X2))|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (r1_tarski(X1,X2)|v3_pre_topc(esk4_3(X3,X1,X2),X4)|~v1_tops_2(X1,X4)|~m1_subset_1(esk4_3(X3,X1,X2),k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))|~l1_pre_topc(X4)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, negated_conjecture, (~v3_pre_topc(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, ![X25, X26]:(~m1_subset_1(X26,k1_zfmisc_1(k1_zfmisc_1(X25)))|m1_subset_1(k5_setfam_1(X25,X26),k1_zfmisc_1(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_setfam_1])])).
cnf(c_0_20, plain, (r1_tarski(X1,X2)|r2_hidden(esk4_3(X3,X1,X2),u1_pre_topc(X4))|~v1_tops_2(X1,X4)|~m1_subset_1(esk4_3(X3,X1,X2),k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))|~l1_pre_topc(X4)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_21, plain, (m1_subset_1(esk4_3(X1,X2,X3),X1)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_22, plain, ![X17]:(~l1_pre_topc(X17)|m1_subset_1(u1_pre_topc(X17),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X17))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_23, negated_conjecture, (~r2_hidden(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),u1_pre_topc(esk1_0))|~m1_subset_1(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])])).
cnf(c_0_24, plain, (m1_subset_1(k5_setfam_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_26, plain, ![X18, X19, X20, X21]:(((r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|~v2_pre_topc(X18)|~l1_pre_topc(X18))&((~m1_subset_1(X19,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|(~r1_tarski(X19,u1_pre_topc(X18))|r2_hidden(k5_setfam_1(u1_struct_0(X18),X19),u1_pre_topc(X18)))|~v2_pre_topc(X18)|~l1_pre_topc(X18))&(~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|(~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X18)))|(~r2_hidden(X20,u1_pre_topc(X18))|~r2_hidden(X21,u1_pre_topc(X18))|r2_hidden(k9_subset_1(u1_struct_0(X18),X20,X21),u1_pre_topc(X18))))|~v2_pre_topc(X18)|~l1_pre_topc(X18))))&(((m1_subset_1(esk6_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|m1_subset_1(esk5_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&((m1_subset_1(esk7_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|m1_subset_1(esk5_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(((r2_hidden(esk6_1(X18),u1_pre_topc(X18))|m1_subset_1(esk5_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(r2_hidden(esk7_1(X18),u1_pre_topc(X18))|m1_subset_1(esk5_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))&(~r2_hidden(k9_subset_1(u1_struct_0(X18),esk6_1(X18),esk7_1(X18)),u1_pre_topc(X18))|m1_subset_1(esk5_1(X18),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X18))))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))))&(((m1_subset_1(esk6_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|r1_tarski(esk5_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&((m1_subset_1(esk7_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|r1_tarski(esk5_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(((r2_hidden(esk6_1(X18),u1_pre_topc(X18))|r1_tarski(esk5_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(r2_hidden(esk7_1(X18),u1_pre_topc(X18))|r1_tarski(esk5_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))&(~r2_hidden(k9_subset_1(u1_struct_0(X18),esk6_1(X18),esk7_1(X18)),u1_pre_topc(X18))|r1_tarski(esk5_1(X18),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))))&((m1_subset_1(esk6_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk5_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&((m1_subset_1(esk7_1(X18),k1_zfmisc_1(u1_struct_0(X18)))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk5_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(((r2_hidden(esk6_1(X18),u1_pre_topc(X18))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk5_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18))&(r2_hidden(esk7_1(X18),u1_pre_topc(X18))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk5_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))&(~r2_hidden(k9_subset_1(u1_struct_0(X18),esk6_1(X18),esk7_1(X18)),u1_pre_topc(X18))|~r2_hidden(k5_setfam_1(u1_struct_0(X18),esk5_1(X18)),u1_pre_topc(X18))|~r2_hidden(u1_struct_0(X18),u1_pre_topc(X18))|v2_pre_topc(X18)|~l1_pre_topc(X18)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
cnf(c_0_27, plain, (r1_tarski(X2,X3)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|r2_hidden(esk4_3(k1_zfmisc_1(u1_struct_0(X3)),X1,X2),u1_pre_topc(X3))|~v1_tops_2(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_29, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (~r2_hidden(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])])).
cnf(c_0_31, plain, (r2_hidden(k5_setfam_1(u1_struct_0(X2),X1),u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~r1_tarski(X1,u1_pre_topc(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_33, plain, (r1_tarski(X1,u1_pre_topc(X2))|~v1_tops_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_34, negated_conjecture, (v1_tops_2(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_35, negated_conjecture, (~r1_tarski(esk2_0,u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_25]), c_0_18]), c_0_32])])).
cnf(c_0_36, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_25]), c_0_18])]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 22
# Proof object formula steps           : 15
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 34
# Processed clauses                    : 97
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 94
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 43
# ...of the previous two non-trivial   : 31
# Contextual simplify-reflections      : 4
# Paramodulations                      : 43
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 60
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 52
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 21
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 1258
# Rec. Clause-clause subsumption calls : 64
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 19
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5556

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28516 pages
