# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t36_funct_2, conjecture, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>((k2_relset_1(X1,X2,X3)=X2&(r2_relset_1(X1,X1,k1_partfun1(X1,X2,X2,X1,X3,X4),k6_partfun1(X1))&v2_funct_1(X3)))=>(X1=k1_xboole_0|(X2=k1_xboole_0|X4=k2_funct_1(X3)))))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t36_funct_2)).
fof(t29_relset_1, axiom, ![X1]:m1_subset_1(k6_relat_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t29_relset_1)).
fof(redefinition_k6_partfun1, axiom, ![X1]:k6_partfun1(X1)=k6_relat_1(X1), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', redefinition_k6_partfun1)).
fof(redefinition_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_relset_1(X1,X2,X3,X4)<=>X3=X4)), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', redefinition_r2_relset_1)).
fof(t24_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>(r2_relset_1(X2,X2,k1_partfun1(X2,X1,X1,X2,X4,X3),k6_partfun1(X2))=>k2_relset_1(X1,X2,X3)=X2))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t24_funct_2)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', dt_k1_partfun1)).
fof(fc4_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v2_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', fc4_funct_1)).
fof(t35_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((k2_relset_1(X1,X2,X3)=X2&v2_funct_1(X3))=>(X2=k1_xboole_0|(k5_relat_1(X3,k2_funct_1(X3))=k6_partfun1(X1)&k5_relat_1(k2_funct_1(X3),X3)=k6_partfun1(X2))))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t35_funct_2)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t71_relat_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', cc1_relset_1)).
fof(t30_funct_2, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X2,X3)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>((v2_funct_1(k1_partfun1(X1,X2,X2,X3,X4,X5))&k2_relset_1(X1,X2,X4)=X2)=>((X3=k1_xboole_0&~(X2=k1_xboole_0))|(v2_funct_1(X4)&v2_funct_1(X5)))))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t30_funct_2)).
fof(t80_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k5_relat_1(X1,k6_relat_1(k2_relat_1(X1)))=X1), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t80_relat_1)).
fof(t58_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k1_relat_1(k5_relat_1(X1,k2_funct_1(X1)))=k1_relat_1(X1)&k2_relat_1(k5_relat_1(X1,k2_funct_1(X1)))=k1_relat_1(X1)))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t58_funct_1)).
fof(t37_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k4_relat_1(X1))&k1_relat_1(X1)=k2_relat_1(k4_relat_1(X1)))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t37_relat_1)).
fof(dt_k4_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>v1_relat_1(k4_relat_1(X1))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', dt_k4_relat_1)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', redefinition_k1_partfun1)).
fof(t55_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>k5_relat_1(k5_relat_1(X1,X2),X3)=k5_relat_1(X1,k5_relat_1(X2,X3))))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t55_relat_1)).
fof(d9_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>k2_funct_1(X1)=k4_relat_1(X1))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', d9_funct_1)).
fof(t78_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k5_relat_1(k6_relat_1(k1_relat_1(X1)),X1)=X1), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', t78_relat_1)).
fof(dt_k2_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k2_funct_1(X1))&v1_funct_1(k2_funct_1(X1)))), file('data/MPTP/proofs/a80cfb6dd36d4f3182403fea002a57a5.E_input', dt_k2_funct_1)).
fof(c_0_20, negated_conjecture, ~(![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X1)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))))=>((k2_relset_1(X1,X2,X3)=X2&(r2_relset_1(X1,X1,k1_partfun1(X1,X2,X2,X1,X3,X4),k6_partfun1(X1))&v2_funct_1(X3)))=>(X1=k1_xboole_0|(X2=k1_xboole_0|X4=k2_funct_1(X3))))))), inference(assume_negation,[status(cth)],[t36_funct_2])).
fof(c_0_21, plain, ![X11]:m1_subset_1(k6_relat_1(X11),k1_zfmisc_1(k2_zfmisc_1(X11,X11))), inference(variable_rename,[status(thm)],[t29_relset_1])).
fof(c_0_22, plain, ![X12]:k6_partfun1(X12)=k6_relat_1(X12), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
fof(c_0_23, plain, ![X17, X18, X19, X20]:((~r2_relset_1(X17,X18,X19,X20)|X19=X20|(~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,X18)))|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X17,X18)))))&(X19!=X20|r2_relset_1(X17,X18,X19,X20)|(~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,X18)))|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X17,X18)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_relset_1])])])).
fof(c_0_24, negated_conjecture, ((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk2_0,esk1_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))))&((k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0&(r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,esk4_0),k6_partfun1(esk1_0))&v2_funct_1(esk3_0)))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&esk4_0!=k2_funct_1(esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
cnf(c_0_25, plain, (m1_subset_1(k6_relat_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, plain, (k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, ![X13, X14, X15, X16]:(~v1_funct_1(X15)|(~v1_funct_2(X15,X13,X14)|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))|(~v1_funct_1(X16)|(~v1_funct_2(X16,X14,X13)|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(X14,X13))))|(~r2_relset_1(X14,X14,k1_partfun1(X14,X13,X13,X14,X16,X15),k6_partfun1(X14))|k2_relset_1(X13,X14,X15)=X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t24_funct_2])])])).
cnf(c_0_28, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,esk4_0),k6_partfun1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_31, plain, ![X28, X29, X30, X31, X32, X33]:((v1_funct_1(k1_partfun1(X28,X29,X30,X31,X32,X33))|(~v1_funct_1(X32)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|(~v1_funct_1(X33)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X30,X31))))))&(m1_subset_1(k1_partfun1(X28,X29,X30,X31,X32,X33),k1_zfmisc_1(k2_zfmisc_1(X28,X31)))|(~v1_funct_1(X32)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|(~v1_funct_1(X33)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X30,X31))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
fof(c_0_32, plain, ![X34]:(v1_relat_1(k6_relat_1(X34))&v2_funct_1(k6_relat_1(X34))), inference(variable_rename,[status(thm)],[fc4_funct_1])).
fof(c_0_33, plain, ![X21, X22, X23]:((k5_relat_1(X23,k2_funct_1(X23))=k6_partfun1(X21)|X22=k1_xboole_0|(k2_relset_1(X21,X22,X23)!=X22|~v2_funct_1(X23))|(~v1_funct_1(X23)|(~v1_funct_2(X23,X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(X21,X22))))))&(k5_relat_1(k2_funct_1(X23),X23)=k6_partfun1(X22)|X22=k1_xboole_0|(k2_relset_1(X21,X22,X23)!=X22|~v2_funct_1(X23))|(~v1_funct_1(X23)|(~v1_funct_2(X23,X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(X21,X22))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_funct_2])])])).
cnf(c_0_34, plain, (k2_relset_1(X2,X3,X1)=X3|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~v1_funct_2(X4,X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~r2_relset_1(X3,X3,k1_partfun1(X3,X2,X2,X3,X4,X1),k6_partfun1(X3))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, negated_conjecture, (v1_funct_2(esk4_0,esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_39, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_40, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_41, plain, ![X24]:(k1_relat_1(k6_relat_1(X24))=X24&k2_relat_1(k6_relat_1(X24))=X24), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_42, plain, ![X25, X26, X27]:(~m1_subset_1(X27,k1_zfmisc_1(k2_zfmisc_1(X25,X26)))|v1_relat_1(X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_43, plain, ![X37, X38, X39, X40, X41]:(((v2_funct_1(X40)|X39=k1_xboole_0|(~v2_funct_1(k1_partfun1(X37,X38,X38,X39,X40,X41))|k2_relset_1(X37,X38,X40)!=X38)|(~v1_funct_1(X41)|(~v1_funct_2(X41,X38,X39)|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))))|(~v1_funct_1(X40)|(~v1_funct_2(X40,X37,X38)|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X37,X38))))))&(v2_funct_1(X41)|X39=k1_xboole_0|(~v2_funct_1(k1_partfun1(X37,X38,X38,X39,X40,X41))|k2_relset_1(X37,X38,X40)!=X38)|(~v1_funct_1(X41)|(~v1_funct_2(X41,X38,X39)|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))))|(~v1_funct_1(X40)|(~v1_funct_2(X40,X37,X38)|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X37,X38)))))))&((v2_funct_1(X40)|X38!=k1_xboole_0|(~v2_funct_1(k1_partfun1(X37,X38,X38,X39,X40,X41))|k2_relset_1(X37,X38,X40)!=X38)|(~v1_funct_1(X41)|(~v1_funct_2(X41,X38,X39)|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))))|(~v1_funct_1(X40)|(~v1_funct_2(X40,X37,X38)|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X37,X38))))))&(v2_funct_1(X41)|X38!=k1_xboole_0|(~v2_funct_1(k1_partfun1(X37,X38,X38,X39,X40,X41))|k2_relset_1(X37,X38,X40)!=X38)|(~v1_funct_1(X41)|(~v1_funct_2(X41,X38,X39)|~m1_subset_1(X41,k1_zfmisc_1(k2_zfmisc_1(X38,X39)))))|(~v1_funct_1(X40)|(~v1_funct_2(X40,X37,X38)|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(X37,X38)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t30_funct_2])])])])])).
cnf(c_0_44, negated_conjecture, (k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,esk4_0)=k6_partfun1(esk1_0)|~m1_subset_1(k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_45, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_46, plain, (v2_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_47, plain, ![X35]:(~v1_relat_1(X35)|k5_relat_1(X35,k6_relat_1(k2_relat_1(X35)))=X35), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t80_relat_1])])).
fof(c_0_48, plain, ![X36]:((k1_relat_1(k5_relat_1(X36,k2_funct_1(X36)))=k1_relat_1(X36)|~v2_funct_1(X36)|(~v1_relat_1(X36)|~v1_funct_1(X36)))&(k2_relat_1(k5_relat_1(X36,k2_funct_1(X36)))=k1_relat_1(X36)|~v2_funct_1(X36)|(~v1_relat_1(X36)|~v1_funct_1(X36)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t58_funct_1])])])).
cnf(c_0_49, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_partfun1(X2)|X3=k1_xboole_0|k2_relset_1(X2,X3,X1)!=X3|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_50, negated_conjecture, (k2_relset_1(esk2_0,esk1_0,esk4_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_29]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])])).
cnf(c_0_51, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_52, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54, plain, (k5_relat_1(k2_funct_1(X1),X1)=k6_partfun1(X2)|X2=k1_xboole_0|k2_relset_1(X3,X2,X1)!=X2|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_55, plain, (v2_funct_1(X1)|X2=k1_xboole_0|~v2_funct_1(k1_partfun1(X3,X4,X4,X2,X5,X1))|k2_relset_1(X3,X4,X5)!=X4|~v1_funct_1(X1)|~v1_funct_2(X1,X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~v1_funct_1(X5)|~v1_funct_2(X5,X3,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_56, negated_conjecture, (k1_partfun1(esk1_0,esk2_0,esk2_0,esk1_0,esk3_0,esk4_0)=k6_partfun1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_36]), c_0_35]), c_0_40]), c_0_39])])).
cnf(c_0_57, negated_conjecture, (k2_relset_1(esk1_0,esk2_0,esk3_0)=esk2_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_58, plain, (v2_funct_1(k6_partfun1(X1))), inference(rw,[status(thm)],[c_0_46, c_0_26])).
cnf(c_0_59, plain, (k5_relat_1(X1,k6_relat_1(k2_relat_1(X1)))=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_60, plain, ![X42]:((k2_relat_1(X42)=k1_relat_1(k4_relat_1(X42))|~v1_relat_1(X42))&(k1_relat_1(X42)=k2_relat_1(k4_relat_1(X42))|~v1_relat_1(X42))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_relat_1])])])).
fof(c_0_61, plain, ![X43]:(~v1_relat_1(X43)|v1_relat_1(k4_relat_1(X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_relat_1])])).
cnf(c_0_62, plain, (k1_relat_1(k5_relat_1(X1,k2_funct_1(X1)))=k1_relat_1(X1)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, negated_conjecture, (k5_relat_1(esk4_0,k2_funct_1(esk4_0))=k6_partfun1(esk2_0)|~v2_funct_1(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_36]), c_0_38]), c_0_40])]), c_0_51])).
cnf(c_0_64, plain, (k1_relat_1(k6_partfun1(X1))=X1), inference(rw,[status(thm)],[c_0_52, c_0_26])).
cnf(c_0_65, negated_conjecture, (v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_53, c_0_36])).
fof(c_0_66, plain, ![X44, X45, X46, X47, X48, X49]:(~v1_funct_1(X48)|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X44,X45)))|(~v1_funct_1(X49)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X46,X47))))|k1_partfun1(X44,X45,X46,X47,X48,X49)=k5_relat_1(X48,X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
fof(c_0_67, plain, ![X50, X51, X52]:(~v1_relat_1(X50)|(~v1_relat_1(X51)|(~v1_relat_1(X52)|k5_relat_1(k5_relat_1(X50,X51),X52)=k5_relat_1(X50,k5_relat_1(X51,X52))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_relat_1])])])).
cnf(c_0_68, negated_conjecture, (k5_relat_1(k2_funct_1(esk4_0),esk4_0)=k6_partfun1(esk1_0)|~v2_funct_1(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_50]), c_0_36]), c_0_38]), c_0_40])]), c_0_51])).
cnf(c_0_69, negated_conjecture, (v2_funct_1(esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57]), c_0_58]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])]), c_0_51])).
cnf(c_0_70, plain, (k5_relat_1(X1,k6_partfun1(k2_relat_1(X1)))=X1|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_59, c_0_26])).
cnf(c_0_71, plain, (k1_relat_1(X1)=k2_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_72, plain, (v1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_73, negated_conjecture, (k1_relat_1(esk4_0)=esk2_0|~v2_funct_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_65]), c_0_40])])).
cnf(c_0_74, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_75, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_76, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_77, plain, (k5_relat_1(k5_relat_1(X1,X2),X3)=k5_relat_1(X1,k5_relat_1(X2,X3))|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_78, negated_conjecture, (k5_relat_1(k2_funct_1(esk4_0),esk4_0)=k6_partfun1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69])])).
cnf(c_0_79, plain, (k5_relat_1(k4_relat_1(X1),k6_partfun1(k1_relat_1(X1)))=k4_relat_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])).
cnf(c_0_80, negated_conjecture, (k1_relat_1(esk4_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_69])])).
fof(c_0_81, plain, ![X53]:(~v1_relat_1(X53)|~v1_funct_1(X53)|(~v2_funct_1(X53)|k2_funct_1(X53)=k4_relat_1(X53))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_funct_1])])).
fof(c_0_82, plain, ![X54]:(~v1_relat_1(X54)|k5_relat_1(k6_relat_1(k1_relat_1(X54)),X54)=X54), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t78_relat_1])])).
cnf(c_0_83, negated_conjecture, (k5_relat_1(esk3_0,k2_funct_1(esk3_0))=k6_partfun1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_57]), c_0_74]), c_0_35]), c_0_37]), c_0_39])]), c_0_75])).
cnf(c_0_84, negated_conjecture, (k5_relat_1(esk3_0,esk4_0)=k6_partfun1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_56]), c_0_36]), c_0_35]), c_0_40]), c_0_39])])).
cnf(c_0_85, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_35])).
cnf(c_0_86, negated_conjecture, (k5_relat_1(k2_funct_1(esk4_0),k5_relat_1(esk4_0,X1))=k5_relat_1(k6_partfun1(esk1_0),X1)|~v1_relat_1(k2_funct_1(esk4_0))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_65])])).
cnf(c_0_87, negated_conjecture, (k5_relat_1(esk4_0,k2_funct_1(esk4_0))=k6_partfun1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_69])])).
cnf(c_0_88, negated_conjecture, (k5_relat_1(k4_relat_1(esk4_0),k6_partfun1(esk2_0))=k4_relat_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_65])])).
cnf(c_0_89, plain, (k2_funct_1(X1)=k4_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_90, plain, (k5_relat_1(k6_relat_1(k1_relat_1(X1)),X1)=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_91, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0|~v1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_83]), c_0_64]), c_0_74]), c_0_39])])).
cnf(c_0_92, negated_conjecture, (k5_relat_1(esk3_0,k5_relat_1(esk4_0,X1))=k5_relat_1(k6_partfun1(esk1_0),X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_84]), c_0_65]), c_0_85])])).
cnf(c_0_93, negated_conjecture, (k5_relat_1(k2_funct_1(esk4_0),k6_partfun1(esk2_0))=k5_relat_1(k6_partfun1(esk1_0),k2_funct_1(esk4_0))|~v1_relat_1(k2_funct_1(esk4_0))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_94, negated_conjecture, (k5_relat_1(k2_funct_1(esk4_0),k6_partfun1(esk2_0))=k2_funct_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_65]), c_0_69]), c_0_40])])).
cnf(c_0_95, plain, (k5_relat_1(k6_partfun1(k1_relat_1(X1)),X1)=X1|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_90, c_0_26])).
cnf(c_0_96, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_85])])).
cnf(c_0_97, negated_conjecture, (k5_relat_1(k6_partfun1(esk1_0),k2_funct_1(esk4_0))=k5_relat_1(esk3_0,k6_partfun1(esk2_0))|~v1_relat_1(k2_funct_1(esk4_0))), inference(spm,[status(thm)],[c_0_92, c_0_87])).
cnf(c_0_98, negated_conjecture, (k5_relat_1(k6_partfun1(esk1_0),k2_funct_1(esk4_0))=k2_funct_1(esk4_0)|~v1_relat_1(k2_funct_1(esk4_0))), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_99, negated_conjecture, (k5_relat_1(esk3_0,k5_relat_1(k2_funct_1(esk3_0),X1))=k5_relat_1(k6_partfun1(esk1_0),X1)|~v1_relat_1(k2_funct_1(esk3_0))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_83]), c_0_85])])).
cnf(c_0_100, negated_conjecture, (k5_relat_1(k2_funct_1(esk3_0),esk3_0)=k6_partfun1(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_57]), c_0_74]), c_0_35]), c_0_37]), c_0_39])]), c_0_75])).
cnf(c_0_101, negated_conjecture, (k5_relat_1(k6_partfun1(esk1_0),esk3_0)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_85])])).
cnf(c_0_102, negated_conjecture, (k5_relat_1(esk3_0,k6_partfun1(esk2_0))=k2_funct_1(esk4_0)|~v1_relat_1(k2_funct_1(esk4_0))), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_103, negated_conjecture, (k5_relat_1(esk3_0,k6_partfun1(esk2_0))=esk3_0|~v1_relat_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101]), c_0_85])])).
fof(c_0_104, plain, ![X55]:((v1_relat_1(k2_funct_1(X55))|(~v1_relat_1(X55)|~v1_funct_1(X55)))&(v1_funct_1(k2_funct_1(X55))|(~v1_relat_1(X55)|~v1_funct_1(X55)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_1])])])).
cnf(c_0_105, negated_conjecture, (k2_funct_1(esk4_0)=esk3_0|~v1_relat_1(k2_funct_1(esk4_0))|~v1_relat_1(k2_funct_1(esk3_0))), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_106, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_107, negated_conjecture, (k2_funct_1(esk4_0)=esk3_0|~v1_relat_1(k2_funct_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_65]), c_0_40])])).
cnf(c_0_108, negated_conjecture, (k5_relat_1(k6_partfun1(esk2_0),esk4_0)=esk4_0|~v2_funct_1(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_73]), c_0_65])])).
cnf(c_0_109, plain, (k2_relat_1(X1)=k1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_110, plain, (k2_relat_1(k2_funct_1(X1))=k1_relat_1(X1)|~v1_relat_1(X1)|~v2_funct_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_71, c_0_89])).
cnf(c_0_111, negated_conjecture, (k2_funct_1(esk4_0)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_106]), c_0_85]), c_0_39])])).
cnf(c_0_112, negated_conjecture, (k5_relat_1(esk4_0,k5_relat_1(k2_funct_1(esk4_0),X1))=k5_relat_1(k6_partfun1(esk2_0),X1)|~v1_relat_1(k2_funct_1(esk4_0))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_87]), c_0_65])])).
cnf(c_0_113, negated_conjecture, (k5_relat_1(k6_partfun1(esk2_0),esk4_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_69])])).
cnf(c_0_114, plain, (k5_relat_1(k6_partfun1(k2_relat_1(X1)),k4_relat_1(X1))=k4_relat_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_109]), c_0_72])).
cnf(c_0_115, negated_conjecture, (k2_relat_1(esk3_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_80]), c_0_65]), c_0_69]), c_0_40])])).
cnf(c_0_116, negated_conjecture, (k5_relat_1(esk4_0,k6_partfun1(esk1_0))=esk4_0|~v1_relat_1(k2_funct_1(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_78]), c_0_113]), c_0_65])])).
cnf(c_0_117, negated_conjecture, (k5_relat_1(k6_partfun1(esk2_0),k4_relat_1(esk3_0))=k4_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_85])])).
cnf(c_0_118, negated_conjecture, (k5_relat_1(esk4_0,k5_relat_1(esk3_0,X1))=k5_relat_1(k6_partfun1(esk2_0),X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_111]), c_0_111]), c_0_85])])).
cnf(c_0_119, negated_conjecture, (k5_relat_1(esk4_0,k6_partfun1(esk1_0))=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_111]), c_0_85])])).
cnf(c_0_120, negated_conjecture, (k5_relat_1(k6_partfun1(esk2_0),k2_funct_1(esk3_0))=k2_funct_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_89]), c_0_85]), c_0_74]), c_0_39])])).
cnf(c_0_121, negated_conjecture, (esk4_0!=k2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_122, negated_conjecture, (~v1_relat_1(k2_funct_1(esk3_0))), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_83]), c_0_119]), c_0_120]), c_0_121])).
cnf(c_0_123, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_106]), c_0_85]), c_0_39])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 124
# Proof object clause steps            : 83
# Proof object formula steps           : 41
# Proof object conjectures             : 57
# Proof object clause conjectures      : 54
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 131
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 41
# Processed clauses                    : 196
# ...of these trivial                  : 19
# ...subsumed                          : 16
# ...remaining for further processing  : 161
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 32
# Generated clauses                    : 246
# ...of the previous two non-trivial   : 227
# Contextual simplify-reflections      : 2
# Paramodulations                      : 243
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 82
#    Positive orientable unit clauses  : 42
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 106
# ...number of literals in the above   : 430
# Current number of archived formulas  : 0
# Current number of archived clauses   : 77
# Clause-clause subsumption calls (NU) : 652
# Rec. Clause-clause subsumption calls : 249
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 20
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8731

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.004 s
# Total time               : 0.015 s
# Maximum resident set size: 28732 pages
