# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C41_F1_AE_CS_SP_PI_S4S
# and selection function SelectNewComplexAHPNS.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t14_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&(r1_tarski(X3,X2)&![X4]:((r1_tarski(X1,X4)&r1_tarski(X3,X4))=>r1_tarski(X2,X4))))=>X2=k2_xboole_0(X1,X3)), file('data/MPTP/proofs/a82a31707d3a48aeb63943cfaef63d03.E_input', t14_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/a82a31707d3a48aeb63943cfaef63d03.E_input', t17_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/a82a31707d3a48aeb63943cfaef63d03.E_input', d10_xboole_0)).
fof(t22_xboole_1, conjecture, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/a82a31707d3a48aeb63943cfaef63d03.E_input', t22_xboole_1)).
fof(c_0_4, plain, ![X7, X8, X9]:(((r1_tarski(X7,esk3_3(X7,X8,X9))|~r1_tarski(X9,X8)|~r1_tarski(X7,X8)|X8=k2_xboole_0(X7,X9))&(r1_tarski(X9,esk3_3(X7,X8,X9))|~r1_tarski(X9,X8)|~r1_tarski(X7,X8)|X8=k2_xboole_0(X7,X9)))&(~r1_tarski(X8,esk3_3(X7,X8,X9))|~r1_tarski(X9,X8)|~r1_tarski(X7,X8)|X8=k2_xboole_0(X7,X9))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_xboole_1])])])])).
fof(c_0_5, plain, ![X13, X14]:r1_tarski(k3_xboole_0(X13,X14),X13), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_6, plain, ![X11, X12]:(((r1_tarski(X11,X12)|X11!=X12)&(r1_tarski(X12,X11)|X11!=X12))&(~r1_tarski(X11,X12)|~r1_tarski(X12,X11)|X11=X12)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_7, plain, (r1_tarski(X1,esk3_3(X1,X2,X3))|X2=k2_xboole_0(X1,X3)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_10, negated_conjecture, ~(![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(assume_negation,[status(cth)],[t22_xboole_1])).
cnf(c_0_11, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X3))=X2|r1_tarski(X1,esk3_3(X1,X2,k3_xboole_0(X2,X3)))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_12, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_9])).
fof(c_0_13, negated_conjecture, k2_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=esk1_0, inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_14, plain, (X1=k2_xboole_0(X2,X3)|~r1_tarski(X1,esk3_3(X2,X1,X3))|~r1_tarski(X3,X1)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1|r1_tarski(X1,esk3_3(X1,X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, negated_conjecture, (k2_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_8]), c_0_12])])).
cnf(c_0_18, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 19
# Proof object clause steps            : 10
# Proof object formula steps           : 9
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 8
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 8
# Processed clauses                    : 21
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 20
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 36
# ...of the previous two non-trivial   : 35
# Contextual simplify-reflections      : 0
# Paramodulations                      : 34
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 14
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 9
# Current number of unprocessed clauses: 21
# ...number of literals in the above   : 65
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4
# Clause-clause subsumption calls (NU) : 43
# Rec. Clause-clause subsumption calls : 41
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 921

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 28688 pages
