# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t27_waybel_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v12_waybel_0(X2,X1)&v12_waybel_0(X3,X1))=>(v12_waybel_0(k9_subset_1(u1_struct_0(X1),X2,X3),X1)&v12_waybel_0(k4_subset_1(u1_struct_0(X1),X2,X3),X1)))))), file('data/MPTP/proofs/a859b0bc456e4f1d8259e73563fc15b9.E_input', t27_waybel_0)).
fof(d19_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X4,X3))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/a859b0bc456e4f1d8259e73563fc15b9.E_input', d19_waybel_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/a859b0bc456e4f1d8259e73563fc15b9.E_input', d4_xboole_0)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/a859b0bc456e4f1d8259e73563fc15b9.E_input', d3_xboole_0)).
fof(dt_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>m1_subset_1(k4_subset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/a859b0bc456e4f1d8259e73563fc15b9.E_input', dt_k4_subset_1)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('data/MPTP/proofs/a859b0bc456e4f1d8259e73563fc15b9.E_input', redefinition_k4_subset_1)).
fof(dt_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>m1_subset_1(k9_subset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('data/MPTP/proofs/a859b0bc456e4f1d8259e73563fc15b9.E_input', dt_k9_subset_1)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/a859b0bc456e4f1d8259e73563fc15b9.E_input', redefinition_k9_subset_1)).
fof(c_0_8, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v12_waybel_0(X2,X1)&v12_waybel_0(X3,X1))=>(v12_waybel_0(k9_subset_1(u1_struct_0(X1),X2,X3),X1)&v12_waybel_0(k4_subset_1(u1_struct_0(X1),X2,X3),X1))))))), inference(assume_negation,[status(cth)],[t27_waybel_0])).
fof(c_0_9, plain, ![X8, X9, X10, X11]:((~v12_waybel_0(X9,X8)|(~m1_subset_1(X10,u1_struct_0(X8))|(~m1_subset_1(X11,u1_struct_0(X8))|(~r2_hidden(X10,X9)|~r1_orders_2(X8,X11,X10)|r2_hidden(X11,X9))))|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_orders_2(X8))&((m1_subset_1(esk4_2(X8,X9),u1_struct_0(X8))|v12_waybel_0(X9,X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_orders_2(X8))&((m1_subset_1(esk5_2(X8,X9),u1_struct_0(X8))|v12_waybel_0(X9,X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_orders_2(X8))&(((r2_hidden(esk4_2(X8,X9),X9)|v12_waybel_0(X9,X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_orders_2(X8))&(r1_orders_2(X8,esk5_2(X8,X9),esk4_2(X8,X9))|v12_waybel_0(X9,X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_orders_2(X8)))&(~r2_hidden(esk5_2(X8,X9),X9)|v12_waybel_0(X9,X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|~l1_orders_2(X8)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_waybel_0])])])])])).
fof(c_0_10, negated_conjecture, (l1_orders_2(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&((v12_waybel_0(esk2_0,esk1_0)&v12_waybel_0(esk3_0,esk1_0))&(~v12_waybel_0(k9_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0),esk1_0)|~v12_waybel_0(k4_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0),esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_11, plain, (r2_hidden(X4,X1)|~v12_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, negated_conjecture, (v12_waybel_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X26, X27, X28, X29, X30, X31, X32, X33]:((((r2_hidden(X29,X26)|~r2_hidden(X29,X28)|X28!=k3_xboole_0(X26,X27))&(r2_hidden(X29,X27)|~r2_hidden(X29,X28)|X28!=k3_xboole_0(X26,X27)))&(~r2_hidden(X30,X26)|~r2_hidden(X30,X27)|r2_hidden(X30,X28)|X28!=k3_xboole_0(X26,X27)))&((~r2_hidden(esk6_3(X31,X32,X33),X33)|(~r2_hidden(esk6_3(X31,X32,X33),X31)|~r2_hidden(esk6_3(X31,X32,X33),X32))|X33=k3_xboole_0(X31,X32))&((r2_hidden(esk6_3(X31,X32,X33),X31)|r2_hidden(esk6_3(X31,X32,X33),X33)|X33=k3_xboole_0(X31,X32))&(r2_hidden(esk6_3(X31,X32,X33),X32)|r2_hidden(esk6_3(X31,X32,X33),X33)|X33=k3_xboole_0(X31,X32))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_16, negated_conjecture, (r2_hidden(X1,esk3_0)|~r1_orders_2(esk1_0,X1,X2)|~r2_hidden(X2,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_17, plain, (r1_orders_2(X1,esk5_2(X1,X2),esk4_2(X1,X2))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_19, negated_conjecture, (v12_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X4!=k3_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk3_0)|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),esk3_0)|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_14])])).
cnf(c_0_22, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_orders_2(esk1_0,X1,X2)|~r2_hidden(X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_18]), c_0_19]), c_0_14])])).
cnf(c_0_24, plain, (v12_waybel_0(X2,X1)|~r2_hidden(esk5_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, plain, (r2_hidden(X1,k3_xboole_0(X2,X3))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_26, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk3_0)|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),esk3_0)|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_14])])).
cnf(c_0_27, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk2_0)|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),esk2_0)|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_17]), c_0_14])])).
cnf(c_0_29, plain, (v12_waybel_0(k3_xboole_0(X1,X2),X3)|~r2_hidden(esk5_2(X3,k3_xboole_0(X1,X2)),X2)|~r2_hidden(esk5_2(X3,k3_xboole_0(X1,X2)),X1)|~m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk3_0)|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_14])])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk2_0)|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),esk2_0)|~m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_22]), c_0_14])])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_33, plain, ![X35, X36, X37, X38, X39, X40, X41, X42]:(((~r2_hidden(X38,X37)|(r2_hidden(X38,X35)|r2_hidden(X38,X36))|X37!=k2_xboole_0(X35,X36))&((~r2_hidden(X39,X35)|r2_hidden(X39,X37)|X37!=k2_xboole_0(X35,X36))&(~r2_hidden(X39,X36)|r2_hidden(X39,X37)|X37!=k2_xboole_0(X35,X36))))&(((~r2_hidden(esk7_3(X40,X41,X42),X40)|~r2_hidden(esk7_3(X40,X41,X42),X42)|X42=k2_xboole_0(X40,X41))&(~r2_hidden(esk7_3(X40,X41,X42),X41)|~r2_hidden(esk7_3(X40,X41,X42),X42)|X42=k2_xboole_0(X40,X41)))&(r2_hidden(esk7_3(X40,X41,X42),X42)|(r2_hidden(esk7_3(X40,X41,X42),X40)|r2_hidden(esk7_3(X40,X41,X42),X41))|X42=k2_xboole_0(X40,X41)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_34, negated_conjecture, (v12_waybel_0(k3_xboole_0(X1,esk3_0),esk1_0)|~r2_hidden(esk4_2(esk1_0,k3_xboole_0(X1,esk3_0)),esk3_0)|~r2_hidden(esk5_2(esk1_0,k3_xboole_0(X1,esk3_0)),X1)|~m1_subset_1(k3_xboole_0(X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_14])])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk2_0)|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_27]), c_0_14])])).
cnf(c_0_36, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (r2_hidden(esk4_2(X1,X2),X2)|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_38, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_39, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_40, plain, ![X14, X15, X16]:(~m1_subset_1(X15,k1_zfmisc_1(X14))|~m1_subset_1(X16,k1_zfmisc_1(X14))|m1_subset_1(k4_subset_1(X14,X15,X16),k1_zfmisc_1(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_subset_1])])).
fof(c_0_41, plain, ![X20, X21, X22]:(~m1_subset_1(X21,k1_zfmisc_1(X20))|~m1_subset_1(X22,k1_zfmisc_1(X20))|k4_subset_1(X20,X21,X22)=k2_xboole_0(X21,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
cnf(c_0_42, negated_conjecture, (v12_waybel_0(k3_xboole_0(esk2_0,esk3_0),esk1_0)|~r2_hidden(esk4_2(esk1_0,k3_xboole_0(esk2_0,esk3_0)),esk3_0)|~r2_hidden(esk4_2(esk1_0,k3_xboole_0(esk2_0,esk3_0)),esk2_0)|~m1_subset_1(k3_xboole_0(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, plain, (r2_hidden(esk4_2(X1,k3_xboole_0(X2,X3)),X3)|v12_waybel_0(k3_xboole_0(X2,X3),X1)|~m1_subset_1(k3_xboole_0(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_38])).
fof(c_0_45, plain, ![X17, X18, X19]:(~m1_subset_1(X19,k1_zfmisc_1(X17))|m1_subset_1(k9_subset_1(X17,X18,X19),k1_zfmisc_1(X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k9_subset_1])])).
fof(c_0_46, plain, ![X23, X24, X25]:(~m1_subset_1(X25,k1_zfmisc_1(X23))|k9_subset_1(X23,X24,X25)=k3_xboole_0(X24,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_48, plain, (m1_subset_1(k4_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50, negated_conjecture, (~v12_waybel_0(k9_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0),esk1_0)|~v12_waybel_0(k4_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_51, negated_conjecture, (v12_waybel_0(k3_xboole_0(esk2_0,esk3_0),esk1_0)|~r2_hidden(esk4_2(esk1_0,k3_xboole_0(esk2_0,esk3_0)),esk2_0)|~m1_subset_1(k3_xboole_0(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_14])])).
cnf(c_0_52, plain, (r2_hidden(esk4_2(X1,k3_xboole_0(X2,X3)),X2)|v12_waybel_0(k3_xboole_0(X2,X3),X1)|~m1_subset_1(k3_xboole_0(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_44, c_0_37])).
cnf(c_0_53, plain, (m1_subset_1(k9_subset_1(X2,X3,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_56, plain, (r2_hidden(esk4_2(X1,k2_xboole_0(X2,X3)),X3)|r2_hidden(esk4_2(X1,k2_xboole_0(X2,X3)),X2)|v12_waybel_0(k2_xboole_0(X2,X3),X1)|~m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_47, c_0_37])).
cnf(c_0_57, plain, (m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_58, negated_conjecture, (~v12_waybel_0(k9_subset_1(u1_struct_0(esk1_0),esk2_0,esk3_0),esk1_0)|~v12_waybel_0(k2_xboole_0(esk2_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_49]), c_0_12]), c_0_18])])).
cnf(c_0_59, negated_conjecture, (v12_waybel_0(k3_xboole_0(esk2_0,esk3_0),esk1_0)|~m1_subset_1(k3_xboole_0(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_14])])).
cnf(c_0_60, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_62, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (r2_hidden(esk4_2(X1,k2_xboole_0(X2,X3)),X2)|r2_hidden(esk4_2(X1,k2_xboole_0(X2,X3)),X3)|v12_waybel_0(k2_xboole_0(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_64, negated_conjecture, (~v12_waybel_0(k3_xboole_0(esk2_0,esk3_0),esk1_0)|~v12_waybel_0(k2_xboole_0(esk2_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_54]), c_0_12])])).
cnf(c_0_65, negated_conjecture, (v12_waybel_0(k3_xboole_0(esk2_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_12])])).
cnf(c_0_66, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_61])).
cnf(c_0_67, plain, (v12_waybel_0(k2_xboole_0(X1,X2),X3)|~r2_hidden(esk5_2(X3,k2_xboole_0(X1,X2)),X2)|~m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_24, c_0_62])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk4_2(esk1_0,k2_xboole_0(X1,esk3_0)),esk3_0)|r2_hidden(esk4_2(esk1_0,k2_xboole_0(X1,esk3_0)),X1)|v12_waybel_0(k2_xboole_0(X1,esk3_0),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_12]), c_0_14])])).
cnf(c_0_69, negated_conjecture, (~v12_waybel_0(k2_xboole_0(esk2_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_65])])).
cnf(c_0_70, plain, (v12_waybel_0(k2_xboole_0(X1,X2),X3)|~r2_hidden(esk5_2(X3,k2_xboole_0(X1,X2)),X1)|~m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_24, c_0_66])).
cnf(c_0_71, negated_conjecture, (v12_waybel_0(k2_xboole_0(X1,esk3_0),esk1_0)|~r2_hidden(esk4_2(esk1_0,k2_xboole_0(X1,esk3_0)),esk3_0)|~m1_subset_1(k2_xboole_0(X1,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_30]), c_0_14])])).
cnf(c_0_72, negated_conjecture, (r2_hidden(esk4_2(esk1_0,k2_xboole_0(esk2_0,esk3_0)),esk2_0)|r2_hidden(esk4_2(esk1_0,k2_xboole_0(esk2_0,esk3_0)),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_18]), c_0_69])).
cnf(c_0_73, negated_conjecture, (v12_waybel_0(k2_xboole_0(esk2_0,X1),esk1_0)|~r2_hidden(esk4_2(esk1_0,k2_xboole_0(esk2_0,X1)),esk2_0)|~m1_subset_1(k2_xboole_0(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_35]), c_0_14])])).
cnf(c_0_74, negated_conjecture, (r2_hidden(esk4_2(esk1_0,k2_xboole_0(esk2_0,esk3_0)),esk2_0)|~m1_subset_1(k2_xboole_0(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_69])).
cnf(c_0_75, negated_conjecture, (~m1_subset_1(k2_xboole_0(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_69])).
cnf(c_0_76, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_57]), c_0_12]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 77
# Proof object clause steps            : 60
# Proof object formula steps           : 17
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 51
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 10037
# ...of these trivial                  : 116
# ...subsumed                          : 8835
# ...remaining for further processing  : 1086
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 5
# Generated clauses                    : 179052
# ...of the previous two non-trivial   : 171903
# Contextual simplify-reflections      : 5
# Paramodulations                      : 177948
# Factorizations                       : 1098
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1024
#    Positive orientable unit clauses  : 88
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 934
# Current number of unprocessed clauses: 161704
# ...number of literals in the above   : 528651
# Current number of archived formulas  : 0
# Current number of archived clauses   : 56
# Clause-clause subsumption calls (NU) : 1004521
# Rec. Clause-clause subsumption calls : 385475
# Non-unit clause-clause subsumptions  : 8863
# Unit Clause-clause subsumption calls : 23888
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 431
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3459657

# -------------------------------------------------
# User time                : 1.599 s
# System time              : 0.073 s
# Total time               : 1.672 s
# Maximum resident set size: 112648 pages
