# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t2_boole)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t48_xboole_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t2_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t3_subset)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', fc1_subset_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t4_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t3_boole)).
fof(t71_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t71_mcart_1)).
fof(t3_ordinal1, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(r2_hidden(X2,X3)&r2_hidden(X3,X1)))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t3_ordinal1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', d3_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', d10_xboole_0)).
fof(t34_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_mcart_1(X3,X4,X5,X6))))))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t34_mcart_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', d1_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t3_xboole_0)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', d4_zfmisc_1)).
fof(t127_zfmisc_1, axiom, ![X1, X2, X3, X4]:((r1_xboole_0(X1,X2)|r1_xboole_0(X3,X4))=>r1_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t127_zfmisc_1)).
fof(t55_mcart_1, axiom, ![X1, X2, X3, X4]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(X4=k1_xboole_0))))<=>~(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0)), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t55_mcart_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t2_xboole_1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t10_mcart_1)).
fof(l139_zfmisc_1, axiom, ![X1, X2, X3]:~((r2_hidden(X1,k2_zfmisc_1(X2,X3))&![X4, X5]:~(k4_tarski(X4,X5)=X1))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', l139_zfmisc_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t7_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', d3_mcart_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t1_subset)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/a94aae28f2ef4978ac4f15319c24ce81.E_input', t50_mcart_1)).
fof(c_0_24, plain, ![X53]:k3_xboole_0(X53,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_25, plain, ![X85, X86]:k4_xboole_0(X85,k4_xboole_0(X85,X86))=k3_xboole_0(X85,X86), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_26, plain, ![X43, X44]:(~m1_subset_1(X43,X44)|(v1_xboole_0(X44)|r2_hidden(X43,X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_27, plain, ![X73, X74]:((~m1_subset_1(X73,k1_zfmisc_1(X74))|r1_tarski(X73,X74))&(~r1_tarski(X73,X74)|m1_subset_1(X73,k1_zfmisc_1(X74)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_28, plain, ![X87]:~v1_xboole_0(k1_zfmisc_1(X87)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_29, plain, ![X23, X24, X26, X27, X28]:((r1_xboole_0(X23,X24)|r2_hidden(esk6_2(X23,X24),k3_xboole_0(X23,X24)))&(~r2_hidden(X28,k3_xboole_0(X26,X27))|~r1_xboole_0(X26,X27))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_30, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X68]:k4_xboole_0(X68,k1_xboole_0)=X68, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_33, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X8))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k7_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t71_mcart_1])).
fof(c_0_34, plain, ![X50, X51, X52]:(~r2_hidden(X50,X51)|(~r2_hidden(X51,X52)|~r2_hidden(X52,X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_ordinal1])])).
cnf(c_0_35, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_41, negated_conjecture, ![X14, X15, X16]:(m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&((~m1_subset_1(X14,esk1_0)|(~m1_subset_1(X15,esk2_0)|(~m1_subset_1(X16,esk3_0)|(esk5_0!=k3_mcart_1(X14,X15,X16)|esk4_0=X16))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&esk4_0!=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])])).
fof(c_0_42, plain, ![X82, X83, X84]:k3_zfmisc_1(X82,X83,X84)=k2_zfmisc_1(k2_zfmisc_1(X82,X83),X84), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
cnf(c_0_43, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
fof(c_0_45, plain, ![X66, X67]:(((r1_tarski(X66,X67)|X66!=X67)&(r1_tarski(X67,X66)|X66!=X67))&(~r1_tarski(X66,X67)|~r1_tarski(X67,X66)|X66=X67)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_46, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_38, c_0_31])).
cnf(c_0_47, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_48, plain, ![X75, X77, X78, X79, X80]:((r2_hidden(esk11_1(X75),X75)|X75=k1_xboole_0)&((~r2_hidden(X77,X75)|esk11_1(X75)!=k4_mcart_1(X77,X78,X79,X80)|X75=k1_xboole_0)&(~r2_hidden(X78,X75)|esk11_1(X75)!=k4_mcart_1(X77,X78,X79,X80)|X75=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_mcart_1])])])])])])])).
fof(c_0_49, plain, ![X57, X58, X59]:((~v1_xboole_0(X57)|~r2_hidden(X58,X57))&(r2_hidden(esk10_1(X59),X59)|v1_xboole_0(X59))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_50, plain, ![X33, X34, X36, X37, X38]:(((r2_hidden(esk7_2(X33,X34),X33)|r1_xboole_0(X33,X34))&(r2_hidden(esk7_2(X33,X34),X34)|r1_xboole_0(X33,X34)))&(~r2_hidden(X38,X36)|~r2_hidden(X38,X37)|~r1_xboole_0(X36,X37))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (~r1_tarski(X1,X2)|~r2_hidden(k1_zfmisc_1(X2),X3)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_54, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_55, plain, ![X19, X20, X21, X22]:k4_zfmisc_1(X19,X20,X21,X22)=k2_zfmisc_1(k3_zfmisc_1(X19,X20,X21),X22), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
cnf(c_0_56, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_40])).
cnf(c_0_57, plain, (r2_hidden(esk11_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_58, plain, ![X29, X30, X31, X32]:((~r1_xboole_0(X29,X30)|r1_xboole_0(k2_zfmisc_1(X29,X31),k2_zfmisc_1(X30,X32)))&(~r1_xboole_0(X31,X32)|r1_xboole_0(k2_zfmisc_1(X29,X31),k2_zfmisc_1(X30,X32)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t127_zfmisc_1])])])).
cnf(c_0_59, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (r2_hidden(esk7_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61, negated_conjecture, (m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_62, plain, (~r1_tarski(k1_zfmisc_1(X1),X2)|~r1_tarski(X3,X1)|~r2_hidden(k1_zfmisc_1(X2),X3)), inference(spm,[status(thm)],[c_0_53, c_0_44])).
cnf(c_0_63, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_54])).
fof(c_0_64, plain, ![X39, X40, X41, X42]:((X39=k1_xboole_0|(X40=k1_xboole_0|(X41=k1_xboole_0|X42=k1_xboole_0))|k4_zfmisc_1(X39,X40,X41,X42)!=k1_xboole_0)&((X39!=k1_xboole_0|k4_zfmisc_1(X39,X40,X41,X42)=k1_xboole_0)&((X40!=k1_xboole_0|k4_zfmisc_1(X39,X40,X41,X42)=k1_xboole_0)&((X41!=k1_xboole_0|k4_zfmisc_1(X39,X40,X41,X42)=k1_xboole_0)&(X42!=k1_xboole_0|k4_zfmisc_1(X39,X40,X41,X42)=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_mcart_1])])])])).
cnf(c_0_65, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_66, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_67, plain, (r1_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_68, plain, (r1_xboole_0(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_35, c_0_61])).
cnf(c_0_70, plain, (~r1_tarski(X1,X2)|~r2_hidden(k1_zfmisc_1(k1_zfmisc_1(X2)),X1)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_71, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k4_zfmisc_1(X1,X2,X3,X4)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_65, c_0_52])).
cnf(c_0_73, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_74, negated_conjecture, (r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r1_xboole_0(X1,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, plain, (~r1_tarski(k1_zfmisc_1(k1_zfmisc_1(X1)),X2)|~r1_tarski(k1_zfmisc_1(X2),X1)), inference(spm,[status(thm)],[c_0_70, c_0_44])).
cnf(c_0_76, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_77, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),X1)=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_78, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_79, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_80, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_81, plain, ![X81]:r1_tarski(k1_xboole_0,X81), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_82, plain, ![X61, X62, X63]:((r2_hidden(k1_mcart_1(X61),X62)|~r2_hidden(X61,k2_zfmisc_1(X62,X63)))&(r2_hidden(k2_mcart_1(X61),X63)|~r2_hidden(X61,k2_zfmisc_1(X62,X63)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_83, plain, (~r1_tarski(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))),X1)), inference(spm,[status(thm)],[c_0_75, c_0_63])).
cnf(c_0_84, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), c_0_79]), c_0_80])).
cnf(c_0_85, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_86, plain, ![X45, X46, X47]:(~r2_hidden(X45,k2_zfmisc_1(X46,X47))|k4_tarski(esk8_1(X45),esk9_1(X45))=X45), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l139_zfmisc_1])])])])])])).
cnf(c_0_87, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_88, negated_conjecture, (r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_85])])).
fof(c_0_89, plain, ![X17, X18]:(k1_mcart_1(k4_tarski(X17,X18))=X17&k2_mcart_1(k4_tarski(X17,X18))=X18), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_90, plain, (k4_tarski(esk8_1(X1),esk9_1(X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_91, negated_conjecture, (r2_hidden(k1_mcart_1(esk5_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_92, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_93, negated_conjecture, (k4_tarski(esk8_1(k1_mcart_1(esk5_0)),esk9_1(k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_94, negated_conjecture, (esk8_1(k1_mcart_1(esk5_0))=k1_mcart_1(k1_mcart_1(esk5_0))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
fof(c_0_95, plain, ![X54, X55, X56]:k3_mcart_1(X54,X55,X56)=k4_tarski(k4_tarski(X54,X55),X56), inference(variable_rename,[status(thm)],[d3_mcart_1])).
cnf(c_0_96, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_97, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk5_0)),esk9_1(k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)), inference(rw,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_98, negated_conjecture, (esk4_0=X3|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|esk5_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_99, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_100, negated_conjecture, (esk9_1(k1_mcart_1(esk5_0))=k2_mcart_1(k1_mcart_1(esk5_0))), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_101, negated_conjecture, (k4_tarski(esk8_1(esk5_0),esk9_1(esk5_0))=esk5_0), inference(spm,[status(thm)],[c_0_90, c_0_88])).
cnf(c_0_102, negated_conjecture, (esk4_0=X3|esk5_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_103, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk5_0)),k2_mcart_1(k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)), inference(rw,[status(thm)],[c_0_97, c_0_100])).
fof(c_0_104, plain, ![X64, X65]:(~r2_hidden(X64,X65)|m1_subset_1(X64,X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_105, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_106, negated_conjecture, (esk8_1(esk5_0)=k1_mcart_1(esk5_0)), inference(spm,[status(thm)],[c_0_92, c_0_101])).
fof(c_0_107, plain, ![X69, X70, X71, X72]:((k5_mcart_1(X69,X70,X71,X72)=k1_mcart_1(k1_mcart_1(X72))|~m1_subset_1(X72,k3_zfmisc_1(X69,X70,X71))|X71=k1_xboole_0|X70=k1_xboole_0|X69=k1_xboole_0)&((k6_mcart_1(X69,X70,X71,X72)=k2_mcart_1(k1_mcart_1(X72))|~m1_subset_1(X72,k3_zfmisc_1(X69,X70,X71))|X71=k1_xboole_0|X70=k1_xboole_0|X69=k1_xboole_0)&(k7_mcart_1(X69,X70,X71,X72)=k2_mcart_1(X72)|~m1_subset_1(X72,k3_zfmisc_1(X69,X70,X71))|X71=k1_xboole_0|X70=k1_xboole_0|X69=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_108, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_109, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_110, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)), inference(spm,[status(thm)],[c_0_105, c_0_91])).
cnf(c_0_111, negated_conjecture, (k4_tarski(k1_mcart_1(esk5_0),esk9_1(esk5_0))=esk5_0), inference(rw,[status(thm)],[c_0_101, c_0_106])).
cnf(c_0_112, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_113, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])])).
cnf(c_0_114, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)), inference(spm,[status(thm)],[c_0_87, c_0_91])).
cnf(c_0_115, negated_conjecture, (esk9_1(esk5_0)=k2_mcart_1(esk5_0)), inference(spm,[status(thm)],[c_0_96, c_0_111])).
cnf(c_0_116, negated_conjecture, (esk4_0!=k7_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_117, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_112, c_0_52])).
cnf(c_0_118, negated_conjecture, (esk4_0=X1|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_109]), c_0_114])])).
cnf(c_0_119, negated_conjecture, (k4_tarski(k1_mcart_1(esk5_0),k2_mcart_1(esk5_0))=esk5_0), inference(rw,[status(thm)],[c_0_111, c_0_115])).
cnf(c_0_120, negated_conjecture, (k2_mcart_1(esk5_0)!=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_61])]), c_0_78]), c_0_79]), c_0_80])).
cnf(c_0_121, negated_conjecture, (~m1_subset_1(k2_mcart_1(esk5_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_120])).
cnf(c_0_122, negated_conjecture, (r2_hidden(k2_mcart_1(esk5_0),esk3_0)), inference(spm,[status(thm)],[c_0_105, c_0_88])).
cnf(c_0_123, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_109]), c_0_122])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 124
# Proof object clause steps            : 75
# Proof object formula steps           : 49
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 47
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 43
# Processed clauses                    : 888
# ...of these trivial                  : 24
# ...subsumed                          : 424
# ...remaining for further processing  : 440
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 61
# Backward-rewritten                   : 30
# Generated clauses                    : 2100
# ...of the previous two non-trivial   : 1750
# Contextual simplify-reflections      : 4
# Paramodulations                      : 2094
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 301
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 27
#    Non-unit-clauses                  : 248
# Current number of unprocessed clauses: 707
# ...number of literals in the above   : 1889
# Current number of archived formulas  : 0
# Current number of archived clauses   : 137
# Clause-clause subsumption calls (NU) : 9401
# Rec. Clause-clause subsumption calls : 7114
# Non-unit clause-clause subsumptions  : 236
# Unit Clause-clause subsumption calls : 546
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 32
# BW rewrite match successes           : 21
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 25201

# -------------------------------------------------
# User time                : 0.029 s
# System time              : 0.003 s
# Total time               : 0.032 s
# Maximum resident set size: 28796 pages
