# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t75_enumset1)).
fof(l3_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2))), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', l3_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t69_enumset1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', l51_zfmisc_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', d3_tarski)).
fof(t95_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t95_xboole_1)).
fof(t43_zfmisc_1, conjecture, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t43_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', idempotence_k3_xboole_0)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', idempotence_k2_xboole_0)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t12_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t7_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t17_xboole_1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t92_xboole_1)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', t91_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', d10_xboole_0)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/a993f6cc7d474010844114747397136d.E_input', commutativity_k5_xboole_0)).
fof(c_0_23, plain, ![X40, X41, X42, X43, X44, X45, X46, X47]:(((~r2_hidden(X43,X42)|(X43=X40|X43=X41)|X42!=k2_tarski(X40,X41))&((X44!=X40|r2_hidden(X44,X42)|X42!=k2_tarski(X40,X41))&(X44!=X41|r2_hidden(X44,X42)|X42!=k2_tarski(X40,X41))))&(((esk1_3(X45,X46,X47)!=X45|~r2_hidden(esk1_3(X45,X46,X47),X47)|X47=k2_tarski(X45,X46))&(esk1_3(X45,X46,X47)!=X46|~r2_hidden(esk1_3(X45,X46,X47),X47)|X47=k2_tarski(X45,X46)))&(r2_hidden(esk1_3(X45,X46,X47),X47)|(esk1_3(X45,X46,X47)=X45|esk1_3(X45,X46,X47)=X46)|X47=k2_tarski(X45,X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_24, plain, ![X11, X12]:k1_enumset1(X11,X11,X12)=k2_tarski(X11,X12), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_25, plain, ![X13, X14, X15]:k2_enumset1(X13,X13,X14,X15)=k1_enumset1(X13,X14,X15), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_26, plain, ![X16, X17, X18, X19]:k3_enumset1(X16,X16,X17,X18,X19)=k2_enumset1(X16,X17,X18,X19), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_27, plain, ![X20, X21, X22, X23, X24]:k4_enumset1(X20,X20,X21,X22,X23,X24)=k3_enumset1(X20,X21,X22,X23,X24), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_28, plain, ![X25, X26, X27, X28, X29, X30]:k5_enumset1(X25,X25,X26,X27,X28,X29,X30)=k4_enumset1(X25,X26,X27,X28,X29,X30), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_29, plain, ![X31, X32, X33, X34, X35, X36, X37]:k6_enumset1(X31,X31,X32,X33,X34,X35,X36,X37)=k5_enumset1(X31,X32,X33,X34,X35,X36,X37), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_30, plain, ![X8, X9]:((~r1_tarski(X8,k1_tarski(X9))|(X8=k1_xboole_0|X8=k1_tarski(X9)))&((X8!=k1_xboole_0|r1_tarski(X8,k1_tarski(X9)))&(X8!=k1_tarski(X9)|r1_tarski(X8,k1_tarski(X9))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_zfmisc_1])])])).
fof(c_0_31, plain, ![X10]:k2_tarski(X10,X10)=k1_tarski(X10), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_32, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (r1_tarski(X1,k1_tarski(X2))|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_41, plain, ![X38, X39]:k3_tarski(k2_tarski(X38,X39))=k2_xboole_0(X38,X39), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
cnf(c_0_42, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_43, plain, ![X69, X70, X71]:((~v1_xboole_0(X69)|~r2_hidden(X70,X69))&(r2_hidden(esk6_1(X71),X71)|v1_xboole_0(X71))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_44, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X4,X4,X4,X4,X4,X4,X4,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38])).
fof(c_0_45, plain, ![X54, X55, X56, X57, X58]:((~r1_tarski(X54,X55)|(~r2_hidden(X56,X54)|r2_hidden(X56,X55)))&((r2_hidden(esk5_2(X57,X58),X57)|r1_tarski(X57,X58))&(~r2_hidden(esk5_2(X57,X58),X58)|r1_tarski(X57,X58)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_46, plain, (r1_tarski(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))|X1!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_40]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38])).
fof(c_0_47, plain, ![X49, X50]:k3_xboole_0(X49,X50)=k5_xboole_0(k5_xboole_0(X49,X50),k2_xboole_0(X49,X50)), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_48, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_49, negated_conjecture, ~(![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0))))))), inference(assume_negation,[status(cth)],[t43_zfmisc_1])).
cnf(c_0_50, plain, (X1=X4|X1=X3|X2!=k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_51, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_44])])).
cnf(c_0_53, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, plain, (r1_tarski(k1_xboole_0,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))), inference(er,[status(thm)],[c_0_46])).
fof(c_0_55, plain, ![X60]:k3_xboole_0(X60,X60)=X60, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_56, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_48, c_0_33])).
fof(c_0_58, plain, ![X61]:k2_xboole_0(X61,X61)=X61, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_59, negated_conjecture, (k1_tarski(esk2_0)=k2_xboole_0(esk3_0,esk4_0)&((esk3_0!=k1_tarski(esk2_0)|esk4_0!=k1_tarski(esk2_0))&((esk3_0!=k1_xboole_0|esk4_0!=k1_tarski(esk2_0))&(esk3_0!=k1_tarski(esk2_0)|esk4_0!=k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_49])])])).
fof(c_0_60, plain, ![X62, X63]:(~r1_tarski(X62,X63)|k2_xboole_0(X62,X63)=X63), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_61, plain, (X1=X2|X1=X3|~r2_hidden(X1,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2))), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_62, plain, (r2_hidden(esk6_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_63, plain, (~v1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_64, plain, (r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
fof(c_0_65, plain, ![X64, X65]:r1_tarski(X64,k2_xboole_0(X64,X65)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_66, plain, ![X66, X67]:r1_tarski(k3_xboole_0(X66,X67),X66), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_67, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_68, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_57]), c_0_34]), c_0_35])).
fof(c_0_69, plain, ![X68]:k5_xboole_0(X68,X68)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
cnf(c_0_70, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_71, negated_conjecture, (k1_tarski(esk2_0)=k2_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_72, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_73, plain, (esk6_1(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))=X2|esk6_1(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
cnf(c_0_74, plain, (X1=X2|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_61, c_0_64])).
cnf(c_0_75, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_76, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_77, plain, ![X73, X74, X75]:k5_xboole_0(k5_xboole_0(X73,X74),X75)=k5_xboole_0(X73,k5_xboole_0(X74,X75)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_78, plain, (k5_xboole_0(k5_xboole_0(X1,X1),k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_79, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_80, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_57]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_81, negated_conjecture, (esk3_0!=k1_tarski(esk2_0)|esk4_0!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_82, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=k3_tarski(k6_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_40]), c_0_33]), c_0_57]), c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38])).
cnf(c_0_83, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_57]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_84, plain, (esk6_1(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))=X1), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_73])])).
cnf(c_0_85, plain, (esk6_1(k1_xboole_0)=X1|v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_74, c_0_62])).
fof(c_0_86, plain, ![X76, X77]:(((r1_tarski(X76,X77)|X76!=X77)&(r1_tarski(X77,X76)|X76!=X77))&(~r1_tarski(X76,X77)|~r1_tarski(X77,X76)|X76=X77)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_87, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|~r1_tarski(X1,k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_88, plain, (r1_tarski(X1,k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_57]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_89, plain, (r1_tarski(k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_68]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_90, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_91, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_79]), c_0_80])).
fof(c_0_92, plain, ![X78, X79]:k5_xboole_0(X78,X79)=k5_xboole_0(X79,X78), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_93, negated_conjecture, (esk3_0!=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)|esk4_0!=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_40]), c_0_40]), c_0_33]), c_0_33]), c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38])).
cnf(c_0_94, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=esk4_0|~r1_tarski(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_95, plain, (esk6_1(esk6_1(k1_xboole_0))=X1|v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_96, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_97, plain, (X1=k1_xboole_0|X1=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2)|~r1_tarski(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_40]), c_0_40]), c_0_33]), c_0_33]), c_0_34]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38])).
cnf(c_0_98, negated_conjecture, (r1_tarski(esk3_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_88, c_0_82])).
cnf(c_0_99, plain, (r1_tarski(k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)))),X1)), inference(rw,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_100, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_79]), c_0_91])).
cnf(c_0_101, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_102, negated_conjecture, (esk3_0!=k1_tarski(esk2_0)|esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_103, negated_conjecture, (esk3_0!=k1_xboole_0|esk4_0!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_104, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_105, negated_conjecture, (esk4_0!=esk3_0|~r1_tarski(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_106, plain, (X1=X2|v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_95, c_0_95])).
cnf(c_0_107, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_96])).
cnf(c_0_108, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)=esk3_0|k1_xboole_0=esk3_0), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_109, negated_conjecture, (r1_tarski(k5_xboole_0(esk3_0,k5_xboole_0(esk4_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))),esk3_0)), inference(spm,[status(thm)],[c_0_99, c_0_82])).
cnf(c_0_110, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_111, negated_conjecture, (esk4_0!=k1_xboole_0|esk3_0!=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_40]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_112, negated_conjecture, (esk3_0!=k1_xboole_0|esk4_0!=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_40]), c_0_33]), c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_113, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_51, c_0_104])).
cnf(c_0_114, negated_conjecture, (v1_xboole_0(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106])]), c_0_107])])).
cnf(c_0_115, negated_conjecture, (k1_xboole_0=esk3_0|X1=k1_xboole_0|X1=esk3_0|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_97, c_0_108])).
cnf(c_0_116, negated_conjecture, (k1_xboole_0=esk3_0|r1_tarski(esk4_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_108]), c_0_110])).
cnf(c_0_117, negated_conjecture, (k1_xboole_0=esk3_0|esk4_0!=esk3_0), inference(spm,[status(thm)],[c_0_93, c_0_108])).
cnf(c_0_118, negated_conjecture, (k1_xboole_0=esk3_0|esk4_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_111, c_0_108])).
cnf(c_0_119, negated_conjecture, (k1_xboole_0!=esk3_0|~r1_tarski(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_112, c_0_94])).
cnf(c_0_120, negated_conjecture, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_121, negated_conjecture, (k1_xboole_0=esk3_0), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117]), c_0_118])).
cnf(c_0_122, negated_conjecture, (~r1_tarski(esk3_0,esk4_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_108]), c_0_119]), c_0_105])).
cnf(c_0_123, negated_conjecture, (r1_tarski(esk3_0,X1)), inference(rw,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_124, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 125
# Proof object clause steps            : 78
# Proof object formula steps           : 47
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 126
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 29
# Processed clauses                    : 303
# ...of these trivial                  : 5
# ...subsumed                          : 119
# ...remaining for further processing  : 179
# Other redundant clauses eliminated   : 94
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 72
# Generated clauses                    : 2226
# ...of the previous two non-trivial   : 1919
# Contextual simplify-reflections      : 7
# Paramodulations                      : 2117
# Factorizations                       : 17
# Equation resolutions                 : 94
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 29
# Current number of unprocessed clauses: 1225
# ...number of literals in the above   : 3973
# Current number of archived formulas  : 0
# Current number of archived clauses   : 131
# Clause-clause subsumption calls (NU) : 1293
# Rec. Clause-clause subsumption calls : 801
# Non-unit clause-clause subsumptions  : 113
# Unit Clause-clause subsumption calls : 107
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 50
# BW rewrite match successes           : 34
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20252

# -------------------------------------------------
# User time                : 0.025 s
# System time              : 0.000 s
# Total time               : 0.025 s
# Maximum resident set size: 26384 pages
