# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t70_enumset1)).
fof(t43_zfmisc_1, conjecture, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t43_zfmisc_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t36_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t100_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t7_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t75_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t69_enumset1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t2_boole)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t5_boole)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t3_xboole_1)).
fof(t108_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(k3_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t108_xboole_1)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t1_xboole_0)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t4_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', d1_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t28_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', commutativity_k3_xboole_0)).
fof(t38_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t38_zfmisc_1)).
fof(l31_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1)), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', l31_zfmisc_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', l13_xboole_0)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', l27_zfmisc_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', l32_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t12_xboole_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/aae56f91021d4f898379a05620cc7c4b.E_input', t10_xboole_1)).
fof(c_0_28, plain, ![X21, X22]:k3_tarski(k2_tarski(X21,X22))=k2_xboole_0(X21,X22), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_29, plain, ![X23, X24]:k1_enumset1(X23,X23,X24)=k2_tarski(X23,X24), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_30, negated_conjecture, ~(![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0))))))), inference(assume_negation,[status(cth)],[t43_zfmisc_1])).
fof(c_0_31, plain, ![X8, X9]:r1_tarski(k4_xboole_0(X8,X9),X8), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_32, plain, ![X10, X11]:k4_xboole_0(X10,X11)=k5_xboole_0(X10,k3_xboole_0(X10,X11)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_33, plain, ![X40, X41]:r1_tarski(X40,k2_xboole_0(X40,X41)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_34, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, plain, ![X42, X43, X44]:k2_enumset1(X42,X42,X43,X44)=k1_enumset1(X42,X43,X44), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_37, plain, ![X45, X46, X47, X48]:k3_enumset1(X45,X45,X46,X47,X48)=k2_enumset1(X45,X46,X47,X48), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_38, plain, ![X49, X50, X51, X52, X53]:k4_enumset1(X49,X49,X50,X51,X52,X53)=k3_enumset1(X49,X50,X51,X52,X53), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_39, plain, ![X54, X55, X56, X57, X58, X59]:k5_enumset1(X54,X54,X55,X56,X57,X58,X59)=k4_enumset1(X54,X55,X56,X57,X58,X59), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_40, plain, ![X60, X61, X62, X63, X64, X65, X66]:k6_enumset1(X60,X60,X61,X62,X63,X64,X65,X66)=k5_enumset1(X60,X61,X62,X63,X64,X65,X66), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_41, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)&((esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0))&((esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0))&(esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
fof(c_0_42, plain, ![X67]:k2_tarski(X67,X67)=k1_tarski(X67), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_43, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_44, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_45, plain, ![X12]:k3_xboole_0(X12,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_46, plain, ![X13]:k5_xboole_0(X13,k1_xboole_0)=X13, inference(variable_rename,[status(thm)],[t5_boole])).
fof(c_0_47, plain, ![X14]:(~r1_tarski(X14,k1_xboole_0)|X14=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
fof(c_0_48, plain, ![X15, X16, X17]:(~r1_tarski(X15,X16)|r1_tarski(k3_xboole_0(X15,X17),X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t108_xboole_1])])).
fof(c_0_49, plain, ![X18, X19, X20]:(((~r2_hidden(X18,X19)|~r2_hidden(X18,X20)|~r2_hidden(X18,k5_xboole_0(X19,X20)))&(r2_hidden(X18,X19)|r2_hidden(X18,X20)|~r2_hidden(X18,k5_xboole_0(X19,X20))))&((~r2_hidden(X18,X19)|r2_hidden(X18,X20)|r2_hidden(X18,k5_xboole_0(X19,X20)))&(~r2_hidden(X18,X20)|r2_hidden(X18,X19)|r2_hidden(X18,k5_xboole_0(X19,X20))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_50, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_51, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_52, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_53, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_54, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_55, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_56, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_57, negated_conjecture, (k1_tarski(esk1_0)=k2_xboole_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_58, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_59, plain, ![X28, X29, X31, X32, X33]:((r1_xboole_0(X28,X29)|r2_hidden(esk4_2(X28,X29),k3_xboole_0(X28,X29)))&(~r2_hidden(X33,k3_xboole_0(X31,X32))|~r1_xboole_0(X31,X32))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_60, plain, ![X34, X35, X36]:((~v1_xboole_0(X34)|~r2_hidden(X35,X34))&(r2_hidden(esk5_1(X36),X36)|v1_xboole_0(X36))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_61, plain, ![X38, X39]:(~r1_tarski(X38,X39)|k3_xboole_0(X38,X39)=X38), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_62, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_63, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_64, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_65, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_66, plain, (r1_tarski(k3_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_67, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X3,X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_68, plain, ![X68, X69]:k3_xboole_0(X68,X69)=k3_xboole_0(X69,X68), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_69, plain, (r1_tarski(X1,k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_70, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_58]), c_0_35]), c_0_51]), c_0_52]), c_0_52]), c_0_53]), c_0_53]), c_0_54]), c_0_54]), c_0_55]), c_0_55]), c_0_56]), c_0_56])).
cnf(c_0_71, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_72, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_73, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_74, plain, (r1_tarski(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])).
cnf(c_0_75, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_76, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_77, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_67, c_0_64])).
cnf(c_0_78, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_79, negated_conjecture, (r1_tarski(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
fof(c_0_80, plain, ![X79, X80, X81]:(((r2_hidden(X79,X81)|~r1_tarski(k2_tarski(X79,X80),X81))&(r2_hidden(X80,X81)|~r1_tarski(k2_tarski(X79,X80),X81)))&(~r2_hidden(X79,X81)|~r2_hidden(X80,X81)|r1_tarski(k2_tarski(X79,X80),X81))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_zfmisc_1])])])).
fof(c_0_81, plain, ![X70, X71]:(~r2_hidden(X70,X71)|k3_xboole_0(X71,k1_tarski(X70))=k1_tarski(X70)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l31_zfmisc_1])])).
cnf(c_0_82, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_83, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_84, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk4_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_85, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_75, c_0_66])).
cnf(c_0_86, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_64]), c_0_77])).
cnf(c_0_87, plain, (r1_tarski(k3_xboole_0(X1,X2),X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_66, c_0_78])).
cnf(c_0_88, negated_conjecture, (k3_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))=esk2_0), inference(spm,[status(thm)],[c_0_73, c_0_79])).
cnf(c_0_89, plain, (r1_tarski(k2_tarski(X1,X3),X2)|~r2_hidden(X1,X2)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_90, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_80])).
fof(c_0_91, plain, ![X76]:(~v1_xboole_0(X76)|X76=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_92, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_71, c_0_78])).
fof(c_0_93, plain, ![X72, X73]:(r2_hidden(X72,X73)|r1_xboole_0(k1_tarski(X72),X73)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
cnf(c_0_94, plain, (k3_xboole_0(X2,k1_tarski(X1))=k1_tarski(X1)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_95, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_96, plain, (v1_xboole_0(X1)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_97, plain, (r1_xboole_0(k3_xboole_0(X1,X2),X3)|~r1_tarski(X1,k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])).
fof(c_0_98, plain, ![X74, X75]:((k4_xboole_0(X74,X75)!=k1_xboole_0|r1_tarski(X74,X75))&(~r1_tarski(X74,X75)|k4_xboole_0(X74,X75)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_99, negated_conjecture, (r1_tarski(esk2_0,X1)|~r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),X1)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_100, plain, (r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X3),X2)|~r2_hidden(X3,X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_35]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_101, plain, (r2_hidden(X1,X2)|~r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X3),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_35]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_102, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_103, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_92, c_0_72])).
cnf(c_0_104, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_105, plain, (k3_xboole_0(X2,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))=k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_58]), c_0_58]), c_0_35]), c_0_35]), c_0_52]), c_0_52]), c_0_53]), c_0_53]), c_0_54]), c_0_54]), c_0_55]), c_0_55]), c_0_56]), c_0_56])).
cnf(c_0_106, plain, (v1_xboole_0(k5_xboole_0(X1,X2))|r2_hidden(esk5_1(k5_xboole_0(X1,X2)),X2)|r2_hidden(esk5_1(k5_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_95, c_0_72])).
fof(c_0_107, plain, ![X77, X78]:(~r1_tarski(X77,X78)|k2_xboole_0(X77,X78)=X78), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_108, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_109, plain, (k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_62]), c_0_78])).
cnf(c_0_110, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_111, negated_conjecture, (r1_tarski(esk2_0,X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_112, plain, (r2_hidden(X1,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(spm,[status(thm)],[c_0_101, c_0_74])).
cnf(c_0_113, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_114, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_115, plain, (r2_hidden(X1,X2)|r1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_58]), c_0_35]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_116, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk2_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_105, c_0_88])).
cnf(c_0_117, plain, (v1_xboole_0(k5_xboole_0(k1_xboole_0,X1))|r2_hidden(esk5_1(k5_xboole_0(k1_xboole_0,X1)),X1)), inference(spm,[status(thm)],[c_0_86, c_0_106])).
cnf(c_0_118, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_92, c_0_84])).
cnf(c_0_119, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_120, plain, (v1_xboole_0(k5_xboole_0(X1,k3_xboole_0(X1,X2)))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_121, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_110, c_0_44])).
cnf(c_0_122, negated_conjecture, (r1_tarski(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1))), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_123, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_113, c_0_44])).
cnf(c_0_124, negated_conjecture, (esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_125, plain, (k3_xboole_0(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))=k1_xboole_0|r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_126, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_102, c_0_82])).
cnf(c_0_127, negated_conjecture, (r1_xboole_0(esk2_0,X1)|r2_hidden(esk1_0,X1)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_128, plain, (v1_xboole_0(k5_xboole_0(k1_xboole_0,k3_xboole_0(X1,X2)))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_71, c_0_117])).
cnf(c_0_129, negated_conjecture, (~r1_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_71, c_0_88])).
cnf(c_0_130, plain, (r1_xboole_0(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))|r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_118, c_0_115])).
cnf(c_0_131, negated_conjecture, (esk2_0!=k1_xboole_0|esk3_0!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_132, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_133, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_102, c_0_120])).
cnf(c_0_134, plain, (r1_tarski(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))|k5_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_121, c_0_109])).
cnf(c_0_135, negated_conjecture, (k3_xboole_0(esk2_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1))=esk2_0), inference(spm,[status(thm)],[c_0_73, c_0_122])).
cnf(c_0_136, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_83]), c_0_74])])).
cnf(c_0_137, negated_conjecture, (esk2_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)|esk3_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_58]), c_0_58]), c_0_35]), c_0_35]), c_0_52]), c_0_52]), c_0_53]), c_0_53]), c_0_54]), c_0_54]), c_0_55]), c_0_55]), c_0_56]), c_0_56])).
cnf(c_0_138, negated_conjecture, (k3_xboole_0(X1,esk2_0)=esk2_0|~r2_hidden(esk1_0,esk2_0)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_105, c_0_116])).
cnf(c_0_139, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_88, c_0_125])).
cnf(c_0_140, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|r2_hidden(esk1_0,X1)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_141, negated_conjecture, (esk2_0!=k1_tarski(esk1_0)|esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_142, plain, (k5_xboole_0(k1_xboole_0,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_102, c_0_128])).
cnf(c_0_143, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_144, plain, (r1_xboole_0(X1,X1)|r2_hidden(esk4_2(X1,X1),X1)), inference(spm,[status(thm)],[c_0_84, c_0_83])).
cnf(c_0_145, negated_conjecture, (esk2_0!=k1_xboole_0|esk3_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_58]), c_0_35]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_146, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk3_0|~r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_70, c_0_132])).
cnf(c_0_147, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_121, c_0_133])).
cnf(c_0_148, negated_conjecture, (r1_tarski(esk2_0,k1_xboole_0)|esk2_0!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_136]), c_0_136]), c_0_64])).
cnf(c_0_149, negated_conjecture, (esk3_0!=esk2_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_137, c_0_116])).
cnf(c_0_150, negated_conjecture, (k3_xboole_0(X1,esk2_0)=esk2_0|esk2_0=k1_xboole_0|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_138, c_0_139])).
cnf(c_0_151, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|esk2_0=k1_xboole_0|r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_140, c_0_139])).
cnf(c_0_152, negated_conjecture, (esk3_0!=k1_xboole_0|esk2_0!=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_141, c_0_58]), c_0_35]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_153, plain, (k5_xboole_0(k1_xboole_0,X1)=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_142, c_0_83])).
cnf(c_0_154, negated_conjecture, (r1_xboole_0(esk2_0,esk2_0)|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_143, c_0_144])).
fof(c_0_155, plain, ![X82, X83, X84]:(~r1_tarski(X82,X83)|r1_tarski(X82,k2_xboole_0(X84,X83))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_156, negated_conjecture, (esk2_0!=k1_xboole_0|~r1_tarski(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_145, c_0_146])).
cnf(c_0_157, negated_conjecture, (r1_tarski(esk2_0,X1)|esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_147, c_0_148])).
cnf(c_0_158, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)|~r1_tarski(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_116]), c_0_149])).
cnf(c_0_159, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k3_xboole_0(X2,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_121, c_0_78])).
cnf(c_0_160, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|k3_xboole_0(X1,esk2_0)=esk2_0|esk2_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_150, c_0_151])).
cnf(c_0_161, negated_conjecture, (esk3_0!=k1_xboole_0|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_152, c_0_116])).
cnf(c_0_162, negated_conjecture, (k5_xboole_0(k1_xboole_0,esk2_0)=k1_xboole_0|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_153, c_0_154])).
cnf(c_0_163, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_164, negated_conjecture, (esk2_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_165, negated_conjecture, (~r1_tarski(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_139]), c_0_156])).
cnf(c_0_166, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|esk2_0=k1_xboole_0|r1_tarski(esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_136])])).
cnf(c_0_167, negated_conjecture, (k5_xboole_0(k1_xboole_0,esk2_0)=k1_xboole_0|esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_168, plain, (r1_tarski(X1,k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_163, c_0_51]), c_0_52]), c_0_53]), c_0_54]), c_0_55]), c_0_56])).
cnf(c_0_169, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(sr,[status(thm)],[c_0_139, c_0_164])).
cnf(c_0_170, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=k1_xboole_0|esk2_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_171, negated_conjecture, (esk3_0!=k1_xboole_0|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_167]), c_0_86])).
cnf(c_0_172, negated_conjecture, (r1_tarski(X1,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_168, c_0_70])).
cnf(c_0_173, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_169])])).
cnf(c_0_174, plain, (k5_xboole_0(X1,k3_xboole_0(X2,X1))=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_123, c_0_78])).
cnf(c_0_175, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=k1_xboole_0), inference(sr,[status(thm)],[c_0_170, c_0_164])).
cnf(c_0_176, negated_conjecture, (esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_171, c_0_169])).
cnf(c_0_177, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk3_0)), inference(rw,[status(thm)],[c_0_172, c_0_173])).
cnf(c_0_178, negated_conjecture, (~r1_tarski(esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175]), c_0_64]), c_0_176])).
cnf(c_0_179, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_74]), c_0_178]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 180
# Proof object clause steps            : 123
# Proof object formula steps           : 57
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 28
# Proof object generating inferences   : 68
# Proof object simplifying inferences  : 118
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 28
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 30
# Processed clauses                    : 2019
# ...of these trivial                  : 31
# ...subsumed                          : 1559
# ...remaining for further processing  : 429
# Other redundant clauses eliminated   : 27
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 28
# Backward-rewritten                   : 76
# Generated clauses                    : 9002
# ...of the previous two non-trivial   : 5787
# Contextual simplify-reflections      : 7
# Paramodulations                      : 8929
# Factorizations                       : 15
# Equation resolutions                 : 27
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 294
#    Positive orientable unit clauses  : 37
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 243
# Current number of unprocessed clauses: 3734
# ...number of literals in the above   : 11851
# Current number of archived formulas  : 0
# Current number of archived clauses   : 144
# Clause-clause subsumption calls (NU) : 12787
# Rec. Clause-clause subsumption calls : 10640
# Non-unit clause-clause subsumptions  : 1282
# Unit Clause-clause subsumption calls : 362
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 46
# BW rewrite match successes           : 15
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 89058

# -------------------------------------------------
# User time                : 0.059 s
# System time              : 0.003 s
# Total time               : 0.063 s
# Maximum resident set size: 28876 pages
