# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t76_mcart_1, conjecture, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(X2))=>![X6]:(m1_subset_1(X6,k1_zfmisc_1(X3))=>![X7]:(m1_subset_1(X7,k3_zfmisc_1(X1,X2,X3))=>(r2_hidden(X7,k3_zfmisc_1(X4,X5,X6))=>(r2_hidden(k5_mcart_1(X1,X2,X3,X7),X4)&(r2_hidden(k6_mcart_1(X1,X2,X3,X7),X5)&r2_hidden(k7_mcart_1(X1,X2,X3,X7),X6)))))))), file('data/MPTP/proofs/ab1330218bca46199f5e30887a7f4f88.E_input', t76_mcart_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/ab1330218bca46199f5e30887a7f4f88.E_input', t50_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/ab1330218bca46199f5e30887a7f4f88.E_input', d3_zfmisc_1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/ab1330218bca46199f5e30887a7f4f88.E_input', t10_mcart_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/ab1330218bca46199f5e30887a7f4f88.E_input', d1_xboole_0)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/ab1330218bca46199f5e30887a7f4f88.E_input', cc1_subset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/ab1330218bca46199f5e30887a7f4f88.E_input', fc1_xboole_0)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(X2))=>![X6]:(m1_subset_1(X6,k1_zfmisc_1(X3))=>![X7]:(m1_subset_1(X7,k3_zfmisc_1(X1,X2,X3))=>(r2_hidden(X7,k3_zfmisc_1(X4,X5,X6))=>(r2_hidden(k5_mcart_1(X1,X2,X3,X7),X4)&(r2_hidden(k6_mcart_1(X1,X2,X3,X7),X5)&r2_hidden(k7_mcart_1(X1,X2,X3,X7),X6))))))))), inference(assume_negation,[status(cth)],[t76_mcart_1])).
fof(c_0_8, plain, ![X18, X19, X20, X21]:((k5_mcart_1(X18,X19,X20,X21)=k1_mcart_1(k1_mcart_1(X21))|~m1_subset_1(X21,k3_zfmisc_1(X18,X19,X20))|X20=k1_xboole_0|X19=k1_xboole_0|X18=k1_xboole_0)&((k6_mcart_1(X18,X19,X20,X21)=k2_mcart_1(k1_mcart_1(X21))|~m1_subset_1(X21,k3_zfmisc_1(X18,X19,X20))|X20=k1_xboole_0|X19=k1_xboole_0|X18=k1_xboole_0)&(k7_mcart_1(X18,X19,X20,X21)=k2_mcart_1(X21)|~m1_subset_1(X21,k3_zfmisc_1(X18,X19,X20))|X20=k1_xboole_0|X19=k1_xboole_0|X18=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
fof(c_0_9, plain, ![X15, X16, X17]:k3_zfmisc_1(X15,X16,X17)=k2_zfmisc_1(k2_zfmisc_1(X15,X16),X17), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_10, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk5_0,k1_zfmisc_1(esk2_0))&(m1_subset_1(esk6_0,k1_zfmisc_1(esk3_0))&(m1_subset_1(esk7_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&(r2_hidden(esk7_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))&(~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk4_0)|(~r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk5_0)|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_11, plain, (k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, plain, ![X24, X25, X26]:((r2_hidden(k1_mcart_1(X24),X25)|~r2_hidden(X24,k2_zfmisc_1(X25,X26)))&(r2_hidden(k2_mcart_1(X24),X26)|~r2_hidden(X24,k2_zfmisc_1(X25,X26)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_14, negated_conjecture, (r2_hidden(esk7_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk7_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (r2_hidden(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(rw,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_19, negated_conjecture, (~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk4_0)|~r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk5_0)|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, plain, (k6_mcart_1(X1,X2,X3,X4)=k6_mcart_1(X5,X6,X7,X4)|X5=k1_xboole_0|X6=k1_xboole_0|X7=k1_xboole_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X5,X6),X7))), inference(spm,[status(thm)],[c_0_15, c_0_15])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_16, c_0_12])).
cnf(c_0_22, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_23, negated_conjecture, (r2_hidden(k1_mcart_1(esk7_0),k2_zfmisc_1(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, negated_conjecture, (esk3_0=k1_xboole_0|esk2_0=k1_xboole_0|esk1_0=k1_xboole_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk4_0)|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)|~r2_hidden(k6_mcart_1(X3,X2,X1,esk7_0),esk5_0)|~m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(X3,X2),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])])).
cnf(c_0_25, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk7_0)),esk5_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_27, plain, ![X27, X28, X29]:((~v1_xboole_0(X27)|~r2_hidden(X28,X27))&(r2_hidden(esk8_1(X29),X29)|v1_xboole_0(X29))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_28, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0|esk3_0=k1_xboole_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk4_0)|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)|~m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_15]), c_0_25])])).
cnf(c_0_29, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_26, c_0_12])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk7_0)),esk4_0)), inference(spm,[status(thm)],[c_0_17, c_0_23])).
cnf(c_0_31, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_32, plain, ![X22, X23]:(~v1_xboole_0(X22)|(~m1_subset_1(X23,k1_zfmisc_1(X22))|v1_xboole_0(X23))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_33, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, negated_conjecture, (r2_hidden(k2_mcart_1(esk7_0),esk6_0)), inference(spm,[status(thm)],[c_0_22, c_0_18])).
cnf(c_0_35, negated_conjecture, (esk3_0=k1_xboole_0|esk2_0=k1_xboole_0|esk1_0=k1_xboole_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)|~m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(X3,X2),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_21])])).
cnf(c_0_36, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_31, c_0_12])).
cnf(c_0_37, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_40, negated_conjecture, (~v1_xboole_0(esk6_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_41, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0|esk3_0=k1_xboole_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_34]), c_0_21])])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43, negated_conjecture, (~v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_33, c_0_25])).
cnf(c_0_44, negated_conjecture, (v1_xboole_0(esk5_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_39]), c_0_40])).
cnf(c_0_46, negated_conjecture, (esk3_0=k1_xboole_0|esk2_0=k1_xboole_0|esk1_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_41, c_0_21])).
cnf(c_0_47, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_48, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_33, c_0_30])).
cnf(c_0_49, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_42])).
cnf(c_0_50, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_52, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_53, negated_conjecture, (esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_47])])).
cnf(c_0_54, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53]), c_0_47])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 41
# Proof object formula steps           : 14
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 15
# Processed clauses                    : 60
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 56
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 6
# Generated clauses                    : 51
# ...of the previous two non-trivial   : 50
# Contextual simplify-reflections      : 0
# Paramodulations                      : 51
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 20
# ...number of literals in the above   : 140
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 52
# Rec. Clause-clause subsumption calls : 20
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 53
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1968

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28800 pages
