# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', t23_relat_1)).
fof(t129_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k1_tarski(X4)))<=>(r2_hidden(X1,X3)&X2=X4)), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', t129_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', t71_enumset1)).
fof(t35_zfmisc_1, axiom, ![X1, X2]:k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2)), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', t35_zfmisc_1)).
fof(t128_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),X4))<=>(X1=X3&r2_hidden(X2,X4))), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', t128_zfmisc_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', d5_relat_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/ab348c7807314a08a2dc1ed88a81a8be.E_input', d4_relat_1)).
fof(c_0_9, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t23_relat_1])).
fof(c_0_10, plain, ![X8, X9, X10, X11]:(((r2_hidden(X8,X10)|~r2_hidden(k4_tarski(X8,X9),k2_zfmisc_1(X10,k1_tarski(X11))))&(X9=X11|~r2_hidden(k4_tarski(X8,X9),k2_zfmisc_1(X10,k1_tarski(X11)))))&(~r2_hidden(X8,X10)|X9!=X11|r2_hidden(k4_tarski(X8,X9),k2_zfmisc_1(X10,k1_tarski(X11))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t129_zfmisc_1])])])).
fof(c_0_11, plain, ![X12]:k2_tarski(X12,X12)=k1_tarski(X12), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_12, plain, ![X13, X14]:k1_enumset1(X13,X13,X14)=k2_tarski(X13,X14), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_13, plain, ![X15, X16, X17]:k2_enumset1(X15,X15,X16,X17)=k1_enumset1(X15,X16,X17), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_14, plain, ![X22, X23]:k2_zfmisc_1(k1_tarski(X22),k1_tarski(X23))=k1_tarski(k4_tarski(X22,X23)), inference(variable_rename,[status(thm)],[t35_zfmisc_1])).
fof(c_0_15, negated_conjecture, (v1_relat_1(esk3_0)&(esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))&(k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
cnf(c_0_16, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, negated_conjecture, (esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_22, plain, ![X18, X19, X20, X21]:(((X18=X20|~r2_hidden(k4_tarski(X18,X19),k2_zfmisc_1(k1_tarski(X20),X21)))&(r2_hidden(X19,X21)|~r2_hidden(k4_tarski(X18,X19),k2_zfmisc_1(k1_tarski(X20),X21))))&(X18!=X20|~r2_hidden(X19,X21)|r2_hidden(k4_tarski(X18,X19),k2_zfmisc_1(k1_tarski(X20),X21)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t128_zfmisc_1])])])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k1_tarski(X4)))|~r2_hidden(X1,X2)|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_24, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k2_enumset1(X2,X2,X2,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_25, plain, (k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X2,X2))=k2_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_17]), c_0_17]), c_0_17]), c_0_18]), c_0_18]), c_0_18]), c_0_19]), c_0_19]), c_0_19])).
cnf(c_0_26, negated_conjecture, (esk3_0=k2_enumset1(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_27, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),X4))|X1!=X2|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(k1_tarski(X4),X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k2_enumset1(X4,X4,X4,X4)))|X3!=X4|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_30, plain, (X1=k4_tarski(X2,X3)|~r2_hidden(k4_tarski(X4,X1),k2_zfmisc_1(X5,k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),k2_enumset1(X3,X3,X3,X3))))), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, negated_conjecture, (k2_zfmisc_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=esk3_0), inference(rw,[status(thm)],[c_0_26, c_0_25])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),X4))|X1!=X2|~r2_hidden(X3,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_33, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(k2_enumset1(X4,X4,X4,X4),X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_34, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k2_enumset1(X2,X2,X2,X2)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_29])).
cnf(c_0_35, negated_conjecture, (X1=k4_tarski(esk1_0,esk2_0)|~r2_hidden(k4_tarski(X2,X1),k2_zfmisc_1(X3,esk3_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),X3))|~r2_hidden(X2,X3)), inference(er,[status(thm)],[c_0_32])).
fof(c_0_37, plain, ![X24, X25, X26, X28, X29, X30, X31, X33]:(((~r2_hidden(X26,X25)|r2_hidden(k4_tarski(esk4_3(X24,X25,X26),X26),X24)|X25!=k2_relat_1(X24))&(~r2_hidden(k4_tarski(X29,X28),X24)|r2_hidden(X28,X25)|X25!=k2_relat_1(X24)))&((~r2_hidden(esk5_2(X30,X31),X31)|~r2_hidden(k4_tarski(X33,esk5_2(X30,X31)),X30)|X31=k2_relat_1(X30))&(r2_hidden(esk5_2(X30,X31),X31)|r2_hidden(k4_tarski(esk6_2(X30,X31),esk5_2(X30,X31)),X30)|X31=k2_relat_1(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_38, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(k2_zfmisc_1(k2_enumset1(X4,X4,X4,X4),k2_enumset1(X5,X5,X5,X5)),X2))), inference(spm,[status(thm)],[c_0_33, c_0_25])).
cnf(c_0_39, plain, (r2_hidden(k4_tarski(X1,k4_tarski(X2,X3)),k2_zfmisc_1(X4,k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),k2_enumset1(X3,X3,X3,X3))))|~r2_hidden(X1,X4)), inference(spm,[status(thm)],[c_0_34, c_0_25])).
cnf(c_0_40, negated_conjecture, (X1=k4_tarski(esk1_0,esk2_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, plain, (r2_hidden(k4_tarski(esk4_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k1_tarski(X4)))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(esk3_0,X2))), inference(spm,[status(thm)],[c_0_38, c_0_31])).
cnf(c_0_44, negated_conjecture, (r2_hidden(k4_tarski(X1,k4_tarski(esk1_0,esk2_0)),k2_zfmisc_1(X2,esk3_0))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_31])).
fof(c_0_45, plain, ![X35, X36, X37, X39, X40, X41, X42, X44]:(((~r2_hidden(X37,X36)|r2_hidden(k4_tarski(X37,esk7_3(X35,X36,X37)),X35)|X36!=k1_relat_1(X35))&(~r2_hidden(k4_tarski(X39,X40),X35)|r2_hidden(X39,X36)|X36!=k1_relat_1(X35)))&((~r2_hidden(esk8_2(X41,X42),X42)|~r2_hidden(k4_tarski(esk8_2(X41,X42),X44),X41)|X42=k1_relat_1(X41))&(r2_hidden(esk8_2(X41,X42),X42)|r2_hidden(k4_tarski(esk8_2(X41,X42),esk9_2(X41,X42)),X41)|X42=k1_relat_1(X41)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_46, negated_conjecture, (k4_tarski(esk4_3(esk3_0,X1,X2),X2)=k4_tarski(esk1_0,esk2_0)|X1!=k2_relat_1(esk3_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k2_enumset1(X4,X4,X4,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_48, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_49, plain, (r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(k4_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_50, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|X1!=k2_relat_1(esk3_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_41, c_0_46])).
cnf(c_0_51, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k2_zfmisc_1(k2_enumset1(X4,X4,X4,X4),k2_enumset1(X5,X5,X5,X5))))), inference(spm,[status(thm)],[c_0_47, c_0_25])).
cnf(c_0_52, negated_conjecture, (k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_53, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))), inference(spm,[status(thm)],[c_0_33, c_0_44])).
cnf(c_0_54, negated_conjecture, (X1=k1_relat_1(esk3_0)|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(esk8_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_55, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(er,[status(thm)],[c_0_50])).
cnf(c_0_56, negated_conjecture, (X1=esk2_0|~r2_hidden(k4_tarski(X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_31])).
cnf(c_0_57, plain, (r2_hidden(esk5_2(X1,X2),X2)|r2_hidden(k4_tarski(esk6_2(X1,X2),esk5_2(X1,X2)),X1)|X2=k2_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_58, negated_conjecture, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,esk3_0))), inference(spm,[status(thm)],[c_0_51, c_0_31])).
cnf(c_0_59, negated_conjecture, (k1_relat_1(esk3_0)!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|k2_relat_1(esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_17]), c_0_17]), c_0_18]), c_0_18]), c_0_19]), c_0_19])).
cnf(c_0_60, negated_conjecture, (k2_enumset1(X1,X1,X1,X1)=k1_relat_1(esk3_0)|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, negated_conjecture, (k2_relat_1(esk3_0)=k1_relat_1(esk3_0)|r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_55, c_0_54])).
cnf(c_0_62, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,X1),esk3_0)|~r2_hidden(X1,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_36, c_0_31])).
cnf(c_0_63, negated_conjecture, (esk5_2(esk3_0,X1)=esk2_0|X1=k2_relat_1(esk3_0)|r2_hidden(esk5_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_64, negated_conjecture, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X1))|~r2_hidden(X2,esk3_0)), inference(spm,[status(thm)],[c_0_58, c_0_36])).
cnf(c_0_65, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_60])).
cnf(c_0_66, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),X4))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_67, plain, (X2=k2_relat_1(X1)|~r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(k4_tarski(X3,esk5_2(X1,X2)),X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_68, negated_conjecture, (esk5_2(esk3_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=esk2_0|k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)=k2_relat_1(esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_56])).
cnf(c_0_69, negated_conjecture, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X1))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_17]), c_0_18]), c_0_19])).
cnf(c_0_71, negated_conjecture, (k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)=k2_relat_1(esk3_0)|~r2_hidden(k4_tarski(X1,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69])])).
cnf(c_0_72, negated_conjecture, (X1=esk1_0|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_70, c_0_31])).
cnf(c_0_73, negated_conjecture, (k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)=k2_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_71, c_0_65])).
cnf(c_0_74, negated_conjecture, (r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~r2_hidden(X1,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_34, c_0_31])).
cnf(c_0_75, negated_conjecture, (esk8_2(esk3_0,X1)=esk1_0|X1=k1_relat_1(esk3_0)|r2_hidden(esk8_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_72, c_0_49])).
cnf(c_0_76, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)!=k1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_73])])).
cnf(c_0_77, plain, (X2=k1_relat_1(X1)|~r2_hidden(esk8_2(X1,X2),X2)|~r2_hidden(k4_tarski(esk8_2(X1,X2),X3),X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_78, negated_conjecture, (esk8_2(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))=esk1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76]), c_0_72])).
cnf(c_0_79, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,X1),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_69])]), c_0_76])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_65, c_0_79]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 62
# Proof object formula steps           : 19
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 18
# Processed clauses                    : 1639
# ...of these trivial                  : 46
# ...subsumed                          : 475
# ...remaining for further processing  : 1118
# Other redundant clauses eliminated   : 746
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 55
# Backward-rewritten                   : 246
# Generated clauses                    : 34720
# ...of the previous two non-trivial   : 32772
# Contextual simplify-reflections      : 13
# Paramodulations                      : 33703
# Factorizations                       : 27
# Equation resolutions                 : 987
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 794
#    Positive orientable unit clauses  : 183
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 609
# Current number of unprocessed clauses: 30500
# ...number of literals in the above   : 113841
# Current number of archived formulas  : 0
# Current number of archived clauses   : 325
# Clause-clause subsumption calls (NU) : 56210
# Rec. Clause-clause subsumption calls : 38238
# Non-unit clause-clause subsumptions  : 542
# Unit Clause-clause subsumption calls : 14957
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2795
# BW rewrite match successes           : 15
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 618586

# -------------------------------------------------
# User time                : 0.399 s
# System time              : 0.010 s
# Total time               : 0.409 s
# Maximum resident set size: 28764 pages
