# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1))))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t41_ordinal1)).
fof(t39_ordinal1, axiom, ![X1]:?[X2]:(v3_ordinal1(X2)&(~(r2_hidden(X2,X1))&![X3]:(v3_ordinal1(X3)=>(~(r2_hidden(X3,X1))=>r1_ordinal1(X2,X3))))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t39_ordinal1)).
fof(s1_xboole_0__e3_53__ordinal1, axiom, ![X1]:?[X2]:![X3]:(r2_hidden(X3,X2)<=>(r2_hidden(X3,X1)&v3_ordinal1(X3))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', s1_xboole_0__e3_53__ordinal1)).
fof(t13_ordinal1, axiom, ![X1, X2]:(r2_hidden(X1,k1_ordinal1(X2))<=>(r2_hidden(X1,X2)|X1=X2)), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t13_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', d1_ordinal1)).
fof(t24_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1))))))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t24_ordinal1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', redefinition_r1_ordinal1)).
fof(t34_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2)))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t34_ordinal1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t7_ordinal1)).
fof(t23_ordinal1, axiom, ![X1, X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>v3_ordinal1(X1))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t23_ordinal1)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t10_ordinal1)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', d4_tarski)).
fof(t35_ordinal1, axiom, ![X1]:(![X2]:(r2_hidden(X2,X1)=>v3_ordinal1(X2))=>v3_ordinal1(k3_tarski(X1))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t35_ordinal1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', d10_xboole_0)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t29_ordinal1)).
fof(d6_ordinal1, axiom, ![X1]:(v4_ordinal1(X1)<=>X1=k3_tarski(X1)), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', d6_ordinal1)).
fof(t14_ordinal1, axiom, ![X1]:~(X1=k1_ordinal1(X1)), file('data/MPTP/proofs/abc875981164403e99b8efd430f6a6b1.E_input', t14_ordinal1)).
fof(c_0_17, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1)))))), inference(assume_negation,[status(cth)],[t41_ordinal1])).
fof(c_0_18, plain, ![X36, X38]:(v3_ordinal1(esk8_1(X36))&(~r2_hidden(esk8_1(X36),X36)&(~v3_ordinal1(X38)|(r2_hidden(X38,X36)|r1_ordinal1(esk8_1(X36),X38))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_ordinal1])])])])])).
fof(c_0_19, plain, ![X8, X10, X11]:(((r2_hidden(X10,X8)|~r2_hidden(X10,esk3_1(X8)))&(v3_ordinal1(X10)|~r2_hidden(X10,esk3_1(X8))))&(~r2_hidden(X11,X8)|~v3_ordinal1(X11)|r2_hidden(X11,esk3_1(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[s1_xboole_0__e3_53__ordinal1])])])])])])).
fof(c_0_20, plain, ![X42, X43]:((~r2_hidden(X42,k1_ordinal1(X43))|(r2_hidden(X42,X43)|X42=X43))&((~r2_hidden(X42,X43)|r2_hidden(X42,k1_ordinal1(X43)))&(X42!=X43|r2_hidden(X42,k1_ordinal1(X43))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_ordinal1])])])).
fof(c_0_21, plain, ![X34]:k1_ordinal1(X34)=k2_xboole_0(X34,k1_tarski(X34)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_22, plain, ![X39, X40]:(~v3_ordinal1(X39)|(~v3_ordinal1(X40)|(r2_hidden(X39,X40)|(X39=X40|r2_hidden(X40,X39))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])).
fof(c_0_23, negated_conjecture, ![X7]:(v3_ordinal1(esk1_0)&(((v3_ordinal1(esk2_0)|~v4_ordinal1(esk1_0))&((r2_hidden(esk2_0,esk1_0)|~v4_ordinal1(esk1_0))&(~r2_hidden(k1_ordinal1(esk2_0),esk1_0)|~v4_ordinal1(esk1_0))))&(v4_ordinal1(esk1_0)|(~v3_ordinal1(X7)|(~r2_hidden(X7,esk1_0)|r2_hidden(k1_ordinal1(X7),esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])])])).
fof(c_0_24, plain, ![X26, X27]:((~r1_ordinal1(X26,X27)|r1_tarski(X26,X27)|(~v3_ordinal1(X26)|~v3_ordinal1(X27)))&(~r1_tarski(X26,X27)|r1_ordinal1(X26,X27)|(~v3_ordinal1(X26)|~v3_ordinal1(X27)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_25, plain, (~r2_hidden(esk8_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (r2_hidden(X1,esk3_1(X2))|~r2_hidden(X1,X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (v3_ordinal1(esk8_1(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (r2_hidden(X1,k1_ordinal1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|X1=X2|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_32, plain, ![X44, X45]:((~r2_hidden(X44,k1_ordinal1(X45))|r1_ordinal1(X44,X45)|~v3_ordinal1(X45)|~v3_ordinal1(X44))&(~r1_ordinal1(X44,X45)|r2_hidden(X44,k1_ordinal1(X45))|~v3_ordinal1(X45)|~v3_ordinal1(X44))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_ordinal1])])])])).
fof(c_0_33, plain, ![X24, X25]:(~r2_hidden(X24,X25)|~r1_tarski(X25,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_34, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, plain, (r2_hidden(X1,X2)|r1_ordinal1(esk8_1(X2),X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_36, plain, (~r2_hidden(esk8_1(esk3_1(X1)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_37, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38, negated_conjecture, (X1=esk1_0|r2_hidden(X1,esk1_0)|r2_hidden(esk1_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, plain, (r1_ordinal1(X1,X2)|~r2_hidden(X1,k1_ordinal1(X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_40, plain, ![X32, X33]:(~v3_ordinal1(X33)|(~r2_hidden(X32,X33)|v3_ordinal1(X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_ordinal1])])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (r1_tarski(esk8_1(X1),X2)|r2_hidden(X2,X1)|~v3_ordinal1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_27])])).
cnf(c_0_43, plain, (~r2_hidden(esk8_1(esk3_1(k2_xboole_0(X1,k1_tarski(X1)))),X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, negated_conjecture, (esk8_1(X1)=esk1_0|r2_hidden(esk1_0,esk8_1(X1))|r2_hidden(esk8_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_38, c_0_27])).
fof(c_0_45, plain, ![X46]:r2_hidden(X46,k1_ordinal1(X46)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
cnf(c_0_46, plain, (r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))), inference(rw,[status(thm)],[c_0_39, c_0_29])).
cnf(c_0_47, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_48, plain, ![X12, X13, X14, X16, X17, X18, X19, X21]:((((r2_hidden(X14,esk4_3(X12,X13,X14))|~r2_hidden(X14,X13)|X13!=k3_tarski(X12))&(r2_hidden(esk4_3(X12,X13,X14),X12)|~r2_hidden(X14,X13)|X13!=k3_tarski(X12)))&(~r2_hidden(X16,X17)|~r2_hidden(X17,X12)|r2_hidden(X16,X13)|X13!=k3_tarski(X12)))&((~r2_hidden(esk5_2(X18,X19),X19)|(~r2_hidden(esk5_2(X18,X19),X21)|~r2_hidden(X21,X18))|X19=k3_tarski(X18))&((r2_hidden(esk5_2(X18,X19),esk6_2(X18,X19))|r2_hidden(esk5_2(X18,X19),X19)|X19=k3_tarski(X18))&(r2_hidden(esk6_2(X18,X19),X18)|r2_hidden(esk5_2(X18,X19),X19)|X19=k3_tarski(X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
cnf(c_0_49, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk8_1(X2))|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50, negated_conjecture, (esk8_1(esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0))))=esk1_0|r2_hidden(esk1_0,esk8_1(esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_51, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (r1_ordinal1(X1,X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_37]), c_0_47])).
cnf(c_0_53, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|X4!=k3_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, negated_conjecture, (esk8_1(esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0))))=esk1_0|r2_hidden(esk1_0,esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_31])])).
cnf(c_0_55, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))), inference(rw,[status(thm)],[c_0_51, c_0_29])).
fof(c_0_56, plain, ![X28]:((r2_hidden(esk7_1(X28),X28)|v3_ordinal1(k3_tarski(X28)))&(~v3_ordinal1(esk7_1(X28))|v3_ordinal1(k3_tarski(X28)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_ordinal1])])])])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_52]), c_0_47])).
cnf(c_0_58, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|X2!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk1_0,esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_54]), c_0_55])])).
cnf(c_0_61, plain, (v3_ordinal1(X1)|~r2_hidden(X1,esk3_1(X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_62, plain, (r2_hidden(esk7_1(X1),X1)|v3_ordinal1(k3_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (v3_ordinal1(k3_tarski(X1))|~v3_ordinal1(esk7_1(X1))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_41, c_0_57])).
cnf(c_0_65, plain, (r2_hidden(esk4_3(X1,k3_tarski(X1),X2),X1)|~r2_hidden(X2,k3_tarski(X1))), inference(er,[status(thm)],[c_0_58])).
cnf(c_0_66, plain, (r2_hidden(X1,esk4_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k3_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_67, negated_conjecture, (r2_hidden(X1,k3_tarski(esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_68, plain, (v3_ordinal1(k3_tarski(esk3_1(X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63])).
cnf(c_0_69, plain, (~r2_hidden(X1,esk4_3(X1,k3_tarski(X1),X2))|~r2_hidden(X2,k3_tarski(X1))|~v3_ordinal1(esk4_3(X1,k3_tarski(X1),X2))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, plain, (r2_hidden(X1,esk4_3(X2,k3_tarski(X2),X1))|~r2_hidden(X1,k3_tarski(X2))), inference(er,[status(thm)],[c_0_66])).
cnf(c_0_71, negated_conjecture, (v4_ordinal1(esk1_0)|r2_hidden(k1_ordinal1(X1),esk1_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_72, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_67]), c_0_68])])).
cnf(c_0_73, plain, (~r2_hidden(X1,k3_tarski(X1))|~v3_ordinal1(esk4_3(X1,k3_tarski(X1),X1))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
fof(c_0_74, plain, ![X30, X31]:(((r1_tarski(X30,X31)|X30!=X31)&(r1_tarski(X31,X30)|X30!=X31))&(~r1_tarski(X30,X31)|~r1_tarski(X31,X30)|X30=X31)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_75, negated_conjecture, (v4_ordinal1(esk1_0)|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[c_0_71, c_0_29])).
cnf(c_0_76, negated_conjecture, (v3_ordinal1(k3_tarski(X1))|~r2_hidden(esk7_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_72])).
cnf(c_0_77, negated_conjecture, (~r2_hidden(esk4_3(X1,k3_tarski(X1),X1),esk1_0)|~r2_hidden(X1,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_73, c_0_72])).
cnf(c_0_78, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_74])).
fof(c_0_79, plain, ![X35]:(~v3_ordinal1(X35)|v3_ordinal1(k1_ordinal1(X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
fof(c_0_80, plain, ![X41]:((~v4_ordinal1(X41)|X41=k3_tarski(X41))&(X41!=k3_tarski(X41)|v4_ordinal1(X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_ordinal1])])).
cnf(c_0_81, negated_conjecture, (r2_hidden(X1,k3_tarski(esk1_0))|v4_ordinal1(esk1_0)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X2,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_75]), c_0_72])).
cnf(c_0_82, negated_conjecture, (v3_ordinal1(k3_tarski(esk1_0))), inference(spm,[status(thm)],[c_0_76, c_0_62])).
cnf(c_0_83, negated_conjecture, (~r2_hidden(esk1_0,k3_tarski(esk1_0))), inference(spm,[status(thm)],[c_0_77, c_0_65])).
cnf(c_0_84, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_78])).
cnf(c_0_85, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_86, negated_conjecture, (v3_ordinal1(esk2_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_87, plain, (v4_ordinal1(X1)|X1!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_88, negated_conjecture, (r2_hidden(X1,k3_tarski(esk1_0))|v4_ordinal1(esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_81, c_0_55])).
cnf(c_0_89, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(k3_tarski(esk1_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_82]), c_0_83])).
cnf(c_0_90, plain, (~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_41, c_0_84])).
cnf(c_0_91, plain, (X1=k3_tarski(X1)|~v4_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_92, negated_conjecture, (~r2_hidden(k1_ordinal1(esk2_0),esk1_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_93, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_85, c_0_29])).
cnf(c_0_94, negated_conjecture, (v3_ordinal1(esk2_0)|k3_tarski(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_95, negated_conjecture, (k3_tarski(esk1_0)=esk1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_91])).
cnf(c_0_96, negated_conjecture, (~v4_ordinal1(esk1_0)|~r2_hidden(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),esk1_0)), inference(rw,[status(thm)],[c_0_92, c_0_29])).
cnf(c_0_97, negated_conjecture, (k2_xboole_0(X1,k1_tarski(X1))=esk1_0|r2_hidden(esk1_0,k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_93])).
cnf(c_0_98, negated_conjecture, (v3_ordinal1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_99, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_100, plain, ![X23]:X23!=k1_ordinal1(X23), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_ordinal1])])).
cnf(c_0_101, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|r2_hidden(esk1_0,k2_xboole_0(esk2_0,k1_tarski(esk2_0)))|~v4_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98])])).
cnf(c_0_102, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|k3_tarski(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_99, c_0_87])).
cnf(c_0_103, plain, (X1!=k1_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_104, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|r1_ordinal1(esk1_0,esk2_0)|~v4_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_101]), c_0_98]), c_0_31])])).
cnf(c_0_105, plain, (r2_hidden(X1,k1_ordinal1(X2))|~r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_106, plain, (r2_hidden(X1,X2)|X1=X2|~r2_hidden(X1,k1_ordinal1(X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_107, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_95])])).
cnf(c_0_108, plain, (X1!=k2_xboole_0(X1,k1_tarski(X1))), inference(rw,[status(thm)],[c_0_103, c_0_29])).
cnf(c_0_109, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|r1_tarski(esk1_0,esk2_0)|~v4_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_104]), c_0_98]), c_0_31])])).
cnf(c_0_110, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r1_ordinal1(X1,X2)), inference(rw,[status(thm)],[c_0_105, c_0_29])).
cnf(c_0_111, plain, (X1=X2|r2_hidden(X1,X2)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))), inference(rw,[status(thm)],[c_0_106, c_0_29])).
cnf(c_0_112, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_107]), c_0_98])])).
cnf(c_0_113, negated_conjecture, (r1_tarski(esk1_0,esk2_0)|esk2_0!=esk1_0|~v4_ordinal1(esk1_0)), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_114, plain, (~r1_ordinal1(esk8_1(esk3_1(k2_xboole_0(X1,k1_tarski(X1)))),X1)|~v3_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_110]), c_0_27])])).
cnf(c_0_115, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|esk2_0=esk1_0|~v4_ordinal1(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_101]), c_0_112])).
cnf(c_0_116, negated_conjecture, (esk8_1(esk3_1(esk1_0))=esk1_0|r2_hidden(esk1_0,esk8_1(esk3_1(esk1_0)))), inference(spm,[status(thm)],[c_0_36, c_0_44])).
cnf(c_0_117, negated_conjecture, (r1_tarski(esk1_0,esk2_0)|esk2_0!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_87]), c_0_95])])).
cnf(c_0_118, negated_conjecture, (esk2_0=esk1_0|~r1_ordinal1(esk8_1(esk3_1(esk1_0)),esk2_0)|~v4_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_98])])).
cnf(c_0_119, negated_conjecture, (esk8_1(esk3_1(esk1_0))=esk1_0|r2_hidden(esk1_0,esk3_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_116]), c_0_31])])).
cnf(c_0_120, negated_conjecture, (esk2_0!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_117]), c_0_107])])).
cnf(c_0_121, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk3_1(X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_122, negated_conjecture, (r2_hidden(esk1_0,esk3_1(esk1_0))|~r1_ordinal1(esk1_0,esk2_0)|~v4_ordinal1(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_120])).
cnf(c_0_123, negated_conjecture, (~r1_ordinal1(esk1_0,esk2_0)|~v4_ordinal1(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_90])).
cnf(c_0_124, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|~v4_ordinal1(esk1_0)), inference(spm,[status(thm)],[c_0_123, c_0_104])).
cnf(c_0_125, plain, (~r2_hidden(esk4_3(X1,k3_tarski(X1),X2),X2)|~r2_hidden(X2,k3_tarski(X1))|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_64, c_0_70])).
cnf(c_0_126, negated_conjecture, (X1=esk2_0|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)|~v4_ordinal1(esk1_0)), inference(spm,[status(thm)],[c_0_111, c_0_124])).
cnf(c_0_127, negated_conjecture, (~r2_hidden(esk4_3(esk1_0,esk1_0,X1),X1)|~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_95]), c_0_72])).
cnf(c_0_128, negated_conjecture, (X1=esk2_0|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_87]), c_0_95])])).
cnf(c_0_129, negated_conjecture, (esk4_3(esk1_0,esk1_0,esk2_0)=esk2_0|~r2_hidden(esk4_3(esk1_0,esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_107])])).
cnf(c_0_130, negated_conjecture, (r2_hidden(esk4_3(esk1_0,esk1_0,X1),esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_65, c_0_95])).
cnf(c_0_131, negated_conjecture, (r2_hidden(X1,esk4_3(esk1_0,esk1_0,X1))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_70, c_0_95])).
cnf(c_0_132, negated_conjecture, (esk4_3(esk1_0,esk1_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130]), c_0_107])])).
cnf(c_0_133, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_107])]), c_0_90]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 134
# Proof object clause steps            : 99
# Proof object formula steps           : 35
# Proof object conjectures             : 52
# Proof object clause conjectures      : 49
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 53
# Proof object simplifying inferences  : 67
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 37
# Processed clauses                    : 1191
# ...of these trivial                  : 5
# ...subsumed                          : 549
# ...remaining for further processing  : 637
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 79
# Generated clauses                    : 4721
# ...of the previous two non-trivial   : 4348
# Contextual simplify-reflections      : 42
# Paramodulations                      : 4711
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 478
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 134
#    Non-unit-clauses                  : 319
# Current number of unprocessed clauses: 3178
# ...number of literals in the above   : 8668
# Current number of archived formulas  : 0
# Current number of archived clauses   : 154
# Clause-clause subsumption calls (NU) : 27382
# Rec. Clause-clause subsumption calls : 11435
# Non-unit clause-clause subsumptions  : 278
# Unit Clause-clause subsumption calls : 7764
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 33
# BW rewrite match successes           : 19
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 81209

# -------------------------------------------------
# User time                : 0.053 s
# System time              : 0.007 s
# Total time               : 0.060 s
# Maximum resident set size: 28756 pages
