# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t60_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X1)&(k1_relat_1(X1)=k2_relat_1(X2)&(k2_relat_1(X1)=k1_relat_1(X2)&![X3, X4]:((r2_hidden(X3,k1_relat_1(X1))&r2_hidden(X4,k1_relat_1(X2)))=>(k1_funct_1(X1,X3)=X4<=>k1_funct_1(X2,X4)=X3)))))=>X2=k2_funct_1(X1)))), file('data/MPTP/proofs/ac6dfd1121af444fad065c6c1275bc1e.E_input', t60_funct_1)).
fof(t54_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(X2=k2_funct_1(X1)<=>(k1_relat_1(X2)=k2_relat_1(X1)&![X3, X4]:(((r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))=>(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4)))&((r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))=>(r2_hidden(X3,k2_relat_1(X1))&X4=k1_funct_1(X2,X3))))))))), file('data/MPTP/proofs/ac6dfd1121af444fad065c6c1275bc1e.E_input', t54_funct_1)).
fof(t12_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X1),k2_relat_1(X2)))), file('data/MPTP/proofs/ac6dfd1121af444fad065c6c1275bc1e.E_input', t12_funct_1)).
fof(c_0_3, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X1)&(k1_relat_1(X1)=k2_relat_1(X2)&(k2_relat_1(X1)=k1_relat_1(X2)&![X3, X4]:((r2_hidden(X3,k1_relat_1(X1))&r2_hidden(X4,k1_relat_1(X2)))=>(k1_funct_1(X1,X3)=X4<=>k1_funct_1(X2,X4)=X3)))))=>X2=k2_funct_1(X1))))), inference(assume_negation,[status(cth)],[t60_funct_1])).
fof(c_0_4, plain, ![X9, X10, X11, X12, X13, X14]:(((k1_relat_1(X10)=k2_relat_1(X9)|X10!=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))&(((r2_hidden(X12,k1_relat_1(X9))|(~r2_hidden(X11,k2_relat_1(X9))|X12!=k1_funct_1(X10,X11))|X10!=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))&(X11=k1_funct_1(X9,X12)|(~r2_hidden(X11,k2_relat_1(X9))|X12!=k1_funct_1(X10,X11))|X10!=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9))))&((r2_hidden(X13,k2_relat_1(X9))|(~r2_hidden(X14,k1_relat_1(X9))|X13!=k1_funct_1(X9,X14))|X10!=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))&(X14=k1_funct_1(X10,X13)|(~r2_hidden(X14,k1_relat_1(X9))|X13!=k1_funct_1(X9,X14))|X10!=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9))))))&(((((r2_hidden(esk6_2(X9,X10),k1_relat_1(X9))|r2_hidden(esk3_2(X9,X10),k2_relat_1(X9))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))&(esk5_2(X9,X10)=k1_funct_1(X9,esk6_2(X9,X10))|r2_hidden(esk3_2(X9,X10),k2_relat_1(X9))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9))))&(~r2_hidden(esk5_2(X9,X10),k2_relat_1(X9))|esk6_2(X9,X10)!=k1_funct_1(X10,esk5_2(X9,X10))|r2_hidden(esk3_2(X9,X10),k2_relat_1(X9))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9))))&(((r2_hidden(esk6_2(X9,X10),k1_relat_1(X9))|esk4_2(X9,X10)=k1_funct_1(X10,esk3_2(X9,X10))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))&(esk5_2(X9,X10)=k1_funct_1(X9,esk6_2(X9,X10))|esk4_2(X9,X10)=k1_funct_1(X10,esk3_2(X9,X10))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9))))&(~r2_hidden(esk5_2(X9,X10),k2_relat_1(X9))|esk6_2(X9,X10)!=k1_funct_1(X10,esk5_2(X9,X10))|esk4_2(X9,X10)=k1_funct_1(X10,esk3_2(X9,X10))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))))&(((r2_hidden(esk6_2(X9,X10),k1_relat_1(X9))|(~r2_hidden(esk4_2(X9,X10),k1_relat_1(X9))|esk3_2(X9,X10)!=k1_funct_1(X9,esk4_2(X9,X10)))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))&(esk5_2(X9,X10)=k1_funct_1(X9,esk6_2(X9,X10))|(~r2_hidden(esk4_2(X9,X10),k1_relat_1(X9))|esk3_2(X9,X10)!=k1_funct_1(X9,esk4_2(X9,X10)))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9))))&(~r2_hidden(esk5_2(X9,X10),k2_relat_1(X9))|esk6_2(X9,X10)!=k1_funct_1(X10,esk5_2(X9,X10))|(~r2_hidden(esk4_2(X9,X10),k1_relat_1(X9))|esk3_2(X9,X10)!=k1_funct_1(X9,esk4_2(X9,X10)))|k1_relat_1(X10)!=k2_relat_1(X9)|X10=k2_funct_1(X9)|(~v1_relat_1(X10)|~v1_funct_1(X10))|~v2_funct_1(X9)|(~v1_relat_1(X9)|~v1_funct_1(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_funct_1])])])])])])).
fof(c_0_5, negated_conjecture, ![X7, X8]:((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v2_funct_1(esk1_0)&(k1_relat_1(esk1_0)=k2_relat_1(esk2_0)&(k2_relat_1(esk1_0)=k1_relat_1(esk2_0)&((k1_funct_1(esk1_0,X7)!=X8|k1_funct_1(esk2_0,X8)=X7|(~r2_hidden(X7,k1_relat_1(esk1_0))|~r2_hidden(X8,k1_relat_1(esk2_0))))&(k1_funct_1(esk2_0,X8)!=X7|k1_funct_1(esk1_0,X7)=X8|(~r2_hidden(X7,k1_relat_1(esk1_0))|~r2_hidden(X8,k1_relat_1(esk2_0))))))))&esk2_0!=k2_funct_1(esk1_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])])).
cnf(c_0_6, plain, (esk5_2(X1,X2)=k1_funct_1(X1,esk6_2(X1,X2))|r2_hidden(esk3_2(X1,X2),k2_relat_1(X1))|X2=k2_funct_1(X1)|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (k2_relat_1(esk1_0)=k1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (v2_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, plain, (r2_hidden(esk6_2(X1,X2),k1_relat_1(X1))|r2_hidden(esk3_2(X1,X2),k2_relat_1(X1))|X2=k2_funct_1(X1)|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_12, plain, ![X19, X20]:(~v1_relat_1(X20)|~v1_funct_1(X20)|(~r2_hidden(X19,k1_relat_1(X20))|r2_hidden(k1_funct_1(X20,X19),k2_relat_1(X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_1])])).
cnf(c_0_13, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,X1))=esk5_2(esk1_0,X1)|X1=k2_funct_1(esk1_0)|r2_hidden(esk3_2(esk1_0,X1),k1_relat_1(esk2_0))|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_6, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_14, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16, negated_conjecture, (esk2_0!=k2_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_17, negated_conjecture, (X1=k2_funct_1(esk1_0)|r2_hidden(esk6_2(esk1_0,X1),k1_relat_1(esk1_0))|r2_hidden(esk3_2(esk1_0,X1),k1_relat_1(esk2_0))|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_18, negated_conjecture, (k1_funct_1(esk1_0,X2)=X1|k1_funct_1(esk2_0,X1)!=X2|~r2_hidden(X2,k1_relat_1(esk1_0))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_19, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_21, negated_conjecture, (k1_relat_1(esk1_0)=k2_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_22, plain, (esk5_2(X1,X2)=k1_funct_1(X1,esk6_2(X1,X2))|esk4_2(X1,X2)=k1_funct_1(X2,esk3_2(X1,X2))|X2=k2_funct_1(X1)|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_23, plain, (r2_hidden(esk6_2(X1,X2),k1_relat_1(X1))|esk4_2(X1,X2)=k1_funct_1(X2,esk3_2(X1,X2))|X2=k2_funct_1(X1)|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_24, negated_conjecture, (k1_funct_1(esk2_0,X2)=X1|k1_funct_1(esk1_0,X1)!=X2|~r2_hidden(X1,k1_relat_1(esk1_0))|~r2_hidden(X2,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))|r2_hidden(esk6_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_17]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_26, negated_conjecture, (k1_funct_1(esk1_0,k1_funct_1(esk2_0,X1))=X1|~r2_hidden(k1_funct_1(esk2_0,X1),k1_relat_1(esk1_0))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_27, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|r2_hidden(k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0)),k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_14]), c_0_15])])).
cnf(c_0_28, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,X1))=esk5_2(esk1_0,X1)|k1_funct_1(X1,esk3_2(esk1_0,X1))=esk4_2(esk1_0,X1)|X1=k2_funct_1(esk1_0)|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_29, negated_conjecture, (k1_funct_1(X1,esk3_2(esk1_0,X1))=esk4_2(esk1_0,X1)|X1=k2_funct_1(esk1_0)|r2_hidden(esk6_2(esk1_0,X1),k1_relat_1(esk1_0))|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_30, negated_conjecture, (k1_funct_1(esk2_0,k1_funct_1(esk1_0,X1))=X1|~r2_hidden(k1_funct_1(esk1_0,X1),k1_relat_1(esk2_0))|~r2_hidden(X1,k1_relat_1(esk1_0))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_31, negated_conjecture, (r2_hidden(k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0)),k1_relat_1(esk2_0))|r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_25]), c_0_7]), c_0_9]), c_0_10])])).
cnf(c_0_32, negated_conjecture, (k1_funct_1(esk1_0,k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0)))=esk3_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_20]), c_0_27])).
cnf(c_0_33, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_28]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_34, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|r2_hidden(esk6_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_29]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_35, negated_conjecture, (k1_funct_1(esk2_0,k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0)))=esk6_2(esk1_0,esk2_0)|r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_25]), c_0_31])).
cnf(c_0_36, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37, plain, (esk4_2(X1,X2)=k1_funct_1(X2,esk3_2(X1,X2))|X2=k2_funct_1(X1)|~r2_hidden(esk5_2(X1,X2),k2_relat_1(X1))|esk6_2(X1,X2)!=k1_funct_1(X2,esk5_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_38, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|r2_hidden(k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0)),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_34]), c_0_7]), c_0_9]), c_0_10])])).
cnf(c_0_39, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)|r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, negated_conjecture, (k1_funct_1(X1,esk3_2(esk1_0,X1))=esk4_2(esk1_0,X1)|X1=k2_funct_1(esk1_0)|k1_funct_1(X1,esk5_2(esk1_0,X1))!=esk6_2(esk1_0,X1)|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~r2_hidden(esk5_2(esk1_0,X1),k1_relat_1(esk2_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_41, negated_conjecture, (k1_funct_1(esk2_0,k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0)))=esk6_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_34]), c_0_38])).
cnf(c_0_42, plain, (r2_hidden(esk3_2(X1,X2),k2_relat_1(X1))|X2=k2_funct_1(X1)|~r2_hidden(esk5_2(X1,X2),k2_relat_1(X1))|esk6_2(X1,X2)!=k1_funct_1(X2,esk5_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_43, negated_conjecture, (k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)|r2_hidden(k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0)),k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_39]), c_0_21]), c_0_14]), c_0_15])])).
cnf(c_0_44, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=esk6_2(esk1_0,esk2_0)|~r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_40]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_45, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_41, c_0_36])).
cnf(c_0_46, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_38, c_0_36])).
cnf(c_0_47, negated_conjecture, (X1=k2_funct_1(esk1_0)|r2_hidden(esk3_2(esk1_0,X1),k1_relat_1(esk2_0))|k1_funct_1(X1,esk5_2(esk1_0,X1))!=esk6_2(esk1_0,X1)|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~r2_hidden(esk5_2(esk1_0,X1),k1_relat_1(esk2_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_48, negated_conjecture, (k1_funct_1(esk1_0,k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0)))=esk3_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_39]), c_0_43])).
cnf(c_0_49, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=esk6_2(esk1_0,esk2_0)|~r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_47]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_51, negated_conjecture, (k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_52, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)|r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))|r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_31, c_0_36])).
cnf(c_0_53, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)|r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_54, plain, (esk5_2(X1,X2)=k1_funct_1(X1,esk6_2(X1,X2))|X2=k2_funct_1(X1)|~r2_hidden(esk4_2(X1,X2),k1_relat_1(X1))|esk3_2(X1,X2)!=k1_funct_1(X1,esk4_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_55, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)|r2_hidden(k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0)),k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_53]), c_0_21]), c_0_14]), c_0_15])])).
cnf(c_0_56, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,X1))=esk5_2(esk1_0,X1)|X1=k2_funct_1(esk1_0)|k1_funct_1(esk1_0,esk4_2(esk1_0,X1))!=esk3_2(esk1_0,X1)|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~r2_hidden(esk4_2(esk1_0,X1),k1_relat_1(esk1_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_57, negated_conjecture, (k1_funct_1(esk1_0,k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0)))=esk3_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_53]), c_0_55])).
cnf(c_0_58, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))!=esk3_2(esk1_0,esk2_0)|~r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_56]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_59, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))=esk3_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_49])).
cnf(c_0_60, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(spm,[status(thm)],[c_0_27, c_0_33])).
cnf(c_0_61, plain, (r2_hidden(esk6_2(X1,X2),k1_relat_1(X1))|X2=k2_funct_1(X1)|~r2_hidden(esk4_2(X1,X2),k1_relat_1(X1))|esk3_2(X1,X2)!=k1_funct_1(X1,esk4_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_62, negated_conjecture, (k1_funct_1(esk1_0,esk6_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])]), c_0_60])).
cnf(c_0_63, negated_conjecture, (X1=k2_funct_1(esk1_0)|r2_hidden(esk6_2(esk1_0,X1),k1_relat_1(esk1_0))|k1_funct_1(esk1_0,esk4_2(esk1_0,X1))!=esk3_2(esk1_0,X1)|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~r2_hidden(esk4_2(esk1_0,X1),k1_relat_1(esk1_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_64, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_41, c_0_62])).
cnf(c_0_65, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(rw,[status(thm)],[c_0_38, c_0_62])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk6_2(esk1_0,esk2_0),k1_relat_1(esk1_0))|k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))!=esk3_2(esk1_0,esk2_0)|~r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_63]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_67, negated_conjecture, (k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)|r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(rw,[status(thm)],[c_0_35, c_0_62])).
cnf(c_0_68, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_64]), c_0_65])).
cnf(c_0_69, negated_conjecture, (r2_hidden(esk6_2(esk1_0,esk2_0),k1_relat_1(esk1_0))|~r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_59])])).
cnf(c_0_70, negated_conjecture, (k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)|r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_67]), c_0_21]), c_0_14]), c_0_15])]), c_0_68])).
cnf(c_0_71, negated_conjecture, (k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)|r2_hidden(esk6_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_72, plain, (X2=k2_funct_1(X1)|~r2_hidden(esk5_2(X1,X2),k2_relat_1(X1))|esk6_2(X1,X2)!=k1_funct_1(X2,esk5_2(X1,X2))|~r2_hidden(esk4_2(X1,X2),k1_relat_1(X1))|esk3_2(X1,X2)!=k1_funct_1(X1,esk4_2(X1,X2))|k1_relat_1(X2)!=k2_relat_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_73, negated_conjecture, (k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)|~r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_71]), c_0_62]), c_0_62])])).
cnf(c_0_74, negated_conjecture, (X1=k2_funct_1(esk1_0)|k1_funct_1(esk1_0,esk4_2(esk1_0,X1))!=esk3_2(esk1_0,X1)|k1_funct_1(X1,esk5_2(esk1_0,X1))!=esk6_2(esk1_0,X1)|k1_relat_1(esk2_0)!=k1_relat_1(X1)|~r2_hidden(esk4_2(esk1_0,X1),k1_relat_1(esk1_0))|~r2_hidden(esk5_2(esk1_0,X1),k1_relat_1(esk2_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_7]), c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_75, negated_conjecture, (k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))=esk6_2(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_71]), c_0_62]), c_0_7]), c_0_9]), c_0_10])]), c_0_73])).
cnf(c_0_76, negated_conjecture, (r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))|r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(rw,[status(thm)],[c_0_31, c_0_62])).
cnf(c_0_77, negated_conjecture, (k1_funct_1(esk1_0,esk4_2(esk1_0,esk2_0))!=esk3_2(esk1_0,esk2_0)|k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=esk6_2(esk1_0,esk2_0)|~r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))|~r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_74]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_78, negated_conjecture, (r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_75])]), c_0_76])).
cnf(c_0_79, negated_conjecture, (k1_funct_1(esk2_0,esk5_2(esk1_0,esk2_0))!=esk6_2(esk1_0,esk2_0)|~r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))|~r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_59])])).
cnf(c_0_80, negated_conjecture, (r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_78]), c_0_68]), c_0_21]), c_0_14]), c_0_15])])).
cnf(c_0_81, negated_conjecture, (~r2_hidden(esk4_2(esk1_0,esk2_0),k1_relat_1(esk1_0))|~r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_75])])).
cnf(c_0_82, negated_conjecture, (r2_hidden(esk6_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_80])])).
cnf(c_0_83, negated_conjecture, (~r2_hidden(esk5_2(esk1_0,esk2_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_80])])).
cnf(c_0_84, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_82]), c_0_62]), c_0_7]), c_0_9]), c_0_10])]), c_0_83]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 85
# Proof object clause steps            : 78
# Proof object formula steps           : 7
# Proof object conjectures             : 71
# Proof object clause conjectures      : 68
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 147
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 348
# ...of these trivial                  : 5
# ...subsumed                          : 165
# ...remaining for further processing  : 178
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 86
# Generated clauses                    : 482
# ...of the previous two non-trivial   : 524
# Contextual simplify-reflections      : 27
# Paramodulations                      : 466
# Factorizations                       : 0
# Equation resolutions                 : 20
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 34
# Current number of unprocessed clauses: 217
# ...number of literals in the above   : 592
# Current number of archived formulas  : 0
# Current number of archived clauses   : 121
# Clause-clause subsumption calls (NU) : 5206
# Rec. Clause-clause subsumption calls : 3954
# Non-unit clause-clause subsumptions  : 203
# Unit Clause-clause subsumption calls : 196
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17092

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.000 s
# Total time               : 0.019 s
# Maximum resident set size: 28840 pages
