# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t149_zfmisc_1, conjecture, ![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', t149_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', t71_enumset1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', d1_tarski)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', commutativity_k3_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', d3_tarski)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', t3_xboole_0)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', t4_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', d4_xboole_0)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', symmetry_r1_xboole_0)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/ad492ca34e8744d4a978b38dcb27960f.E_input', t70_xboole_1)).
fof(c_0_12, negated_conjecture, ~(![X1, X2, X3, X4]:((r1_tarski(k3_xboole_0(X1,X2),k1_tarski(X4))&(r2_hidden(X4,X3)&r1_xboole_0(X3,X2)))=>r1_xboole_0(k2_xboole_0(X1,X3),X2))), inference(assume_negation,[status(cth)],[t149_zfmisc_1])).
fof(c_0_13, negated_conjecture, ((r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))&(r2_hidden(esk4_0,esk3_0)&r1_xboole_0(esk3_0,esk2_0)))&~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
fof(c_0_14, plain, ![X21]:k2_tarski(X21,X21)=k1_tarski(X21), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_15, plain, ![X48, X49]:k1_enumset1(X48,X48,X49)=k2_tarski(X48,X49), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_16, plain, ![X9, X10, X11]:k2_enumset1(X9,X9,X10,X11)=k1_enumset1(X9,X10,X11), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_17, plain, ![X14, X15, X16, X17, X18, X19]:(((~r2_hidden(X16,X15)|X16=X14|X15!=k1_tarski(X14))&(X17!=X14|r2_hidden(X17,X15)|X15!=k1_tarski(X14)))&((~r2_hidden(esk5_2(X18,X19),X19)|esk5_2(X18,X19)!=X18|X19=k1_tarski(X18))&(r2_hidden(esk5_2(X18,X19),X19)|esk5_2(X18,X19)=X18|X19=k1_tarski(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_18, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k1_tarski(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X12, X13]:k3_xboole_0(X12,X13)=k3_xboole_0(X13,X12), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_23, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_24, plain, ![X50, X51, X52, X53, X54]:((~r1_tarski(X50,X51)|(~r2_hidden(X52,X50)|r2_hidden(X52,X51)))&((r2_hidden(esk9_2(X53,X54),X53)|r1_tarski(X53,X54))&(~r2_hidden(esk9_2(X53,X54),X54)|r1_tarski(X53,X54)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_25, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,esk2_0),k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])).
cnf(c_0_26, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, ![X30, X31, X33, X34, X35]:(((r2_hidden(esk7_2(X30,X31),X30)|r1_xboole_0(X30,X31))&(r2_hidden(esk7_2(X30,X31),X31)|r1_xboole_0(X30,X31)))&(~r2_hidden(X35,X33)|~r2_hidden(X35,X34)|~r1_xboole_0(X33,X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_28, plain, (X1=X3|X2!=k2_enumset1(X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_19]), c_0_20]), c_0_21])).
cnf(c_0_29, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (r1_tarski(k3_xboole_0(esk2_0,esk1_0),k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0))), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, negated_conjecture, (r1_xboole_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_33, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (r2_hidden(X1,k2_enumset1(esk4_0,esk4_0,esk4_0,esk4_0))|~r2_hidden(X1,k3_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_35, plain, ![X22, X23, X25, X26, X27]:((r1_xboole_0(X22,X23)|r2_hidden(esk6_2(X22,X23),k3_xboole_0(X22,X23)))&(~r2_hidden(X27,k3_xboole_0(X25,X26))|~r1_xboole_0(X25,X26))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_36, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, negated_conjecture, (X1=esk4_0|~r2_hidden(X1,k3_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, plain, (r2_hidden(esk7_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_40, plain, ![X39, X40, X41, X42, X43, X44, X45, X46]:((((r2_hidden(X42,X39)|~r2_hidden(X42,X41)|X41!=k3_xboole_0(X39,X40))&(r2_hidden(X42,X40)|~r2_hidden(X42,X41)|X41!=k3_xboole_0(X39,X40)))&(~r2_hidden(X43,X39)|~r2_hidden(X43,X40)|r2_hidden(X43,X41)|X41!=k3_xboole_0(X39,X40)))&((~r2_hidden(esk8_3(X44,X45,X46),X46)|(~r2_hidden(esk8_3(X44,X45,X46),X44)|~r2_hidden(esk8_3(X44,X45,X46),X45))|X46=k3_xboole_0(X44,X45))&((r2_hidden(esk8_3(X44,X45,X46),X44)|r2_hidden(esk8_3(X44,X45,X46),X46)|X46=k3_xboole_0(X44,X45))&(r2_hidden(esk8_3(X44,X45,X46),X45)|r2_hidden(esk8_3(X44,X45,X46),X46)|X46=k3_xboole_0(X44,X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_41, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk6_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, negated_conjecture, (r1_xboole_0(esk2_0,X1)|~r2_hidden(esk7_2(esk2_0,X1),esk3_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43, negated_conjecture, (esk7_2(X1,k3_xboole_0(esk2_0,esk1_0))=esk4_0|r1_xboole_0(X1,k3_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_45, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_46, plain, ![X28, X29]:(~r1_xboole_0(X28,X29)|r1_xboole_0(X29,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_47, plain, ![X36, X37, X38]:((r1_xboole_0(X36,k2_xboole_0(X37,X38))|(~r1_xboole_0(X36,X37)|~r1_xboole_0(X36,X38)))&((r1_xboole_0(X36,X37)|~r1_xboole_0(X36,k2_xboole_0(X37,X38)))&(r1_xboole_0(X36,X38)|~r1_xboole_0(X36,k2_xboole_0(X37,X38))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
cnf(c_0_48, negated_conjecture, (esk6_2(esk2_0,esk1_0)=esk4_0|r1_xboole_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_38, c_0_41])).
cnf(c_0_49, negated_conjecture, (r1_xboole_0(esk2_0,k3_xboole_0(esk2_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (r1_xboole_0(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53, negated_conjecture, (r1_xboole_0(esk2_0,esk1_0)|r2_hidden(esk4_0,k3_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_48])).
cnf(c_0_54, negated_conjecture, (~r2_hidden(X1,k3_xboole_0(esk2_0,esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_49]), c_0_50])).
cnf(c_0_55, negated_conjecture, (~r1_xboole_0(k2_xboole_0(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_56, plain, (r1_xboole_0(k2_xboole_0(X1,X2),X3)|~r1_xboole_0(X3,X2)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_57, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_32])).
cnf(c_0_58, negated_conjecture, (r1_xboole_0(esk2_0,esk1_0)), inference(sr,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_59, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57]), c_0_58])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 35
# Proof object formula steps           : 25
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 27
# Processed clauses                    : 193
# ...of these trivial                  : 11
# ...subsumed                          : 51
# ...remaining for further processing  : 131
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 7
# Generated clauses                    : 525
# ...of the previous two non-trivial   : 473
# Contextual simplify-reflections      : 1
# Paramodulations                      : 513
# Factorizations                       : 4
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 82
#    Positive orientable unit clauses  : 23
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 56
# Current number of unprocessed clauses: 316
# ...number of literals in the above   : 928
# Current number of archived formulas  : 0
# Current number of archived clauses   : 47
# Clause-clause subsumption calls (NU) : 583
# Rec. Clause-clause subsumption calls : 518
# Non-unit clause-clause subsumptions  : 49
# Unit Clause-clause subsumption calls : 116
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 24
# BW rewrite match successes           : 15
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7315

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28680 pages
