# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t2_xboole_1)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t2_tarski)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t10_mcart_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', d1_xboole_0)).
fof(t69_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X6))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k5_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t69_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', d3_zfmisc_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', d4_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', d2_subset_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t71_enumset1)).
fof(t55_mcart_1, axiom, ![X1, X2, X3, X4]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(X4=k1_xboole_0))))<=>~(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0)), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t55_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', d3_mcart_1)).
fof(t23_mcart_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1)))), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t23_mcart_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', fc6_relat_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t50_mcart_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/adbaee60113546c090a6d080dd6466de.E_input', t1_subset)).
fof(c_0_18, plain, ![X17, X18]:(~r2_hidden(X17,X18)|~r1_tarski(X18,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_19, plain, ![X40]:r1_tarski(k1_xboole_0,X40), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X104, X105]:((~r2_hidden(esk16_2(X104,X105),X104)|~r2_hidden(esk16_2(X104,X105),X105)|X104=X105)&(r2_hidden(esk16_2(X104,X105),X104)|r2_hidden(esk16_2(X104,X105),X105)|X104=X105)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
fof(c_0_23, plain, ![X75, X76, X77]:((r2_hidden(k1_mcart_1(X75),X76)|~r2_hidden(X75,k2_zfmisc_1(X76,X77)))&(r2_hidden(k2_mcart_1(X75),X77)|~r2_hidden(X75,k2_zfmisc_1(X76,X77)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
fof(c_0_24, plain, ![X19, X20, X21]:((~v1_xboole_0(X19)|~r2_hidden(X20,X19))&(r2_hidden(esk6_1(X21),X21)|v1_xboole_0(X21))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_25, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X6))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k5_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t69_mcart_1])).
cnf(c_0_26, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (r2_hidden(esk16_2(X1,X2),X1)|r2_hidden(esk16_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r2_hidden(esk6_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, negated_conjecture, ![X14, X15, X16]:(m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&((~m1_subset_1(X14,esk1_0)|(~m1_subset_1(X15,esk2_0)|(~m1_subset_1(X16,esk3_0)|(esk5_0!=k3_mcart_1(X14,X15,X16)|esk4_0=X14))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&esk4_0!=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_25])])])])).
fof(c_0_31, plain, ![X70, X71, X72]:k3_zfmisc_1(X70,X71,X72)=k2_zfmisc_1(k2_zfmisc_1(X70,X71),X72), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_32, plain, ![X91, X92, X93, X94]:k4_zfmisc_1(X91,X92,X93,X94)=k2_zfmisc_1(k3_zfmisc_1(X91,X92,X93),X94), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
cnf(c_0_33, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (k1_xboole_0=X1|r2_hidden(esk16_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|r2_hidden(k1_mcart_1(esk6_1(k2_zfmisc_1(X1,X2))),X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_36, plain, ![X45, X46]:(((~m1_subset_1(X46,X45)|r2_hidden(X46,X45)|v1_xboole_0(X45))&(~r2_hidden(X46,X45)|m1_subset_1(X46,X45)|v1_xboole_0(X45)))&((~m1_subset_1(X46,X45)|v1_xboole_0(X46)|~v1_xboole_0(X45))&(~v1_xboole_0(X46)|m1_subset_1(X46,X45)|~v1_xboole_0(X45)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, ![X31, X32, X33, X34, X35, X36, X37, X38]:(((~r2_hidden(X34,X33)|(X34=X31|X34=X32)|X33!=k2_tarski(X31,X32))&((X35!=X31|r2_hidden(X35,X33)|X33!=k2_tarski(X31,X32))&(X35!=X32|r2_hidden(X35,X33)|X33!=k2_tarski(X31,X32))))&(((esk10_3(X36,X37,X38)!=X36|~r2_hidden(esk10_3(X36,X37,X38),X38)|X38=k2_tarski(X36,X37))&(esk10_3(X36,X37,X38)!=X37|~r2_hidden(esk10_3(X36,X37,X38),X38)|X38=k2_tarski(X36,X37)))&(r2_hidden(esk10_3(X36,X37,X38),X38)|(esk10_3(X36,X37,X38)=X36|esk10_3(X36,X37,X38)=X37)|X38=k2_tarski(X36,X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_40, plain, ![X100, X101]:k1_enumset1(X100,X100,X101)=k2_tarski(X100,X101), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_41, plain, ![X82, X83, X84]:k2_enumset1(X82,X82,X83,X84)=k1_enumset1(X82,X83,X84), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_42, plain, ![X66, X67, X68, X69]:((X66=k1_xboole_0|(X67=k1_xboole_0|(X68=k1_xboole_0|X69=k1_xboole_0))|k4_zfmisc_1(X66,X67,X68,X69)!=k1_xboole_0)&((X66!=k1_xboole_0|k4_zfmisc_1(X66,X67,X68,X69)=k1_xboole_0)&((X67!=k1_xboole_0|k4_zfmisc_1(X66,X67,X68,X69)=k1_xboole_0)&((X68!=k1_xboole_0|k4_zfmisc_1(X66,X67,X68,X69)=k1_xboole_0)&(X69!=k1_xboole_0|k4_zfmisc_1(X66,X67,X68,X69)=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_mcart_1])])])])).
cnf(c_0_43, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_44, plain, (k1_xboole_0=X1|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_45, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_33, c_0_35])).
cnf(c_0_46, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_48, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k4_zfmisc_1(X1,X2,X3,X4)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_43, c_0_38])).
cnf(c_0_53, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X2,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_56, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_57, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),X1)=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_58, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_59, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_60, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_61, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_55])])).
cnf(c_0_62, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), c_0_59]), c_0_60])).
fof(c_0_63, plain, ![X107, X108, X109]:k3_mcart_1(X107,X108,X109)=k4_tarski(k4_tarski(X107,X108),X109), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_64, plain, ![X80, X81]:(~v1_relat_1(X81)|(~r2_hidden(X80,X81)|X80=k4_tarski(k1_mcart_1(X80),k2_mcart_1(X80)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_mcart_1])])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_26])).
fof(c_0_66, plain, ![X64, X65]:v1_relat_1(k2_zfmisc_1(X64,X65)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_67, plain, ![X87, X88, X89, X90]:((k5_mcart_1(X87,X88,X89,X90)=k1_mcart_1(k1_mcart_1(X90))|~m1_subset_1(X90,k3_zfmisc_1(X87,X88,X89))|X89=k1_xboole_0|X88=k1_xboole_0|X87=k1_xboole_0)&((k6_mcart_1(X87,X88,X89,X90)=k2_mcart_1(k1_mcart_1(X90))|~m1_subset_1(X90,k3_zfmisc_1(X87,X88,X89))|X89=k1_xboole_0|X88=k1_xboole_0|X87=k1_xboole_0)&(k7_mcart_1(X87,X88,X89,X90)=k2_mcart_1(X90)|~m1_subset_1(X90,k3_zfmisc_1(X87,X88,X89))|X89=k1_xboole_0|X88=k1_xboole_0|X87=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_68, negated_conjecture, (esk4_0=X1|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|esk5_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_69, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_70, plain, (X2=k4_tarski(k1_mcart_1(X2),k2_mcart_1(X2))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_71, negated_conjecture, (r2_hidden(k1_mcart_1(esk5_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_28, c_0_65])).
cnf(c_0_72, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_73, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_74, negated_conjecture, (esk4_0=X1|esk5_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk5_0)),k2_mcart_1(k1_mcart_1(esk5_0)))=k1_mcart_1(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])])).
fof(c_0_76, plain, ![X73, X74]:(~r2_hidden(X73,X74)|m1_subset_1(X73,X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_77, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_78, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_73, c_0_38])).
cnf(c_0_79, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk5_0))=esk4_0|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)|~m1_subset_1(X1,esk3_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_80, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_81, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk5_0)),esk2_0)), inference(spm,[status(thm)],[c_0_77, c_0_71])).
cnf(c_0_82, negated_conjecture, (esk4_0!=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_83, plain, (k5_mcart_1(X1,X2,X3,X4)=k5_mcart_1(X5,X6,X7,X4)|X5=k1_xboole_0|X6=k1_xboole_0|X7=k1_xboole_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X5,X6),X7))), inference(spm,[status(thm)],[c_0_78, c_0_78])).
cnf(c_0_84, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk5_0))=esk4_0|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])])).
cnf(c_0_85, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk5_0)),esk1_0)), inference(spm,[status(thm)],[c_0_28, c_0_71])).
cnf(c_0_86, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|k5_mcart_1(X3,X2,X1,esk5_0)!=esk4_0|~m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(X3,X2),X1))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_47])]), c_0_60]), c_0_59]), c_0_58])).
cnf(c_0_87, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk5_0))=esk4_0|k4_tarski(k1_mcart_1(esk5_0),X1)!=esk5_0|~m1_subset_1(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_80]), c_0_85])])).
cnf(c_0_88, negated_conjecture, (k4_tarski(k1_mcart_1(esk5_0),k2_mcart_1(esk5_0))=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_65]), c_0_72])])).
cnf(c_0_89, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|k1_mcart_1(k1_mcart_1(esk5_0))!=esk4_0|~m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_86, c_0_78])).
cnf(c_0_90, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk5_0))=esk4_0|~m1_subset_1(k2_mcart_1(esk5_0),esk3_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_91, negated_conjecture, (X1=k1_xboole_0|r2_hidden(k2_mcart_1(esk5_0),esk3_0)), inference(spm,[status(thm)],[c_0_77, c_0_62])).
cnf(c_0_92, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(X3,X2),X1))|~m1_subset_1(k2_mcart_1(esk5_0),esk3_0)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_93, negated_conjecture, (r2_hidden(k2_mcart_1(esk5_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_91]), c_0_26])).
cnf(c_0_94, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_80]), c_0_93])])).
cnf(c_0_95, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_47]), c_0_60]), c_0_59]), c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 96
# Proof object clause steps            : 59
# Proof object formula steps           : 37
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 59
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 54
# Processed clauses                    : 35025
# ...of these trivial                  : 74
# ...subsumed                          : 31731
# ...remaining for further processing  : 3220
# Other redundant clauses eliminated   : 1173
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 49
# Backward-rewritten                   : 79
# Generated clauses                    : 572617
# ...of the previous two non-trivial   : 420420
# Contextual simplify-reflections      : 22
# Paramodulations                      : 571287
# Factorizations                       : 160
# Equation resolutions                 : 1173
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3027
#    Positive orientable unit clauses  : 132
#    Positive unorientable unit clauses: 6
#    Negative unit clauses             : 71
#    Non-unit-clauses                  : 2818
# Current number of unprocessed clauses: 383527
# ...number of literals in the above   : 2254123
# Current number of archived formulas  : 0
# Current number of archived clauses   : 185
# Clause-clause subsumption calls (NU) : 1381617
# Rec. Clause-clause subsumption calls : 752634
# Non-unit clause-clause subsumptions  : 23408
# Unit Clause-clause subsumption calls : 1694
# Rewrite failures with RHS unbound    : 586
# BW rewrite match attempts            : 549
# BW rewrite match successes           : 68
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12320156

# -------------------------------------------------
# User time                : 8.568 s
# System time              : 0.126 s
# Total time               : 8.695 s
# Maximum resident set size: 317404 pages
