# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_waybel_7, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2))))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', t8_waybel_7)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', dt_k1_yellow_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', t1_subset)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', d3_tarski)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', d9_yellow_0)).
fof(t6_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_lattice3(X1,k1_xboole_0,X2)&r1_lattice3(X1,k1_xboole_0,X2)))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', t6_yellow_0)).
fof(d20_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v13_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X3,X4))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', d20_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', t4_subset)).
fof(t42_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))=>(r1_yellow_0(X1,k1_xboole_0)&r2_yellow_0(X1,u1_struct_0(X1)))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', t42_yellow_0)).
fof(d11_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', d11_yellow_0)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', d7_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', d10_xboole_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/add938693cd34ef7bc2a2c06e4f19688.E_input', t2_subset)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&(v1_yellow_0(X1)&l1_orders_2(X1))))))=>![X2]:((~(v1_xboole_0(X2))&(v2_waybel_0(X2,X1)&(v13_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v1_subset_1(X2,u1_struct_0(X1))<=>~(r2_hidden(k3_yellow_0(X1),X2)))))), inference(assume_negation,[status(cth)],[t8_waybel_7])).
fof(c_0_15, plain, ![X31, X32]:(~l1_orders_2(X31)|m1_subset_1(k1_yellow_0(X31,X32),u1_struct_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v1_yellow_0(esk1_0)&l1_orders_2(esk1_0))))))&((~v1_xboole_0(esk2_0)&(v2_waybel_0(esk2_0,esk1_0)&(v13_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v1_subset_1(esk2_0,u1_struct_0(esk1_0))|r2_hidden(k3_yellow_0(esk1_0),esk2_0))&(v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_17, plain, ![X15, X16]:(~r2_hidden(X15,X16)|m1_subset_1(X15,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_18, plain, ![X24, X25, X26, X27, X28]:((~r1_tarski(X24,X25)|(~r2_hidden(X26,X24)|r2_hidden(X26,X25)))&((r2_hidden(esk5_2(X27,X28),X27)|r1_tarski(X27,X28))&(~r2_hidden(esk5_2(X27,X28),X28)|r1_tarski(X27,X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_19, plain, ![X38, X39, X40, X41]:(((r2_lattice3(X38,X39,X40)|X40!=k1_yellow_0(X38,X39)|~r1_yellow_0(X38,X39)|~m1_subset_1(X40,u1_struct_0(X38))|~l1_orders_2(X38))&(~m1_subset_1(X41,u1_struct_0(X38))|(~r2_lattice3(X38,X39,X41)|r1_orders_2(X38,X40,X41))|X40!=k1_yellow_0(X38,X39)|~r1_yellow_0(X38,X39)|~m1_subset_1(X40,u1_struct_0(X38))|~l1_orders_2(X38)))&((m1_subset_1(esk6_3(X38,X39,X40),u1_struct_0(X38))|~r2_lattice3(X38,X39,X40)|X40=k1_yellow_0(X38,X39)|~r1_yellow_0(X38,X39)|~m1_subset_1(X40,u1_struct_0(X38))|~l1_orders_2(X38))&((r2_lattice3(X38,X39,esk6_3(X38,X39,X40))|~r2_lattice3(X38,X39,X40)|X40=k1_yellow_0(X38,X39)|~r1_yellow_0(X38,X39)|~m1_subset_1(X40,u1_struct_0(X38))|~l1_orders_2(X38))&(~r1_orders_2(X38,X40,esk6_3(X38,X39,X40))|~r2_lattice3(X38,X39,X40)|X40=k1_yellow_0(X38,X39)|~r1_yellow_0(X38,X39)|~m1_subset_1(X40,u1_struct_0(X38))|~l1_orders_2(X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
cnf(c_0_20, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X7, X8]:((r2_lattice3(X7,k1_xboole_0,X8)|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7))&(r1_lattice3(X7,k1_xboole_0,X8)|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_yellow_0])])])])).
cnf(c_0_23, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, ![X9, X10, X11, X12]:((~v13_waybel_0(X10,X9)|(~m1_subset_1(X11,u1_struct_0(X9))|(~m1_subset_1(X12,u1_struct_0(X9))|(~r2_hidden(X11,X10)|~r1_orders_2(X9,X11,X12)|r2_hidden(X12,X10))))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&((m1_subset_1(esk3_2(X9,X10),u1_struct_0(X9))|v13_waybel_0(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&((m1_subset_1(esk4_2(X9,X10),u1_struct_0(X9))|v13_waybel_0(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&(((r2_hidden(esk3_2(X9,X10),X10)|v13_waybel_0(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9))&(r1_orders_2(X9,esk3_2(X9,X10),esk4_2(X9,X10))|v13_waybel_0(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9)))&(~r2_hidden(esk4_2(X9,X10),X10)|v13_waybel_0(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|~l1_orders_2(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d20_waybel_0])])])])])).
fof(c_0_26, plain, ![X33, X34, X35]:(~r2_hidden(X33,X34)|~m1_subset_1(X34,k1_zfmisc_1(X35))|m1_subset_1(X33,X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_27, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_29, plain, ![X21]:((r1_yellow_0(X21,k1_xboole_0)|(v2_struct_0(X21)|(~v5_orders_2(X21)|(~v1_yellow_0(X21)|~l1_orders_2(X21)))))&(r2_yellow_0(X21,u1_struct_0(X21))|(v2_struct_0(X21)|(~v5_orders_2(X21)|(~v1_yellow_0(X21)|~l1_orders_2(X21)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t42_yellow_0])])])])).
cnf(c_0_30, plain, (r2_lattice3(X1,k1_xboole_0,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (r1_tarski(X1,X2)|m1_subset_1(esk5_2(X1,X2),X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, plain, (r2_hidden(X4,X1)|~v13_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X3,X4)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),X2)|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,X3)|~r1_yellow_0(esk1_0,X3)|~r2_lattice3(esk1_0,X3,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_21])])).
cnf(c_0_35, plain, (r1_yellow_0(X1,k1_xboole_0)|v2_struct_0(X1)|~v5_orders_2(X1)|~v1_yellow_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (v1_yellow_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_37, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_38, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_39, plain, ![X30]:(~l1_orders_2(X30)|k3_yellow_0(X30)=k1_yellow_0(X30,k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_yellow_0])])).
cnf(c_0_40, plain, (r1_tarski(u1_struct_0(X1),X2)|r2_lattice3(X1,k1_xboole_0,esk5_2(u1_struct_0(X1),X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|~r1_orders_2(X3,X4,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~v13_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(csr,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_42, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_43, negated_conjecture, (v13_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_44, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),k1_yellow_0(esk1_0,X2))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,X3)|~r1_yellow_0(esk1_0,X3)|~r2_lattice3(esk1_0,X3,k1_yellow_0(esk1_0,X2))), inference(spm,[status(thm)],[c_0_34, c_0_28])).
cnf(c_0_45, negated_conjecture, (r2_lattice3(esk1_0,k1_xboole_0,k1_yellow_0(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_28]), c_0_21])])).
cnf(c_0_46, negated_conjecture, (r1_yellow_0(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_21]), c_0_36]), c_0_37])]), c_0_38])).
cnf(c_0_47, plain, (k3_yellow_0(X1)=k1_yellow_0(X1,k1_xboole_0)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_48, plain, ![X22, X23]:((~v1_subset_1(X23,X22)|X23!=X22|~m1_subset_1(X23,k1_zfmisc_1(X22)))&(X23=X22|v1_subset_1(X23,X22)|~m1_subset_1(X23,k1_zfmisc_1(X22)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_49, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),X1)|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X2),esk5_2(u1_struct_0(esk1_0),X1))|k1_yellow_0(esk1_0,X2)!=k1_yellow_0(esk1_0,X3)|~r1_yellow_0(esk1_0,X3)|~r2_lattice3(esk1_0,X3,esk5_2(u1_struct_0(esk1_0),X1))), inference(spm,[status(thm)],[c_0_34, c_0_31])).
cnf(c_0_50, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),X1)|r2_lattice3(esk1_0,k1_xboole_0,esk5_2(u1_struct_0(esk1_0),X1))), inference(spm,[status(thm)],[c_0_40, c_0_21])).
cnf(c_0_51, negated_conjecture, (r2_hidden(X1,esk2_0)|~r1_orders_2(esk1_0,X2,X1)|~r2_hidden(X2,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_21])])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),k1_yellow_0(esk1_0,X2))|k1_yellow_0(esk1_0,X1)!=k1_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])])).
cnf(c_0_53, negated_conjecture, (r2_hidden(k3_yellow_0(esk1_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_54, negated_conjecture, (k3_yellow_0(esk1_0)=k1_yellow_0(esk1_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_47, c_0_21])).
cnf(c_0_55, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),X1)|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X2),esk5_2(u1_struct_0(esk1_0),X1))|k1_yellow_0(esk1_0,X2)!=k1_yellow_0(esk1_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_46])])).
cnf(c_0_57, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,X2,k1_yellow_0(esk1_0,X1))|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_51, c_0_28])).
cnf(c_0_58, negated_conjecture, (r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),k1_yellow_0(esk1_0,X1))), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_59, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)|~v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_60, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|v1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_55, c_0_42])).
fof(c_0_61, plain, ![X36, X37]:((~m1_subset_1(X36,k1_zfmisc_1(X37))|r1_tarski(X36,X37))&(~r1_tarski(X36,X37)|m1_subset_1(X36,k1_zfmisc_1(X37)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_62, plain, ![X17, X18]:(((r1_tarski(X17,X18)|X17!=X18)&(r1_tarski(X18,X17)|X17!=X18))&(~r1_tarski(X17,X18)|~r1_tarski(X18,X17)|X17=X18)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_63, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),X1)|r2_hidden(esk5_2(u1_struct_0(esk1_0),X1),esk2_0)|~r1_orders_2(esk1_0,X2,esk5_2(u1_struct_0(esk1_0),X1))|~r2_hidden(X2,esk2_0)), inference(spm,[status(thm)],[c_0_51, c_0_31])).
cnf(c_0_64, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),X1)|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,k1_xboole_0),esk5_2(u1_struct_0(esk1_0),X1))), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_65, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),esk2_0)|~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_66, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_67, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),X1)|r2_hidden(esk5_2(u1_struct_0(esk1_0),X1),esk2_0)|~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_70, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r2_hidden(k1_yellow_0(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_71, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_72, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_67, c_0_42])).
cnf(c_0_73, plain, (~v1_subset_1(X1,X2)|X1!=X2|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_74, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_75, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_68])).
fof(c_0_76, plain, ![X19, X20]:(~m1_subset_1(X19,X20)|(v1_xboole_0(X20)|r2_hidden(X19,X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_77, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k3_yellow_0(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_78, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_79, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|r1_tarski(u1_struct_0(esk1_0),X1)|r2_hidden(esk5_2(u1_struct_0(esk1_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_80, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0|~r1_tarski(u1_struct_0(esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_81, plain, (~v1_subset_1(X1,X1)|~m1_subset_1(X1,k1_zfmisc_1(X1))), inference(er,[status(thm)],[c_0_73])).
cnf(c_0_82, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_83, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_84, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(rw,[status(thm)],[c_0_77, c_0_54])).
cnf(c_0_85, negated_conjecture, (u1_struct_0(esk1_0)=esk2_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])).
cnf(c_0_86, plain, (~v1_subset_1(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_82])])).
cnf(c_0_87, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_83, c_0_28])).
cnf(c_0_88, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_89, negated_conjecture, (~r2_hidden(k1_yellow_0(esk1_0,k1_xboole_0),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_85]), c_0_86])).
cnf(c_0_90, negated_conjecture, (r2_hidden(k1_yellow_0(esk1_0,X1),esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_85]), c_0_85]), c_0_88])).
cnf(c_0_91, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_90])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 92
# Proof object clause steps            : 63
# Proof object formula steps           : 29
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 42
# Processed clauses                    : 370
# ...of these trivial                  : 1
# ...subsumed                          : 63
# ...remaining for further processing  : 306
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 150
# Generated clauses                    : 445
# ...of the previous two non-trivial   : 482
# Contextual simplify-reflections      : 6
# Paramodulations                      : 433
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 101
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 81
# Current number of unprocessed clauses: 193
# ...number of literals in the above   : 937
# Current number of archived formulas  : 0
# Current number of archived clauses   : 202
# Clause-clause subsumption calls (NU) : 8362
# Rec. Clause-clause subsumption calls : 1507
# Non-unit clause-clause subsumptions  : 79
# Unit Clause-clause subsumption calls : 121
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17811

# -------------------------------------------------
# User time                : 0.026 s
# System time              : 0.000 s
# Total time               : 0.026 s
# Maximum resident set size: 28752 pages
