# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t4_xboole_0)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t2_boole)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t3_xboole_0)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t66_xboole_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t39_xboole_1)).
fof(t37_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t37_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t36_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t3_boole)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', idempotence_k2_xboole_0)).
fof(t18_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k3_xboole_0(X2,X3))=>r1_tarski(X1,X2)), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t18_xboole_1)).
fof(t24_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3)), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t24_xboole_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t16_xboole_1)).
fof(t71_xboole_1, conjecture, ![X1, X2, X3]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X2)&(r1_xboole_0(X1,X2)&r1_xboole_0(X3,X2)))=>X1=X3), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t71_xboole_1)).
fof(t31_xboole_1, axiom, ![X1, X2, X3]:r1_tarski(k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)),k2_xboole_0(X2,X3)), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t31_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', d10_xboole_0)).
fof(t20_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&(r1_tarski(X1,X3)&![X4]:((r1_tarski(X4,X2)&r1_tarski(X4,X3))=>r1_tarski(X4,X1))))=>X1=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t20_xboole_1)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t44_xboole_1)).
fof(t45_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1))), file('data/MPTP/proofs/ae08008f8c4546569afd13c44de3591f.E_input', t45_xboole_1)).
fof(c_0_18, plain, ![X44, X45, X47, X48, X49]:((r1_xboole_0(X44,X45)|r2_hidden(esk6_2(X44,X45),k3_xboole_0(X44,X45)))&(~r2_hidden(X49,k3_xboole_0(X47,X48))|~r1_xboole_0(X47,X48))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_19, plain, ![X23]:k3_xboole_0(X23,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_20, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X17, X18, X20, X21, X22]:(((r2_hidden(esk4_2(X17,X18),X17)|r1_xboole_0(X17,X18))&(r2_hidden(esk4_2(X17,X18),X18)|r1_xboole_0(X17,X18)))&(~r2_hidden(X22,X20)|~r2_hidden(X22,X21)|~r1_xboole_0(X20,X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_23, plain, ![X24]:((r1_xboole_0(X24,X24)|X24!=k1_xboole_0)&(X24=k1_xboole_0|~r1_xboole_0(X24,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_24, plain, (~r2_hidden(X1,k1_xboole_0)|~r1_xboole_0(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, plain, (r2_hidden(esk4_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (r1_xboole_0(X1,k1_xboole_0)|~r1_xboole_0(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_26])).
fof(c_0_29, plain, ![X10, X11]:k2_xboole_0(X10,k4_xboole_0(X11,X10))=k2_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_30, plain, ![X15, X16]:((k4_xboole_0(X15,X16)!=k1_xboole_0|r1_tarski(X15,X16))&(~r1_tarski(X15,X16)|k4_xboole_0(X15,X16)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_xboole_1])])).
fof(c_0_31, plain, ![X32, X33]:r1_tarski(k4_xboole_0(X32,X33),X32), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_32, plain, ![X52]:k4_xboole_0(X52,k1_xboole_0)=X52, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_33, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_34, plain, ![X28]:k2_xboole_0(X28,X28)=X28, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_35, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_20, c_0_25])).
cnf(c_0_41, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_33])])).
cnf(c_0_42, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_43, plain, ![X34, X35, X36]:(~r1_tarski(X34,k3_xboole_0(X35,X36))|r1_tarski(X34,X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_xboole_1])])).
fof(c_0_44, plain, ![X29, X30, X31]:k2_xboole_0(X29,k3_xboole_0(X30,X31))=k3_xboole_0(k2_xboole_0(X29,X30),k2_xboole_0(X29,X31)), inference(variable_rename,[status(thm)],[t24_xboole_1])).
cnf(c_0_45, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, plain, (k2_xboole_0(X1,k1_xboole_0)=k2_xboole_0(X1,X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_47, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_48, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_50, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_51, plain, ![X12, X13, X14]:k3_xboole_0(k3_xboole_0(X12,X13),X14)=k3_xboole_0(X12,k3_xboole_0(X13,X14)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_52, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k2_xboole_0(X1,X2),k2_xboole_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_53, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_47])])).
cnf(c_0_54, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_48, c_0_49])).
fof(c_0_55, negated_conjecture, ~(![X1, X2, X3]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X2)&(r1_xboole_0(X1,X2)&r1_xboole_0(X3,X2)))=>X1=X3)), inference(assume_negation,[status(cth)],[t71_xboole_1])).
cnf(c_0_56, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_50, c_0_47])).
cnf(c_0_57, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_58, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54]), c_0_53])).
fof(c_0_59, negated_conjecture, ((k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk3_0,esk2_0)&(r1_xboole_0(esk1_0,esk2_0)&r1_xboole_0(esk3_0,esk2_0)))&esk1_0!=esk3_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_55])])])).
fof(c_0_60, plain, ![X41, X42, X43]:r1_tarski(k2_xboole_0(k3_xboole_0(X41,X42),k3_xboole_0(X41,X43)),k2_xboole_0(X42,X43)), inference(variable_rename,[status(thm)],[t31_xboole_1])).
fof(c_0_61, plain, ![X8, X9]:(((r1_tarski(X8,X9)|X8!=X9)&(r1_tarski(X9,X8)|X8!=X9))&(~r1_tarski(X8,X9)|~r1_tarski(X9,X8)|X8=X9)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_62, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_58, c_0_53])).
cnf(c_0_64, negated_conjecture, (k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65, plain, (r1_tarski(k2_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X3)),k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_67, plain, (r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(X1,k3_xboole_0(X2,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_57])).
cnf(c_0_68, negated_conjecture, (k3_xboole_0(k2_xboole_0(esk1_0,esk2_0),k2_xboole_0(esk3_0,X1))=k2_xboole_0(esk3_0,k3_xboole_0(esk2_0,X1))), inference(spm,[status(thm)],[c_0_52, c_0_64])).
cnf(c_0_69, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_45]), c_0_45])).
cnf(c_0_70, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_62])])).
cnf(c_0_71, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_72, negated_conjecture, (k3_xboole_0(k2_xboole_0(esk1_0,esk2_0),esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_53]), c_0_21]), c_0_53])).
cnf(c_0_73, plain, (r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_74, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk4_2(X1,X2),X3)|~r1_xboole_0(X3,X1)), inference(spm,[status(thm)],[c_0_71, c_0_42])).
cnf(c_0_75, negated_conjecture, (r1_tarski(k3_xboole_0(X1,esk3_0),k3_xboole_0(X1,k2_xboole_0(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_62, c_0_72])).
cnf(c_0_76, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_73]), c_0_73])])).
cnf(c_0_77, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_74, c_0_25])).
cnf(c_0_78, negated_conjecture, (r1_xboole_0(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_79, plain, (k3_xboole_0(k2_xboole_0(X1,X2),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_21]), c_0_53])).
cnf(c_0_80, negated_conjecture, (r1_tarski(k3_xboole_0(X1,esk3_0),k3_xboole_0(k2_xboole_0(esk1_0,esk2_0),X1))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_81, negated_conjecture, (r1_xboole_0(esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
fof(c_0_82, plain, ![X37, X38, X39]:(((r1_tarski(esk5_3(X37,X38,X39),X38)|~r1_tarski(X37,X39)|~r1_tarski(X37,X38)|X37=k3_xboole_0(X38,X39))&(r1_tarski(esk5_3(X37,X38,X39),X39)|~r1_tarski(X37,X39)|~r1_tarski(X37,X38)|X37=k3_xboole_0(X38,X39)))&(~r1_tarski(esk5_3(X37,X38,X39),X37)|~r1_tarski(X37,X39)|~r1_tarski(X37,X38)|X37=k3_xboole_0(X38,X39))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_xboole_1])])])])).
fof(c_0_83, plain, ![X25, X26, X27]:(~r1_tarski(k4_xboole_0(X25,X26),X27)|r1_tarski(X25,k2_xboole_0(X26,X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
cnf(c_0_84, plain, (r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(X1,k2_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_62, c_0_79])).
cnf(c_0_85, negated_conjecture, (r1_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_86, negated_conjecture, (r1_tarski(k3_xboole_0(esk3_0,k2_xboole_0(esk1_0,X1)),k2_xboole_0(esk1_0,k3_xboole_0(esk2_0,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_52]), c_0_76])).
cnf(c_0_87, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_48, c_0_81])).
cnf(c_0_88, plain, (X1=k3_xboole_0(X2,X3)|~r1_tarski(esk5_3(X1,X2,X3),X1)|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_89, plain, (r1_tarski(esk5_3(X1,X2,X3),X2)|X1=k3_xboole_0(X2,X3)|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_90, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_91, plain, (r1_tarski(k3_xboole_0(X1,X2),k3_xboole_0(k2_xboole_0(X2,X3),X1))), inference(spm,[status(thm)],[c_0_84, c_0_76])).
cnf(c_0_92, negated_conjecture, (r1_xboole_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_77, c_0_85])).
fof(c_0_93, plain, ![X50, X51]:(~r1_tarski(X50,X51)|X51=k2_xboole_0(X50,k4_xboole_0(X51,X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
cnf(c_0_94, negated_conjecture, (r1_tarski(k3_xboole_0(esk3_0,k2_xboole_0(esk1_0,esk3_0)),esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_53])).
cnf(c_0_95, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_47])])).
cnf(c_0_96, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_47]), c_0_35])).
cnf(c_0_97, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,k2_xboole_0(esk3_0,X1)),k2_xboole_0(esk3_0,k3_xboole_0(esk2_0,X1)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_68]), c_0_76])).
cnf(c_0_98, negated_conjecture, (k3_xboole_0(esk2_0,esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_48, c_0_92])).
cnf(c_0_99, plain, (X2=k2_xboole_0(X1,k4_xboole_0(X2,X1))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_100, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])])).
cnf(c_0_101, negated_conjecture, (esk1_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_102, negated_conjecture, (r1_tarski(k3_xboole_0(esk1_0,k2_xboole_0(esk3_0,esk1_0)),esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_53])).
cnf(c_0_103, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_99, c_0_35])).
cnf(c_0_104, negated_conjecture, (~r1_tarski(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_100]), c_0_101])).
cnf(c_0_105, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_63]), c_0_100])]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 69
# Proof object formula steps           : 37
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 41
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 3153
# ...of these trivial                  : 584
# ...subsumed                          : 1683
# ...remaining for further processing  : 886
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 18
# Backward-rewritten                   : 148
# Generated clauses                    : 48101
# ...of the previous two non-trivial   : 18921
# Contextual simplify-reflections      : 3
# Paramodulations                      : 48097
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 717
#    Positive orientable unit clauses  : 454
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 255
# Current number of unprocessed clauses: 15314
# ...number of literals in the above   : 26596
# Current number of archived formulas  : 0
# Current number of archived clauses   : 166
# Clause-clause subsumption calls (NU) : 20450
# Rec. Clause-clause subsumption calls : 10396
# Non-unit clause-clause subsumptions  : 1472
# Unit Clause-clause subsumption calls : 888
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2796
# BW rewrite match successes           : 687
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 423757

# -------------------------------------------------
# User time                : 0.220 s
# System time              : 0.010 s
# Total time               : 0.230 s
# Maximum resident set size: 28760 pages
