# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_SE_CS_SP_PI_PS_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t21_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3)))))), file('data/MPTP/proofs/ae4e57d1cffb40aab7637ab0e67e1ac9.E_input', t21_waybel_0)).
fof(dt_k5_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/ae4e57d1cffb40aab7637ab0e67e1ac9.E_input', dt_k5_waybel_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/ae4e57d1cffb40aab7637ab0e67e1ac9.E_input', t4_subset)).
fof(t18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/ae4e57d1cffb40aab7637ab0e67e1ac9.E_input', t18_waybel_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/ae4e57d1cffb40aab7637ab0e67e1ac9.E_input', d3_tarski)).
fof(t17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/ae4e57d1cffb40aab7637ab0e67e1ac9.E_input', t17_waybel_0)).
fof(d8_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X3,X4)))))), file('data/MPTP/proofs/ae4e57d1cffb40aab7637ab0e67e1ac9.E_input', d8_lattice3)).
fof(t12_yellow_0, axiom, ![X1]:((v4_orders_2(X1)&l1_orders_2(X1))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X3)&r1_orders_2(X1,X4,X3))=>r1_lattice3(X1,X2,X4))))), file('data/MPTP/proofs/ae4e57d1cffb40aab7637ab0e67e1ac9.E_input', t12_yellow_0)).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v4_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)=>r1_tarski(k5_waybel_0(X1,X2),k5_waybel_0(X1,X3))))))), inference(assume_negation,[status(cth)],[t21_waybel_0])).
fof(c_0_9, plain, ![X13, X14]:(v2_struct_0(X13)|~l1_orders_2(X13)|~m1_subset_1(X14,u1_struct_0(X13))|m1_subset_1(k5_waybel_0(X13,X14),k1_zfmisc_1(u1_struct_0(X13)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
fof(c_0_10, negated_conjecture, ((~v2_struct_0(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(r1_orders_2(esk1_0,esk2_0,esk3_0)&~r1_tarski(k5_waybel_0(esk1_0,esk2_0),k5_waybel_0(esk1_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
fof(c_0_11, plain, ![X21, X22, X23]:(~r2_hidden(X21,X22)|~m1_subset_1(X22,k1_zfmisc_1(X23))|m1_subset_1(X21,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_12, plain, (v2_struct_0(X1)|m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_16, plain, ![X28, X29, X30]:((~r2_hidden(X30,k6_waybel_0(X28,X29))|r1_orders_2(X28,X29,X30)|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|~l1_orders_2(X28)))&(~r1_orders_2(X28,X29,X30)|r2_hidden(X30,k6_waybel_0(X28,X29))|~m1_subset_1(X30,u1_struct_0(X28))|~m1_subset_1(X29,u1_struct_0(X28))|(v2_struct_0(X28)|~l1_orders_2(X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])).
cnf(c_0_17, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(k5_waybel_0(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])]), c_0_15])).
fof(c_0_19, plain, ![X15, X16, X17, X18, X19]:((~r1_tarski(X15,X16)|(~r2_hidden(X17,X15)|r2_hidden(X17,X16)))&((r2_hidden(esk5_2(X18,X19),X18)|r1_tarski(X18,X19))&(~r2_hidden(esk5_2(X18,X19),X19)|r1_tarski(X18,X19)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_20, plain, ![X31, X32, X33]:((~r2_hidden(X33,k5_waybel_0(X31,X32))|r1_orders_2(X31,X33,X32)|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|~l1_orders_2(X31)))&(~r1_orders_2(X31,X33,X32)|r2_hidden(X33,k5_waybel_0(X31,X32))|~m1_subset_1(X33,u1_struct_0(X31))|~m1_subset_1(X32,u1_struct_0(X31))|(v2_struct_0(X31)|~l1_orders_2(X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])).
fof(c_0_21, plain, ![X8, X9, X10, X11]:((~r1_lattice3(X8,X9,X10)|(~m1_subset_1(X11,u1_struct_0(X8))|(~r2_hidden(X11,X9)|r1_orders_2(X8,X10,X11)))|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&((m1_subset_1(esk4_3(X8,X9,X10),u1_struct_0(X8))|r1_lattice3(X8,X9,X10)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&((r2_hidden(esk4_3(X8,X9,X10),X9)|r1_lattice3(X8,X9,X10)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))&(~r1_orders_2(X8,X10,esk4_3(X8,X9,X10))|r1_lattice3(X8,X9,X10)|~m1_subset_1(X10,u1_struct_0(X8))|~l1_orders_2(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_lattice3])])])])])).
cnf(c_0_22, plain, (r2_hidden(X3,k6_waybel_0(X1,X2))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_24, plain, ![X24, X25, X26, X27]:(~v4_orders_2(X24)|~l1_orders_2(X24)|(~m1_subset_1(X26,u1_struct_0(X24))|(~m1_subset_1(X27,u1_struct_0(X24))|(~r1_lattice3(X24,X25,X26)|~r1_orders_2(X24,X27,X26)|r1_lattice3(X24,X25,X27))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_yellow_0])])])).
cnf(c_0_25, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k5_waybel_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (r1_orders_2(X2,X1,X3)|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (r1_orders_2(X2,X3,X1)|v2_struct_0(X2)|~r2_hidden(X1,k6_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_29, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (r1_lattice3(X1,X3,X2)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (r1_orders_2(X1,X3,X4)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk3_0,k6_waybel_0(esk1_0,X1))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_14])]), c_0_15])).
cnf(c_0_33, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_34, plain, (r1_lattice3(X1,X4,X3)|~v4_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattice3(X1,X4,X2)|~r1_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, negated_conjecture, (r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|m1_subset_1(esk5_2(k5_waybel_0(esk1_0,esk2_0),X1),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_36, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|~r2_hidden(X1,k5_waybel_0(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_13]), c_0_14])]), c_0_15]), c_0_25])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,X1)|~r2_hidden(X1,k6_waybel_0(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_13]), c_0_14])]), c_0_15])).
cnf(c_0_39, negated_conjecture, (r1_lattice3(esk1_0,X1,esk2_0)|m1_subset_1(esk4_3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_13]), c_0_14])])).
cnf(c_0_40, negated_conjecture, (r1_lattice3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk4_3(esk1_0,X1,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_13]), c_0_14])])).
cnf(c_0_41, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_42, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|~r2_hidden(esk3_0,X2)|~r1_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_23]), c_0_14])])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk3_0,k6_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_13])])).
cnf(c_0_44, negated_conjecture, (r1_lattice3(esk1_0,X1,esk5_2(k5_waybel_0(esk1_0,esk2_0),X2))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X2)|~r1_lattice3(esk1_0,X1,X3)|~r1_orders_2(esk1_0,esk5_2(k5_waybel_0(esk1_0,esk2_0),X2),X3)|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_14]), c_0_36])])).
cnf(c_0_45, negated_conjecture, (r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk5_2(k5_waybel_0(esk1_0,esk2_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_37, c_0_26])).
cnf(c_0_46, negated_conjecture, (r1_lattice3(esk1_0,X1,esk2_0)|~r2_hidden(esk4_3(esk1_0,X1,esk2_0),k6_waybel_0(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk4_3(esk1_0,X1,esk2_0),X1)|r1_lattice3(esk1_0,X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_13]), c_0_14])])).
cnf(c_0_48, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|~r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, negated_conjecture, (r1_lattice3(esk1_0,X1,esk5_2(k5_waybel_0(esk1_0,esk2_0),X2))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X2)|~r1_lattice3(esk1_0,X1,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_13]), c_0_45])).
cnf(c_0_50, negated_conjecture, (r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51, plain, (r2_hidden(X2,k5_waybel_0(X1,X3))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_52, negated_conjecture, (r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk5_2(k5_waybel_0(esk1_0,esk2_0),X1),esk3_0)|~r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk5_2(k5_waybel_0(esk1_0,esk2_0),X1))), inference(spm,[status(thm)],[c_0_48, c_0_35])).
cnf(c_0_53, negated_conjecture, (r1_lattice3(esk1_0,k6_waybel_0(esk1_0,esk2_0),esk5_2(k5_waybel_0(esk1_0,esk2_0),X1))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_23]), c_0_14])]), c_0_15])).
cnf(c_0_55, negated_conjecture, (r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)|r1_orders_2(esk1_0,esk5_2(k5_waybel_0(esk1_0,esk2_0),X1),esk3_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_56, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk5_2(k5_waybel_0(esk1_0,esk2_0),X1),k5_waybel_0(esk1_0,esk3_0))|r1_tarski(k5_waybel_0(esk1_0,esk2_0),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_35])).
cnf(c_0_58, negated_conjecture, (~r1_tarski(k5_waybel_0(esk1_0,esk2_0),k5_waybel_0(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_59, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 60
# Proof object clause steps            : 43
# Proof object formula steps           : 17
# Proof object conjectures             : 33
# Proof object clause conjectures      : 30
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 275
# ...of these trivial                  : 0
# ...subsumed                          : 40
# ...remaining for further processing  : 235
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 1
# Generated clauses                    : 676
# ...of the previous two non-trivial   : 664
# Contextual simplify-reflections      : 23
# Paramodulations                      : 676
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 209
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 194
# Current number of unprocessed clauses: 429
# ...number of literals in the above   : 1954
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 6103
# Rec. Clause-clause subsumption calls : 3414
# Non-unit clause-clause subsumptions  : 67
# Unit Clause-clause subsumption calls : 54
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 24065

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.003 s
# Total time               : 0.022 s
# Maximum resident set size: 28792 pages
