# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t90_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k4_xboole_0(X1,k3_xboole_0(X1,X2)),X2), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t90_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t48_xboole_1)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t47_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', d7_xboole_0)).
fof(t75_xboole_1, axiom, ![X1, X2]:~((~(r1_xboole_0(X1,X2))&r1_xboole_0(k3_xboole_0(X1,X2),X2))), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t75_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', d10_xboole_0)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t36_xboole_1)).
fof(t86_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_xboole_0(X1,X3))=>r1_tarski(X1,k4_xboole_0(X2,X3))), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t86_xboole_1)).
fof(t106_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k4_xboole_0(X2,X3))=>(r1_tarski(X1,X2)&r1_xboole_0(X1,X3))), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t106_xboole_1)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t4_boole)).
fof(t155_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>r1_tarski(k6_subset_1(k9_relat_1(X3,X1),k9_relat_1(X3,X2)),k9_relat_1(X3,k6_subset_1(X1,X2)))), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t155_relat_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', redefinition_k6_subset_1)).
fof(t109_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(k4_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t109_xboole_1)).
fof(t123_funct_1, conjecture, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(v2_funct_1(X3)=>k9_relat_1(X3,k6_subset_1(X1,X2))=k6_subset_1(k9_relat_1(X3,X1),k9_relat_1(X3,X2)))), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t123_funct_1)).
fof(t149_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k9_relat_1(X1,k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t149_relat_1)).
fof(t121_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(v2_funct_1(X3)=>k9_relat_1(X3,k3_xboole_0(X1,X2))=k3_xboole_0(k9_relat_1(X3,X1),k9_relat_1(X3,X2)))), file('data/MPTP/proofs/aedf400dbdaf426084acdea9d71b67ee.E_input', t121_funct_1)).
fof(c_0_16, plain, ![X8, X9]:r1_xboole_0(k4_xboole_0(X8,k3_xboole_0(X8,X9)),X9), inference(variable_rename,[status(thm)],[t90_xboole_1])).
fof(c_0_17, plain, ![X4, X5]:k4_xboole_0(X4,k4_xboole_0(X4,X5))=k3_xboole_0(X4,X5), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_18, plain, ![X10, X11]:k4_xboole_0(X10,k3_xboole_0(X10,X11))=k4_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t47_xboole_1])).
fof(c_0_19, plain, ![X12, X13]:((~r1_xboole_0(X12,X13)|k3_xboole_0(X12,X13)=k1_xboole_0)&(k3_xboole_0(X12,X13)!=k1_xboole_0|r1_xboole_0(X12,X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_20, plain, (r1_xboole_0(k4_xboole_0(X1,k3_xboole_0(X1,X2)),X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2))),X2)), inference(rw,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_21])).
fof(c_0_27, plain, ![X14, X15]:(r1_xboole_0(X14,X15)|~r1_xboole_0(k3_xboole_0(X14,X15),X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t75_xboole_1])])])).
cnf(c_0_28, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,k4_xboole_0(X1,X2))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_23, c_0_21])).
cnf(c_0_29, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_24, c_0_21])).
cnf(c_0_30, plain, (r1_xboole_0(k4_xboole_0(X1,X2),X2)), inference(rw,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k3_xboole_0(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (r1_xboole_0(X1,k4_xboole_0(X1,X2))|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_26])).
cnf(c_0_33, plain, (k4_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(k4_xboole_0(X1,X2),X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_34, plain, ![X25, X26]:(((r1_tarski(X25,X26)|X25!=X26)&(r1_tarski(X26,X25)|X25!=X26))&(~r1_tarski(X25,X26)|~r1_tarski(X26,X25)|X25=X26)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_35, plain, ![X16, X17]:r1_tarski(k4_xboole_0(X16,X17),X16), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_36, plain, ![X18, X19, X20]:(~r1_tarski(X18,X19)|~r1_xboole_0(X18,X20)|r1_tarski(X18,k4_xboole_0(X19,X20))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t86_xboole_1])])).
cnf(c_0_37, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(rw,[status(thm)],[c_0_31, c_0_21])).
cnf(c_0_38, plain, (r1_xboole_0(k4_xboole_0(X1,X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_39, plain, ![X21, X22, X23]:((r1_tarski(X21,X22)|~r1_tarski(X21,k4_xboole_0(X22,X23)))&(r1_xboole_0(X21,X23)|~r1_tarski(X21,k4_xboole_0(X22,X23)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t106_xboole_1])])])).
fof(c_0_40, plain, ![X24]:k4_xboole_0(k1_xboole_0,X24)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_41, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (r1_tarski(X1,k4_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_46, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, plain, (k4_xboole_0(X1,X2)=X1|~r1_tarski(X1,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_49, plain, (r1_tarski(X1,k4_xboole_0(X2,k1_xboole_0))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (r1_xboole_0(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_29, c_0_44])).
cnf(c_0_53, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_54, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_55, plain, ![X33, X34, X35]:(~v1_relat_1(X35)|r1_tarski(k6_subset_1(k9_relat_1(X35,X33),k9_relat_1(X35,X34)),k9_relat_1(X35,k6_subset_1(X33,X34)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t155_relat_1])])).
fof(c_0_56, plain, ![X6, X7]:k6_subset_1(X6,X7)=k4_xboole_0(X6,X7), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_57, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_29, c_0_51])).
fof(c_0_58, plain, ![X27, X28, X29]:(~r1_tarski(X27,X28)|r1_tarski(k4_xboole_0(X27,X29),X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t109_xboole_1])])).
fof(c_0_59, negated_conjecture, ~(![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(v2_funct_1(X3)=>k9_relat_1(X3,k6_subset_1(X1,X2))=k6_subset_1(k9_relat_1(X3,X1),k9_relat_1(X3,X2))))), inference(assume_negation,[status(cth)],[t123_funct_1])).
cnf(c_0_60, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X3,X2))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_43, c_0_30])).
cnf(c_0_61, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_62, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_54, c_0_42])).
cnf(c_0_63, plain, (r1_tarski(k6_subset_1(k9_relat_1(X1,X2),k9_relat_1(X1,X3)),k9_relat_1(X1,k6_subset_1(X2,X3)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_65, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_26, c_0_57])).
cnf(c_0_66, plain, (r1_tarski(k4_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_67, plain, ![X36]:(~v1_relat_1(X36)|k9_relat_1(X36,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t149_relat_1])])).
fof(c_0_68, negated_conjecture, ((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&(v2_funct_1(esk3_0)&k9_relat_1(esk3_0,k6_subset_1(esk1_0,esk2_0))!=k6_subset_1(k9_relat_1(esk3_0,esk1_0),k9_relat_1(esk3_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_59])])])).
cnf(c_0_69, plain, (r1_tarski(k4_xboole_0(X1,X2),k1_xboole_0)|~r1_tarski(k4_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_70, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(k4_xboole_0(X1,X2),X3),X4),X1)), inference(spm,[status(thm)],[c_0_54, c_0_62])).
cnf(c_0_71, plain, (r1_tarski(k4_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,X3)),k9_relat_1(X1,k4_xboole_0(X2,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64]), c_0_64])).
cnf(c_0_72, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_73, plain, (k9_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_74, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_75, plain, (r1_tarski(k4_xboole_0(X1,X2),k1_xboole_0)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_69, c_0_66])).
cnf(c_0_76, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_70, c_0_33])).
cnf(c_0_77, plain, (r1_tarski(k4_xboole_0(k9_relat_1(X1,k4_xboole_0(X2,X3)),k9_relat_1(X1,X4)),k9_relat_1(X1,k1_xboole_0))|~v1_relat_1(X1)|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_78, negated_conjecture, (k9_relat_1(esk3_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_79, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_75]), c_0_76])])).
cnf(c_0_80, negated_conjecture, (r1_tarski(k9_relat_1(esk3_0,k4_xboole_0(X1,X2)),k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_53]), c_0_78]), c_0_74])])).
cnf(c_0_81, negated_conjecture, (k9_relat_1(esk3_0,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_53])).
cnf(c_0_82, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_54, c_0_52])).
cnf(c_0_83, plain, (r1_tarski(X1,k4_xboole_0(X2,k4_xboole_0(X1,X3)))|k4_xboole_0(X1,X3)!=k1_xboole_0|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_43, c_0_32])).
cnf(c_0_84, plain, (k4_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,X3))=k9_relat_1(X1,k4_xboole_0(X2,X3))|~v1_relat_1(X1)|~r1_tarski(k9_relat_1(X1,k4_xboole_0(X2,X3)),k4_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,X3)))), inference(spm,[status(thm)],[c_0_41, c_0_71])).
cnf(c_0_85, negated_conjecture, (k9_relat_1(esk3_0,X1)=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_81, c_0_53])).
cnf(c_0_86, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_82, c_0_75])).
cnf(c_0_87, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_83]), c_0_50])])).
cnf(c_0_88, negated_conjecture, (k4_xboole_0(k9_relat_1(esk3_0,X1),k9_relat_1(esk3_0,X2))=k1_xboole_0|~r1_tarski(k4_xboole_0(X1,X2),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_74]), c_0_76])])).
cnf(c_0_89, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X3)!=k1_xboole_0|~r1_tarski(X1,k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_90, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_60]), c_0_50])])).
fof(c_0_91, plain, ![X37, X38, X39]:(~v1_relat_1(X39)|~v1_funct_1(X39)|(~v2_funct_1(X39)|k9_relat_1(X39,k3_xboole_0(X37,X38))=k3_xboole_0(k9_relat_1(X39,X37),k9_relat_1(X39,X38)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t121_funct_1])])).
cnf(c_0_92, negated_conjecture, (k9_relat_1(esk3_0,X1)=k1_xboole_0|~r1_tarski(k4_xboole_0(X1,X2),k1_xboole_0)|~r1_tarski(X2,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_85]), c_0_53])).
cnf(c_0_93, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_50])])).
cnf(c_0_94, plain, (k9_relat_1(X1,k3_xboole_0(X2,X3))=k3_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_95, negated_conjecture, (k9_relat_1(esk3_0,X1)=k1_xboole_0|k4_xboole_0(X1,X2)!=k1_xboole_0|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_96, plain, (k9_relat_1(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))=k4_xboole_0(k9_relat_1(X1,X2),k4_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,X3)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_funct_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_94, c_0_21]), c_0_21])).
cnf(c_0_97, negated_conjecture, (k9_relat_1(esk3_0,X1)=k1_xboole_0|k4_xboole_0(X1,k4_xboole_0(X2,X3))!=k1_xboole_0|k4_xboole_0(X2,X3)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_95, c_0_93])).
cnf(c_0_98, negated_conjecture, (v2_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_99, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_100, plain, (r1_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,X3))|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xboole_0(k9_relat_1(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3))),k9_relat_1(X1,X3))), inference(spm,[status(thm)],[c_0_37, c_0_96])).
cnf(c_0_101, negated_conjecture, (k9_relat_1(esk3_0,k4_xboole_0(X1,X2))=k1_xboole_0|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_97, c_0_61])).
cnf(c_0_102, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_30, c_0_47])).
cnf(c_0_103, plain, (r1_tarski(k4_xboole_0(k9_relat_1(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3))),X4),k9_relat_1(X1,X2))|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_62, c_0_96])).
cnf(c_0_104, plain, (k4_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3))))=k4_xboole_0(k9_relat_1(X1,X2),k9_relat_1(X1,X3))|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_26, c_0_96])).
cnf(c_0_105, negated_conjecture, (k9_relat_1(esk3_0,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=k1_xboole_0|~r1_tarski(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_85]), c_0_53]), c_0_61]), c_0_98]), c_0_99]), c_0_74])])).
cnf(c_0_106, negated_conjecture, (r1_xboole_0(k9_relat_1(esk3_0,X1),k9_relat_1(esk3_0,X2))|k4_xboole_0(X1,k4_xboole_0(X1,X2))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_98]), c_0_99]), c_0_74]), c_0_102])])).
cnf(c_0_107, plain, (r1_tarski(k4_xboole_0(k9_relat_1(X1,k4_xboole_0(X2,X3)),X4),k9_relat_1(X1,X2))|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_103, c_0_26])).
cnf(c_0_108, negated_conjecture, (k4_xboole_0(k9_relat_1(esk3_0,X1),k9_relat_1(esk3_0,X2))=k9_relat_1(esk3_0,X1)|~r1_tarski(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_53]), c_0_98]), c_0_99]), c_0_74])])).
cnf(c_0_109, negated_conjecture, (r1_xboole_0(k9_relat_1(esk3_0,k4_xboole_0(X1,X2)),k9_relat_1(esk3_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_90]), c_0_61])])).
cnf(c_0_110, negated_conjecture, (r1_tarski(k9_relat_1(esk3_0,k4_xboole_0(X1,X2)),k9_relat_1(esk3_0,X1))|~r1_tarski(X3,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_98]), c_0_99]), c_0_74])])).
cnf(c_0_111, negated_conjecture, (k9_relat_1(esk3_0,k6_subset_1(esk1_0,esk2_0))!=k6_subset_1(k9_relat_1(esk3_0,esk1_0),k9_relat_1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_112, negated_conjecture, (r1_tarski(k9_relat_1(esk3_0,k4_xboole_0(X1,X2)),k4_xboole_0(X3,k9_relat_1(esk3_0,X2)))|~r1_tarski(k9_relat_1(esk3_0,k4_xboole_0(X1,X2)),X3)), inference(spm,[status(thm)],[c_0_43, c_0_109])).
cnf(c_0_113, negated_conjecture, (r1_tarski(k9_relat_1(esk3_0,k4_xboole_0(X1,X2)),k9_relat_1(esk3_0,X1))), inference(spm,[status(thm)],[c_0_110, c_0_76])).
cnf(c_0_114, negated_conjecture, (k9_relat_1(esk3_0,k4_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(k9_relat_1(esk3_0,esk1_0),k9_relat_1(esk3_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_64]), c_0_64])).
cnf(c_0_115, negated_conjecture, (k4_xboole_0(k9_relat_1(esk3_0,X1),k9_relat_1(esk3_0,X2))=k9_relat_1(esk3_0,k4_xboole_0(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_112]), c_0_74]), c_0_113])])).
cnf(c_0_116, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_115])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 117
# Proof object clause steps            : 84
# Proof object formula steps           : 33
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 50
# Proof object simplifying inferences  : 60
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 21
# Processed clauses                    : 34525
# ...of these trivial                  : 336
# ...subsumed                          : 32380
# ...remaining for further processing  : 1809
# Other redundant clauses eliminated   : 11832
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 177
# Backward-rewritten                   : 140
# Generated clauses                    : 763758
# ...of the previous two non-trivial   : 514894
# Contextual simplify-reflections      : 21
# Paramodulations                      : 751926
# Factorizations                       : 0
# Equation resolutions                 : 11832
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1470
#    Positive orientable unit clauses  : 50
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 1414
# Current number of unprocessed clauses: 477831
# ...number of literals in the above   : 2315033
# Current number of archived formulas  : 0
# Current number of archived clauses   : 339
# Clause-clause subsumption calls (NU) : 323887
# Rec. Clause-clause subsumption calls : 215629
# Non-unit clause-clause subsumptions  : 20197
# Unit Clause-clause subsumption calls : 609
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 447
# BW rewrite match successes           : 57
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 18018775

# -------------------------------------------------
# User time                : 7.069 s
# System time              : 0.153 s
# Total time               : 7.221 s
# Maximum resident set size: 358816 pages
