# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t42_wellord1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&(r2_hidden(X2,k3_relat_1(X3))&![X4]:(r2_hidden(X4,k1_wellord1(X3,X1))=>(r2_hidden(k4_tarski(X4,X2),X3)&~(X4=X2))))))=>r2_hidden(k4_tarski(X1,X2),X3))), file('data/MPTP/proofs/af7503fadb294f2c85d060e3fc713660.E_input', t42_wellord1)).
fof(d4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v2_wellord1(X1)<=>(v1_relat_2(X1)&(v8_relat_2(X1)&(v4_relat_2(X1)&(v6_relat_2(X1)&v1_wellord1(X1))))))), file('data/MPTP/proofs/af7503fadb294f2c85d060e3fc713660.E_input', d4_wellord1)).
fof(l4_wellord1, axiom, ![X1]:(v1_relat_1(X1)=>(v6_relat_2(X1)<=>![X2, X3]:~((r2_hidden(X2,k3_relat_1(X1))&(r2_hidden(X3,k3_relat_1(X1))&(~(X2=X3)&(~(r2_hidden(k4_tarski(X2,X3),X1))&~(r2_hidden(k4_tarski(X3,X2),X1))))))))), file('data/MPTP/proofs/af7503fadb294f2c85d060e3fc713660.E_input', l4_wellord1)).
fof(t38_wellord1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))<=>(X1=X2|r2_hidden(X1,k1_wellord1(X3,X2)))))), file('data/MPTP/proofs/af7503fadb294f2c85d060e3fc713660.E_input', t38_wellord1)).
fof(t37_wellord1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))))), file('data/MPTP/proofs/af7503fadb294f2c85d060e3fc713660.E_input', t37_wellord1)).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>((v2_wellord1(X3)&(r2_hidden(X1,k3_relat_1(X3))&(r2_hidden(X2,k3_relat_1(X3))&![X4]:(r2_hidden(X4,k1_wellord1(X3,X1))=>(r2_hidden(k4_tarski(X4,X2),X3)&~(X4=X2))))))=>r2_hidden(k4_tarski(X1,X2),X3)))), inference(assume_negation,[status(cth)],[t42_wellord1])).
fof(c_0_6, plain, ![X9]:(((v1_relat_2(X9)|~v2_wellord1(X9)|~v1_relat_1(X9))&((v8_relat_2(X9)|~v2_wellord1(X9)|~v1_relat_1(X9))&((v4_relat_2(X9)|~v2_wellord1(X9)|~v1_relat_1(X9))&((v6_relat_2(X9)|~v2_wellord1(X9)|~v1_relat_1(X9))&(v1_wellord1(X9)|~v2_wellord1(X9)|~v1_relat_1(X9))))))&(~v1_relat_2(X9)|(~v8_relat_2(X9)|(~v4_relat_2(X9)|(~v6_relat_2(X9)|~v1_wellord1(X9))))|v2_wellord1(X9)|~v1_relat_1(X9))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_wellord1])])])).
fof(c_0_7, negated_conjecture, ![X8]:(v1_relat_1(esk3_0)&((v2_wellord1(esk3_0)&(r2_hidden(esk1_0,k3_relat_1(esk3_0))&(r2_hidden(esk2_0,k3_relat_1(esk3_0))&((r2_hidden(k4_tarski(X8,esk2_0),esk3_0)|~r2_hidden(X8,k1_wellord1(esk3_0,esk1_0)))&(X8!=esk2_0|~r2_hidden(X8,k1_wellord1(esk3_0,esk1_0)))))))&~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])])])).
fof(c_0_8, plain, ![X10, X11, X12]:((~v6_relat_2(X10)|(~r2_hidden(X11,k3_relat_1(X10))|(~r2_hidden(X12,k3_relat_1(X10))|(X11=X12|(r2_hidden(k4_tarski(X11,X12),X10)|r2_hidden(k4_tarski(X12,X11),X10)))))|~v1_relat_1(X10))&((r2_hidden(esk4_1(X10),k3_relat_1(X10))|v6_relat_2(X10)|~v1_relat_1(X10))&((r2_hidden(esk5_1(X10),k3_relat_1(X10))|v6_relat_2(X10)|~v1_relat_1(X10))&((esk4_1(X10)!=esk5_1(X10)|v6_relat_2(X10)|~v1_relat_1(X10))&((~r2_hidden(k4_tarski(esk4_1(X10),esk5_1(X10)),X10)|v6_relat_2(X10)|~v1_relat_1(X10))&(~r2_hidden(k4_tarski(esk5_1(X10),esk4_1(X10)),X10)|v6_relat_2(X10)|~v1_relat_1(X10))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l4_wellord1])])])])])])])).
cnf(c_0_9, plain, (v6_relat_2(X1)|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (v2_wellord1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_12, plain, ![X15, X16, X17]:((~r1_tarski(k1_wellord1(X17,X15),k1_wellord1(X17,X16))|(X15=X16|r2_hidden(X15,k1_wellord1(X17,X16)))|(~v2_wellord1(X17)|(~r2_hidden(X15,k3_relat_1(X17))|~r2_hidden(X16,k3_relat_1(X17))))|~v1_relat_1(X17))&((X15!=X16|r1_tarski(k1_wellord1(X17,X15),k1_wellord1(X17,X16))|(~v2_wellord1(X17)|(~r2_hidden(X15,k3_relat_1(X17))|~r2_hidden(X16,k3_relat_1(X17))))|~v1_relat_1(X17))&(~r2_hidden(X15,k1_wellord1(X17,X16))|r1_tarski(k1_wellord1(X17,X15),k1_wellord1(X17,X16))|(~v2_wellord1(X17)|(~r2_hidden(X15,k3_relat_1(X17))|~r2_hidden(X16,k3_relat_1(X17))))|~v1_relat_1(X17)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_wellord1])])])).
fof(c_0_13, plain, ![X18, X19, X20]:((~r2_hidden(k4_tarski(X18,X19),X20)|r1_tarski(k1_wellord1(X20,X18),k1_wellord1(X20,X19))|(~v2_wellord1(X20)|(~r2_hidden(X18,k3_relat_1(X20))|~r2_hidden(X19,k3_relat_1(X20))))|~v1_relat_1(X20))&(~r1_tarski(k1_wellord1(X20,X18),k1_wellord1(X20,X19))|r2_hidden(k4_tarski(X18,X19),X20)|(~v2_wellord1(X20)|(~r2_hidden(X18,k3_relat_1(X20))|~r2_hidden(X19,k3_relat_1(X20))))|~v1_relat_1(X20))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t37_wellord1])])])).
cnf(c_0_14, plain, (X2=X3|r2_hidden(k4_tarski(X2,X3),X1)|r2_hidden(k4_tarski(X3,X2),X1)|~v6_relat_2(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, negated_conjecture, (r2_hidden(esk2_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, negated_conjecture, (v6_relat_2(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_10]), c_0_11])])).
cnf(c_0_17, plain, (X2=X3|r2_hidden(X2,k1_wellord1(X1,X3))|~r1_tarski(k1_wellord1(X1,X2),k1_wellord1(X1,X3))|~v2_wellord1(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~v2_wellord1(X3)|~r2_hidden(X1,k3_relat_1(X3))|~r2_hidden(X2,k3_relat_1(X3))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, negated_conjecture, (X1=esk2_0|r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|r2_hidden(k4_tarski(esk2_0,X1),esk3_0)|~r2_hidden(X1,k3_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_11])]), c_0_16])])).
cnf(c_0_20, negated_conjecture, (r2_hidden(esk1_0,k3_relat_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, plain, (X1=X2|r2_hidden(X1,k1_wellord1(X3,X2))|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,k3_relat_1(X3))|~r2_hidden(X1,k3_relat_1(X3))|~v2_wellord1(X3)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, negated_conjecture, (esk2_0=esk1_0|r2_hidden(k4_tarski(esk2_0,esk1_0),esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_24, negated_conjecture, (X1!=esk2_0|~r2_hidden(X1,k1_wellord1(esk3_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,k1_wellord1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_20]), c_0_15]), c_0_10]), c_0_11])])).
cnf(c_0_26, plain, (r1_tarski(k1_wellord1(X3,X1),k1_wellord1(X3,X2))|X1!=X2|~v2_wellord1(X3)|~r2_hidden(X1,k3_relat_1(X3))|~r2_hidden(X2,k3_relat_1(X3))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_27, negated_conjecture, (esk2_0=esk1_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(X2,X3),X1)|~r1_tarski(k1_wellord1(X1,X2),k1_wellord1(X1,X3))|~v2_wellord1(X1)|~r2_hidden(X2,k3_relat_1(X1))|~r2_hidden(X3,k3_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_29, plain, (r1_tarski(k1_wellord1(X1,X2),k1_wellord1(X1,X2))|~r2_hidden(X2,k3_relat_1(X1))|~v2_wellord1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_30, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk1_0),esk3_0)), inference(rw,[status(thm)],[c_0_21, c_0_27])).
cnf(c_0_31, plain, (r2_hidden(k4_tarski(X1,X1),X2)|~r2_hidden(X1,k3_relat_1(X2))|~v2_wellord1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_32, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_20]), c_0_10]), c_0_11])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 22
# Proof object formula steps           : 11
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 75
# ...of these trivial                  : 1
# ...subsumed                          : 4
# ...remaining for further processing  : 70
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 11
# Generated clauses                    : 57
# ...of the previous two non-trivial   : 50
# Contextual simplify-reflections      : 0
# Paramodulations                      : 56
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 34
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 22
# ...number of literals in the above   : 155
# Current number of archived formulas  : 0
# Current number of archived clauses   : 35
# Clause-clause subsumption calls (NU) : 634
# Rec. Clause-clause subsumption calls : 73
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 31
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3621

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28540 pages
