# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t42_enumset1, axiom, ![X1, X2, X3]:k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_tarski(X2,X3)), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t42_enumset1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t69_enumset1)).
fof(t46_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4)), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t46_enumset1)).
fof(t58_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_enumset1(X4,X5,X6,X7)), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t58_enumset1)).
fof(l68_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_enumset1(X5,X6,X7)), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', l68_enumset1)).
fof(t41_enumset1, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2)), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t41_enumset1)).
fof(t75_enumset1, conjecture, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t75_enumset1)).
fof(t56_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_tarski(X1),k4_enumset1(X2,X3,X4,X5,X6,X7)), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t56_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t74_enumset1)).
fof(t62_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X3,X4,X5,X6,X7,X8)), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', t62_enumset1)).
fof(l75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8)), file('data/MPTP/proofs/afb6991715494cff8d0675470d232d53.E_input', l75_enumset1)).
fof(c_0_11, plain, ![X9, X10, X11]:k1_enumset1(X9,X10,X11)=k2_xboole_0(k1_tarski(X9),k2_tarski(X10,X11)), inference(variable_rename,[status(thm)],[t42_enumset1])).
fof(c_0_12, plain, ![X12]:k2_tarski(X12,X12)=k1_tarski(X12), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_13, plain, ![X13, X14, X15, X16]:k2_enumset1(X13,X14,X15,X16)=k2_xboole_0(k1_enumset1(X13,X14,X15),k1_tarski(X16)), inference(variable_rename,[status(thm)],[t46_enumset1])).
cnf(c_0_14, plain, (k1_enumset1(X1,X2,X3)=k2_xboole_0(k1_tarski(X1),k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X17, X18, X19, X20, X21, X22, X23]:k5_enumset1(X17,X18,X19,X20,X21,X22,X23)=k2_xboole_0(k1_enumset1(X17,X18,X19),k2_enumset1(X20,X21,X22,X23)), inference(variable_rename,[status(thm)],[t58_enumset1])).
cnf(c_0_17, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X4))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (k1_enumset1(X1,X2,X3)=k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3))), inference(rw,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_19, plain, ![X24, X25, X26, X27, X28, X29, X30]:k5_enumset1(X24,X25,X26,X27,X28,X29,X30)=k2_xboole_0(k2_enumset1(X24,X25,X26,X27),k1_enumset1(X28,X29,X30)), inference(variable_rename,[status(thm)],[l68_enumset1])).
cnf(c_0_20, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_enumset1(X1,X2,X3),k2_enumset1(X4,X5,X6,X7))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (k2_enumset1(X1,X2,X3,X4)=k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_tarski(X4,X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_15]), c_0_18])).
fof(c_0_22, plain, ![X31, X32]:k2_tarski(X31,X32)=k2_xboole_0(k1_tarski(X31),k1_tarski(X32)), inference(variable_rename,[status(thm)],[t41_enumset1])).
cnf(c_0_23, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k1_enumset1(X5,X6,X7))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_xboole_0(k2_xboole_0(k2_tarski(X4,X4),k2_tarski(X5,X6)),k2_tarski(X7,X7)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_18]), c_0_21])).
cnf(c_0_25, plain, (k2_tarski(X1,X2)=k2_xboole_0(k1_tarski(X1),k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_xboole_0(k2_xboole_0(k2_tarski(X4,X4),k2_tarski(X5,X6)),k2_tarski(X7,X7)))=k2_xboole_0(k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_tarski(X4,X4)),k2_xboole_0(k2_tarski(X5,X5),k2_tarski(X6,X7)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_18]), c_0_21]), c_0_24])).
cnf(c_0_27, plain, (k2_tarski(X1,X2)=k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_15]), c_0_15])).
cnf(c_0_28, plain, (k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_xboole_0(k2_xboole_0(k2_tarski(X4,X4),k2_tarski(X5,X6)),k2_tarski(X6,X6)))=k2_xboole_0(k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_tarski(X4,X4)),k2_tarski(X5,X6))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_29, plain, (k2_xboole_0(k2_tarski(X1,X2),k2_xboole_0(k2_xboole_0(k2_tarski(X3,X3),k2_tarski(X4,X5)),k2_tarski(X5,X5)))=k2_xboole_0(k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X3)),k2_tarski(X4,X5))), inference(spm,[status(thm)],[c_0_28, c_0_27])).
cnf(c_0_30, plain, (k2_xboole_0(k2_tarski(X1,X2),k2_xboole_0(k2_tarski(X3,X4),k2_tarski(X4,X4)))=k2_xboole_0(k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X3)),k2_tarski(X4,X4))), inference(spm,[status(thm)],[c_0_29, c_0_27])).
fof(c_0_31, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(assume_negation,[status(cth)],[t75_enumset1])).
fof(c_0_32, plain, ![X33, X34, X35, X36, X37, X38, X39]:k5_enumset1(X33,X34,X35,X36,X37,X38,X39)=k2_xboole_0(k1_tarski(X33),k4_enumset1(X34,X35,X36,X37,X38,X39)), inference(variable_rename,[status(thm)],[t56_enumset1])).
fof(c_0_33, plain, ![X40, X41, X42, X43, X44, X45]:k5_enumset1(X40,X40,X41,X42,X43,X44,X45)=k4_enumset1(X40,X41,X42,X43,X44,X45), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_34, plain, ![X46, X47, X48, X49, X50, X51, X52, X53]:k6_enumset1(X46,X47,X48,X49,X50,X51,X52,X53)=k2_xboole_0(k1_tarski(X46),k5_enumset1(X47,X48,X49,X50,X51,X52,X53)), inference(variable_rename,[status(thm)],[t62_enumset1])).
fof(c_0_35, plain, ![X61, X62, X63, X64, X65, X66, X67, X68]:k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68)=k2_xboole_0(k2_enumset1(X61,X62,X63,X64),k2_enumset1(X65,X66,X67,X68)), inference(variable_rename,[status(thm)],[l75_enumset1])).
cnf(c_0_36, plain, (k2_xboole_0(k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X3)),k2_tarski(X4,X4))=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_27]), c_0_27])).
fof(c_0_37, negated_conjecture, k6_enumset1(esk1_0,esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0)!=k5_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])).
cnf(c_0_38, plain, (k5_enumset1(X1,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k1_tarski(X1),k4_enumset1(X2,X3,X4,X5,X6,X7))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k1_tarski(X1),k5_enumset1(X2,X3,X4,X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X3))=k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3))), inference(spm,[status(thm)],[c_0_36, c_0_27])).
cnf(c_0_43, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0)!=k5_enumset1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, plain, (k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_xboole_0(k2_xboole_0(k2_tarski(X4,X4),k2_tarski(X5,X6)),k2_tarski(X7,X7)))=k2_xboole_0(k2_tarski(X1,X1),k2_xboole_0(k2_xboole_0(k2_tarski(X2,X2),k2_tarski(X2,X3)),k2_xboole_0(k2_xboole_0(k2_tarski(X4,X4),k2_tarski(X5,X6)),k2_tarski(X7,X7))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_15]), c_0_39]), c_0_24]), c_0_24])).
cnf(c_0_45, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_tarski(X1,X1),k2_xboole_0(k2_xboole_0(k2_tarski(X2,X2),k2_tarski(X3,X4)),k2_xboole_0(k2_xboole_0(k2_tarski(X5,X5),k2_tarski(X6,X7)),k2_tarski(X8,X8))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_15]), c_0_24])).
cnf(c_0_46, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_tarski(X4,X4)),k2_xboole_0(k2_xboole_0(k2_tarski(X5,X5),k2_tarski(X6,X7)),k2_tarski(X8,X8)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_21]), c_0_21])).
cnf(c_0_47, plain, (k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_tarski(X4,X4))=k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4))), inference(spm,[status(thm)],[c_0_36, c_0_42])).
cnf(c_0_48, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0)!=k2_xboole_0(k2_xboole_0(k2_tarski(esk1_0,esk1_0),k2_tarski(esk2_0,esk3_0)),k2_xboole_0(k2_xboole_0(k2_tarski(esk4_0,esk4_0),k2_tarski(esk5_0,esk6_0)),k2_tarski(esk7_0,esk7_0)))), inference(rw,[status(thm)],[c_0_43, c_0_24])).
cnf(c_0_49, plain, (k6_enumset1(X1,X2,X2,X3,X4,X5,X6,X7)=k2_xboole_0(k2_xboole_0(k2_tarski(X1,X1),k2_tarski(X2,X3)),k2_xboole_0(k2_xboole_0(k2_tarski(X4,X4),k2_tarski(X5,X6)),k2_tarski(X7,X7)))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, plain, (k2_xboole_0(k2_xboole_0(k2_tarski(X1,X2),k2_tarski(X3,X4)),k2_xboole_0(k2_tarski(X5,X6),k2_tarski(X7,X8)))=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_47])).
cnf(c_0_51, negated_conjecture, (k6_enumset1(esk1_0,esk2_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0)!=k6_enumset1(esk1_0,esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0,esk7_0)), inference(rw,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_52, plain, (k6_enumset1(X1,X2,X2,X3,X4,X5,X6,X7)=k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_47]), c_0_50])).
cnf(c_0_53, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 31
# Proof object formula steps           : 23
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 6
# Processed clauses                    : 94
# ...of these trivial                  : 11
# ...subsumed                          : 47
# ...remaining for further processing  : 36
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 11
# Generated clauses                    : 329
# ...of the previous two non-trivial   : 262
# Contextual simplify-reflections      : 0
# Paramodulations                      : 329
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 19
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 6
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 175
# ...number of literals in the above   : 175
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 34
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 154
# BW rewrite match successes           : 45
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5121

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28788 pages
