# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', d5_xboole_0)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', t100_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', l32_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', t36_xboole_1)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', t44_xboole_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', t70_enumset1)).
fof(t38_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', t38_zfmisc_1)).
fof(t74_zfmisc_1, conjecture, ![X1, X2, X3]:~((~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0)&(~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1))&(~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X2))&~(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2)))))), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', t74_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/afc67726b9ee441d8428e13a698410e0.E_input', t69_enumset1)).
fof(c_0_10, plain, ![X11, X12, X13, X14, X15, X16, X17, X18]:((((r2_hidden(X14,X11)|~r2_hidden(X14,X13)|X13!=k4_xboole_0(X11,X12))&(~r2_hidden(X14,X12)|~r2_hidden(X14,X13)|X13!=k4_xboole_0(X11,X12)))&(~r2_hidden(X15,X11)|r2_hidden(X15,X12)|r2_hidden(X15,X13)|X13!=k4_xboole_0(X11,X12)))&((~r2_hidden(esk4_3(X16,X17,X18),X18)|(~r2_hidden(esk4_3(X16,X17,X18),X16)|r2_hidden(esk4_3(X16,X17,X18),X17))|X18=k4_xboole_0(X16,X17))&((r2_hidden(esk4_3(X16,X17,X18),X16)|r2_hidden(esk4_3(X16,X17,X18),X18)|X18=k4_xboole_0(X16,X17))&(~r2_hidden(esk4_3(X16,X17,X18),X17)|r2_hidden(esk4_3(X16,X17,X18),X18)|X18=k4_xboole_0(X16,X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_11, plain, ![X21, X22]:k4_xboole_0(X21,X22)=k5_xboole_0(X21,k3_xboole_0(X21,X22)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_12, plain, ![X30, X31]:((k4_xboole_0(X30,X31)!=k1_xboole_0|r1_tarski(X30,X31))&(~r1_tarski(X30,X31)|k4_xboole_0(X30,X31)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_13, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_15, c_0_14])).
cnf(c_0_18, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)|~r2_hidden(X2,X3)|~r1_tarski(X4,X3)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_19, plain, (r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_20, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)|~r1_tarski(X3,X2)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (X3=k5_xboole_0(X1,k3_xboole_0(X1,X2))|r2_hidden(esk4_3(X1,X2,X3),X3)|r2_hidden(esk4_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_19, c_0_14])).
fof(c_0_22, plain, ![X23, X24]:r1_tarski(k4_xboole_0(X23,X24),X23), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_23, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|r2_hidden(esk4_3(X1,X2,k1_xboole_0),X1)|~r2_hidden(esk4_3(X1,X2,k1_xboole_0),X3)|~r1_tarski(X4,X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_25, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|r2_hidden(esk4_3(X1,X2,k1_xboole_0),X1)|~r1_tarski(X3,k1_xboole_0)), inference(spm,[status(thm)],[c_0_23, c_0_21])).
cnf(c_0_26, plain, (r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_24, c_0_14])).
cnf(c_0_27, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|r2_hidden(esk4_3(X1,X2,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_28, plain, (k5_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,X1))=k1_xboole_0|~r2_hidden(esk4_3(k1_xboole_0,X1,k1_xboole_0),X2)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_20, c_0_27])).
fof(c_0_29, plain, ![X27, X28, X29]:(~r1_tarski(k4_xboole_0(X27,X28),X29)|r1_tarski(X27,k2_xboole_0(X28,X29))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
fof(c_0_30, plain, ![X32, X33, X34, X35, X36, X37, X38, X39]:(((~r2_hidden(X35,X34)|(X35=X32|X35=X33)|X34!=k2_tarski(X32,X33))&((X36!=X32|r2_hidden(X36,X34)|X34!=k2_tarski(X32,X33))&(X36!=X33|r2_hidden(X36,X34)|X34!=k2_tarski(X32,X33))))&(((esk5_3(X37,X38,X39)!=X37|~r2_hidden(esk5_3(X37,X38,X39),X39)|X39=k2_tarski(X37,X38))&(esk5_3(X37,X38,X39)!=X38|~r2_hidden(esk5_3(X37,X38,X39),X39)|X39=k2_tarski(X37,X38)))&(r2_hidden(esk5_3(X37,X38,X39),X39)|(esk5_3(X37,X38,X39)=X37|esk5_3(X37,X38,X39)=X38)|X39=k2_tarski(X37,X38)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_31, plain, ![X25, X26]:k1_enumset1(X25,X25,X26)=k2_tarski(X25,X26), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_32, plain, (k5_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,X1))=k1_xboole_0|~r1_tarski(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_27])).
cnf(c_0_33, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_34, plain, ![X8, X9, X10]:(((r2_hidden(X8,X10)|~r1_tarski(k2_tarski(X8,X9),X10))&(r2_hidden(X9,X10)|~r1_tarski(k2_tarski(X8,X9),X10)))&(~r2_hidden(X8,X10)|~r2_hidden(X9,X10)|r1_tarski(k2_tarski(X8,X9),X10))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_zfmisc_1])])])).
cnf(c_0_35, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_36, negated_conjecture, ~(![X1, X2, X3]:~((~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_xboole_0)&(~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X1))&(~(k4_xboole_0(k2_tarski(X1,X2),X3)=k1_tarski(X2))&~(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2))))))), inference(assume_negation,[status(cth)],[t74_zfmisc_1])).
cnf(c_0_37, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_40, plain, (k5_xboole_0(k1_xboole_0,k3_xboole_0(k1_xboole_0,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_32, c_0_26])).
cnf(c_0_41, plain, (X3=k5_xboole_0(X1,k3_xboole_0(X1,X2))|r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_33, c_0_14])).
cnf(c_0_42, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k5_xboole_0(X1,k3_xboole_0(X1,X2)),X3)), inference(rw,[status(thm)],[c_0_35, c_0_14])).
fof(c_0_44, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_xboole_0&(k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_tarski(esk1_0)&(k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_tarski(esk2_0)&k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k2_tarski(esk1_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_36])])])])).
fof(c_0_45, plain, ![X20]:k2_tarski(X20,X20)=k1_tarski(X20), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_46, plain, (X1=X4|X1=X3|X2!=k1_enumset1(X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_48, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X1))=k1_xboole_0|r2_hidden(esk4_3(X1,X1,k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_41, c_0_27])).
cnf(c_0_49, plain, (r2_hidden(X1,X2)|~r1_tarski(k1_enumset1(X1,X1,X3),X2)), inference(rw,[status(thm)],[c_0_42, c_0_38])).
cnf(c_0_50, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_43, c_0_26])).
cnf(c_0_51, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_53, plain, (X1=X2|X1=X3|~r2_hidden(X1,k1_enumset1(X2,X2,X3))), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_54, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k2_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_56, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_57, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X1))=k1_xboole_0|~r2_hidden(esk4_3(X1,X1,k1_xboole_0),X2)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_58, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_enumset1(X1,X1,X3)))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59, negated_conjecture, (k5_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),k3_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0))!=k1_enumset1(esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52]), c_0_38]), c_0_38]), c_0_14])).
cnf(c_0_60, plain, (X1=k5_xboole_0(k1_enumset1(X2,X2,X3),k3_xboole_0(k1_enumset1(X2,X2,X3),X4))|esk4_3(k1_enumset1(X2,X2,X3),X4,X1)=X3|esk4_3(k1_enumset1(X2,X2,X3),X4,X1)=X2|r2_hidden(esk4_3(k1_enumset1(X2,X2,X3),X4,X1),X1)), inference(spm,[status(thm)],[c_0_53, c_0_21])).
cnf(c_0_61, negated_conjecture, (k5_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),k3_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0))!=k1_enumset1(esk1_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_38]), c_0_38]), c_0_14])).
cnf(c_0_62, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_55, c_0_38])).
cnf(c_0_63, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_64, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_65, plain, (r2_hidden(X1,X2)|X3!=k5_xboole_0(X2,k3_xboole_0(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_56, c_0_14])).
cnf(c_0_66, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_67, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_16, c_0_40])).
cnf(c_0_68, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1)=esk1_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1)=esk2_0|r2_hidden(esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1),X1)|X1!=k1_enumset1(esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_69, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_70, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1)=esk1_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1)=esk2_0|X1!=k1_enumset1(esk1_0,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_60]), c_0_46])).
cnf(c_0_71, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X3,X1)), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_72, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X2,X4)), inference(rw,[status(thm)],[c_0_63, c_0_38])).
cnf(c_0_73, negated_conjecture, (k5_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),k3_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_38]), c_0_14])).
cnf(c_0_74, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67])).
cnf(c_0_75, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=esk2_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=esk1_0|r2_hidden(esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk1_0)),k1_enumset1(esk1_0,esk1_0,esk1_0))), inference(er,[status(thm)],[c_0_68])).
cnf(c_0_76, plain, (X3=k5_xboole_0(X1,k3_xboole_0(X1,X2))|r2_hidden(esk4_3(X1,X2,X3),X2)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_69, c_0_14])).
cnf(c_0_77, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk2_0))=esk2_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk2_0))=esk1_0), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_78, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_79, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X1,X1,X3)), inference(er,[status(thm)],[c_0_72])).
cnf(c_0_80, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1)=esk1_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1)=esk2_0|X1!=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_60]), c_0_74])).
cnf(c_0_81, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=esk2_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=esk1_0), inference(spm,[status(thm)],[c_0_53, c_0_75])).
cnf(c_0_82, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk2_0))=esk1_0|r2_hidden(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])]), c_0_61])).
cnf(c_0_83, plain, (r2_hidden(X1,k1_enumset1(X1,X1,X2))), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_84, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_xboole_0)=esk2_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_xboole_0)=esk1_0), inference(er,[status(thm)],[c_0_80])).
cnf(c_0_85, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))), inference(er,[status(thm)],[c_0_65])).
cnf(c_0_86, negated_conjecture, (k4_xboole_0(k2_tarski(esk1_0,esk2_0),esk3_0)!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_87, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=esk1_0|r2_hidden(esk2_0,k1_enumset1(esk1_0,esk1_0,esk1_0))|~r2_hidden(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_81]), c_0_59])).
cnf(c_0_88, negated_conjecture, (r2_hidden(esk2_0,esk3_0)|r2_hidden(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_82]), c_0_83])]), c_0_61])).
cnf(c_0_89, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_xboole_0)=esk1_0|esk2_0!=esk1_0), inference(ef,[status(thm)],[c_0_84])).
cnf(c_0_90, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_66]), c_0_47])).
cnf(c_0_91, negated_conjecture, (k5_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),k3_xboole_0(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0))!=k1_enumset1(esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_52]), c_0_38]), c_0_38]), c_0_14])).
cnf(c_0_92, plain, (r1_tarski(k2_tarski(X1,X3),X2)|~r2_hidden(X1,X2)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_93, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=esk1_0|r2_hidden(esk2_0,k1_enumset1(esk1_0,esk1_0,esk1_0))|r2_hidden(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_94, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk1_0,esk1_0,esk1_0))=esk1_0|esk2_0!=esk1_0), inference(ef,[status(thm)],[c_0_81])).
cnf(c_0_95, negated_conjecture, (esk2_0!=esk1_0|~r2_hidden(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_89]), c_0_73]), c_0_90])).
cnf(c_0_96, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1)=esk1_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1)=esk2_0|r2_hidden(esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,X1),X1)|X1!=k1_enumset1(esk2_0,esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_91, c_0_60])).
cnf(c_0_97, negated_conjecture, (~r1_tarski(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_73, c_0_17])).
cnf(c_0_98, plain, (r1_tarski(k1_enumset1(X1,X1,X3),X2)|~r2_hidden(X3,X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_92, c_0_38])).
cnf(c_0_99, negated_conjecture, (r2_hidden(esk2_0,k1_enumset1(esk1_0,esk1_0,esk1_0))|r2_hidden(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_93]), c_0_83]), c_0_83])]), c_0_59])).
cnf(c_0_100, negated_conjecture, (esk2_0!=esk1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_94]), c_0_83]), c_0_83])]), c_0_59]), c_0_95])).
cnf(c_0_101, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk2_0,esk2_0,esk2_0))=esk2_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk2_0,esk2_0,esk2_0))=esk1_0|r2_hidden(esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk2_0,esk2_0,esk2_0)),k1_enumset1(esk2_0,esk2_0,esk2_0))), inference(er,[status(thm)],[c_0_96])).
cnf(c_0_102, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)|~r2_hidden(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_103, negated_conjecture, (r2_hidden(esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_99]), c_0_100])).
cnf(c_0_104, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk2_0,esk2_0,esk2_0))=esk1_0|esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk2_0,esk2_0,esk2_0))=esk2_0), inference(spm,[status(thm)],[c_0_53, c_0_101])).
cnf(c_0_105, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_106, negated_conjecture, (esk4_3(k1_enumset1(esk1_0,esk1_0,esk2_0),esk3_0,k1_enumset1(esk2_0,esk2_0,esk2_0))=esk1_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_104]), c_0_83]), c_0_78])]), c_0_91]), c_0_105])).
cnf(c_0_107, negated_conjecture, (r2_hidden(esk1_0,k1_enumset1(esk2_0,esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_106]), c_0_103])]), c_0_91])).
cnf(c_0_108, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_107]), c_0_100]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 109
# Proof object clause steps            : 88
# Proof object formula steps           : 21
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 48
# Proof object simplifying inferences  : 62
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 23
# Processed clauses                    : 399
# ...of these trivial                  : 14
# ...subsumed                          : 168
# ...remaining for further processing  : 217
# Other redundant clauses eliminated   : 52
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 12
# Generated clauses                    : 2785
# ...of the previous two non-trivial   : 2617
# Contextual simplify-reflections      : 9
# Paramodulations                      : 2691
# Factorizations                       : 24
# Equation resolutions                 : 69
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 164
#    Positive orientable unit clauses  : 33
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 21
#    Non-unit-clauses                  : 110
# Current number of unprocessed clauses: 2199
# ...number of literals in the above   : 11744
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 3829
# Rec. Clause-clause subsumption calls : 2483
# Non-unit clause-clause subsumptions  : 120
# Unit Clause-clause subsumption calls : 1458
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 69
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 49512

# -------------------------------------------------
# User time                : 0.046 s
# System time              : 0.003 s
# Total time               : 0.049 s
# Maximum resident set size: 28788 pages
