# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t41_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1))))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t41_ordinal1)).
fof(t24_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1))))))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t24_ordinal1)).
fof(s1_xboole_0__e3_53__ordinal1, axiom, ![X1]:?[X2]:![X3]:(r2_hidden(X3,X2)<=>(r2_hidden(X3,X1)&v3_ordinal1(X3))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', s1_xboole_0__e3_53__ordinal1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', d3_tarski)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', redefinition_r1_ordinal1)).
fof(t39_ordinal1, axiom, ![X1]:?[X2]:(v3_ordinal1(X2)&(~(r2_hidden(X2,X1))&![X3]:(v3_ordinal1(X3)=>(~(r2_hidden(X3,X1))=>r1_ordinal1(X2,X3))))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t39_ordinal1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t7_ordinal1)).
fof(t34_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2)))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t34_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', d1_ordinal1)).
fof(t13_ordinal1, axiom, ![X1, X2]:(r2_hidden(X1,k1_ordinal1(X2))<=>(r2_hidden(X1,X2)|X1=X2)), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t13_ordinal1)).
fof(t38_ordinal1, axiom, ![X1]:~(![X2]:(v3_ordinal1(X2)=>r2_hidden(X2,X1))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t38_ordinal1)).
fof(t23_ordinal1, axiom, ![X1, X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)=>v3_ordinal1(X1))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t23_ordinal1)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', d4_tarski)).
fof(t35_ordinal1, axiom, ![X1]:(![X2]:(r2_hidden(X2,X1)=>v3_ordinal1(X2))=>v3_ordinal1(k3_tarski(X1))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t35_ordinal1)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t10_ordinal1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', d10_xboole_0)).
fof(d6_ordinal1, axiom, ![X1]:(v4_ordinal1(X1)<=>X1=k3_tarski(X1)), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', d6_ordinal1)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/b0b4fa4322f94e9eb41f40ae944728eb.E_input', t29_ordinal1)).
fof(c_0_18, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1)))))), inference(assume_negation,[status(cth)],[t41_ordinal1])).
fof(c_0_19, plain, ![X8, X9]:(~v3_ordinal1(X8)|(~v3_ordinal1(X9)|(r2_hidden(X8,X9)|(X8=X9|r2_hidden(X9,X8))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])).
fof(c_0_20, negated_conjecture, ![X7]:(v3_ordinal1(esk1_0)&(((v3_ordinal1(esk2_0)|~v4_ordinal1(esk1_0))&((r2_hidden(esk2_0,esk1_0)|~v4_ordinal1(esk1_0))&(~r2_hidden(k1_ordinal1(esk2_0),esk1_0)|~v4_ordinal1(esk1_0))))&(v4_ordinal1(esk1_0)|(~v3_ordinal1(X7)|(~r2_hidden(X7,esk1_0)|r2_hidden(k1_ordinal1(X7),esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])])).
fof(c_0_21, plain, ![X15, X17, X18]:(((r2_hidden(X17,X15)|~r2_hidden(X17,esk4_1(X15)))&(v3_ordinal1(X17)|~r2_hidden(X17,esk4_1(X15))))&(~r2_hidden(X18,X15)|~v3_ordinal1(X18)|r2_hidden(X18,esk4_1(X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[s1_xboole_0__e3_53__ordinal1])])])])])])).
fof(c_0_22, plain, ![X19, X20, X21, X22, X23]:((~r1_tarski(X19,X20)|(~r2_hidden(X21,X19)|r2_hidden(X21,X20)))&((r2_hidden(esk5_2(X22,X23),X22)|r1_tarski(X22,X23))&(~r2_hidden(esk5_2(X22,X23),X23)|r1_tarski(X22,X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_23, plain, ![X10, X11]:((~r1_ordinal1(X10,X11)|r1_tarski(X10,X11)|(~v3_ordinal1(X10)|~v3_ordinal1(X11)))&(~r1_tarski(X10,X11)|r1_ordinal1(X10,X11)|(~v3_ordinal1(X10)|~v3_ordinal1(X11)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
fof(c_0_24, plain, ![X12, X14]:(v3_ordinal1(esk3_1(X12))&(~r2_hidden(esk3_1(X12),X12)&(~v3_ordinal1(X14)|(r2_hidden(X14,X12)|r1_ordinal1(esk3_1(X12),X14))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_ordinal1])])])])])).
cnf(c_0_25, plain, (r2_hidden(X1,X2)|X1=X2|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk4_1(X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X25, X26]:(~r2_hidden(X25,X26)|~r1_tarski(X26,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_30, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (r2_hidden(X1,X2)|r1_ordinal1(esk3_1(X2),X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (v3_ordinal1(esk3_1(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (~r2_hidden(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (r2_hidden(X1,esk4_1(X2))|~r2_hidden(X1,X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_35, negated_conjecture, (X1=esk1_0|r2_hidden(X1,esk1_0)|r2_hidden(esk1_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_36, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_37, plain, (r1_tarski(esk4_1(X1),X2)|r2_hidden(esk5_2(esk4_1(X1),X2),X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_38, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (r1_tarski(esk3_1(X1),X2)|r2_hidden(X2,X1)|~v3_ordinal1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])])).
cnf(c_0_40, plain, (~r2_hidden(esk3_1(esk4_1(X1)),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_32])])).
cnf(c_0_41, negated_conjecture, (esk3_1(X1)=esk1_0|r2_hidden(esk1_0,esk3_1(X1))|r2_hidden(esk3_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_32])).
cnf(c_0_42, plain, (r1_tarski(esk4_1(X1),X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk3_1(X2))|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44, negated_conjecture, (esk3_1(esk4_1(esk1_0))=esk1_0|r2_hidden(esk1_0,esk3_1(esk4_1(esk1_0)))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45, plain, (~r2_hidden(X1,esk4_1(X1))), inference(spm,[status(thm)],[c_0_38, c_0_42])).
fof(c_0_46, plain, ![X27, X28]:((~r2_hidden(X27,k1_ordinal1(X28))|r1_ordinal1(X27,X28)|~v3_ordinal1(X28)|~v3_ordinal1(X27))&(~r1_ordinal1(X27,X28)|r2_hidden(X27,k1_ordinal1(X28))|~v3_ordinal1(X28)|~v3_ordinal1(X27))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_ordinal1])])])])).
fof(c_0_47, plain, ![X29]:k1_ordinal1(X29)=k2_xboole_0(X29,k1_tarski(X29)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_48, plain, ![X30, X31]:((~r2_hidden(X30,k1_ordinal1(X31))|(r2_hidden(X30,X31)|X30=X31))&((~r2_hidden(X30,X31)|r2_hidden(X30,k1_ordinal1(X31)))&(X30!=X31|r2_hidden(X30,k1_ordinal1(X31))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_ordinal1])])])).
fof(c_0_49, plain, ![X32]:(v3_ordinal1(esk6_1(X32))&~r2_hidden(esk6_1(X32),X32)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_ordinal1])])])).
cnf(c_0_50, negated_conjecture, (esk3_1(esk4_1(esk1_0))=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_26])]), c_0_45])).
cnf(c_0_51, plain, (r1_ordinal1(X1,X2)|~r2_hidden(X1,k1_ordinal1(X2))|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53, plain, (r2_hidden(X1,k1_ordinal1(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_54, plain, ![X34, X35]:(~v3_ordinal1(X35)|(~r2_hidden(X34,X35)|v3_ordinal1(X34))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_ordinal1])])).
cnf(c_0_55, plain, (~r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, negated_conjecture, (r1_ordinal1(esk1_0,X1)|r2_hidden(X1,esk4_1(esk1_0))|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_31, c_0_50])).
cnf(c_0_57, plain, (v3_ordinal1(esk6_1(X1))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_58, plain, (r1_ordinal1(X1,X2)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))), inference(rw,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59, plain, (r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_53, c_0_52])).
cnf(c_0_60, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, negated_conjecture, (r1_ordinal1(esk1_0,esk6_1(esk4_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])])).
cnf(c_0_62, plain, (r1_ordinal1(X1,X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60])).
fof(c_0_63, plain, ![X36, X37, X38, X40, X41, X42, X43, X45]:((((r2_hidden(X38,esk7_3(X36,X37,X38))|~r2_hidden(X38,X37)|X37!=k3_tarski(X36))&(r2_hidden(esk7_3(X36,X37,X38),X36)|~r2_hidden(X38,X37)|X37!=k3_tarski(X36)))&(~r2_hidden(X40,X41)|~r2_hidden(X41,X36)|r2_hidden(X40,X37)|X37!=k3_tarski(X36)))&((~r2_hidden(esk8_2(X42,X43),X43)|(~r2_hidden(esk8_2(X42,X43),X45)|~r2_hidden(X45,X42))|X43=k3_tarski(X42))&((r2_hidden(esk8_2(X42,X43),esk9_2(X42,X43))|r2_hidden(esk8_2(X42,X43),X43)|X43=k3_tarski(X42))&(r2_hidden(esk9_2(X42,X43),X42)|r2_hidden(esk8_2(X42,X43),X43)|X43=k3_tarski(X42))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
cnf(c_0_64, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_65, negated_conjecture, (r1_tarski(esk1_0,esk6_1(esk4_1(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_61]), c_0_57]), c_0_26])])).
cnf(c_0_66, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,X2)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_62]), c_0_60])).
cnf(c_0_67, plain, (r2_hidden(esk7_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|X2!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_68, plain, ![X47]:((r2_hidden(esk10_1(X47),X47)|v3_ordinal1(k3_tarski(X47)))&(~v3_ordinal1(esk10_1(X47))|v3_ordinal1(k3_tarski(X47)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_ordinal1])])])])).
cnf(c_0_69, negated_conjecture, (r2_hidden(X1,esk6_1(esk4_1(esk1_0)))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_66])).
cnf(c_0_71, plain, (r2_hidden(esk7_3(X1,k3_tarski(X1),X2),X1)|~r2_hidden(X2,k3_tarski(X1))), inference(er,[status(thm)],[c_0_67])).
cnf(c_0_72, plain, (r2_hidden(X1,esk7_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k3_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_73, plain, (v3_ordinal1(k3_tarski(X1))|~v3_ordinal1(esk10_1(X1))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_74, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_69]), c_0_57])])).
cnf(c_0_75, plain, (~r2_hidden(X1,esk7_3(X1,k3_tarski(X1),X2))|~r2_hidden(X2,k3_tarski(X1))|~v3_ordinal1(esk7_3(X1,k3_tarski(X1),X2))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_76, plain, (r2_hidden(X1,esk7_3(X2,k3_tarski(X2),X1))|~r2_hidden(X1,k3_tarski(X2))), inference(er,[status(thm)],[c_0_72])).
cnf(c_0_77, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|X4!=k3_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_78, negated_conjecture, (v4_ordinal1(esk1_0)|r2_hidden(k1_ordinal1(X1),esk1_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_79, plain, ![X49]:r2_hidden(X49,k1_ordinal1(X49)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
cnf(c_0_80, negated_conjecture, (v3_ordinal1(k3_tarski(X1))|~r2_hidden(esk10_1(X1),esk1_0)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_81, plain, (r2_hidden(esk10_1(X1),X1)|v3_ordinal1(k3_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_82, plain, (~r2_hidden(X1,k3_tarski(X1))|~v3_ordinal1(esk7_3(X1,k3_tarski(X1),X1))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
fof(c_0_83, plain, ![X50, X51]:(((r1_tarski(X50,X51)|X50!=X51)&(r1_tarski(X51,X50)|X50!=X51))&(~r1_tarski(X50,X51)|~r1_tarski(X51,X50)|X50=X51)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_84, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_77])).
cnf(c_0_85, negated_conjecture, (v4_ordinal1(esk1_0)|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[c_0_78, c_0_52])).
cnf(c_0_86, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_87, negated_conjecture, (v3_ordinal1(k3_tarski(esk1_0))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_88, negated_conjecture, (~r2_hidden(esk7_3(X1,k3_tarski(X1),X1),esk1_0)|~r2_hidden(X1,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_82, c_0_74])).
cnf(c_0_89, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_83])).
fof(c_0_90, plain, ![X55]:((~v4_ordinal1(X55)|X55=k3_tarski(X55))&(X55!=k3_tarski(X55)|v4_ordinal1(X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_ordinal1])])).
cnf(c_0_91, negated_conjecture, (r2_hidden(X1,k3_tarski(esk1_0))|v4_ordinal1(esk1_0)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))|~r2_hidden(X2,esk1_0)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92, plain, (r2_hidden(X1,k2_xboole_0(X1,k1_tarski(X1)))), inference(rw,[status(thm)],[c_0_86, c_0_52])).
cnf(c_0_93, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(esk1_0,k3_tarski(esk1_0))|r2_hidden(k3_tarski(esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_87])).
cnf(c_0_94, negated_conjecture, (~r2_hidden(esk1_0,k3_tarski(esk1_0))), inference(spm,[status(thm)],[c_0_88, c_0_71])).
cnf(c_0_95, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_89])).
cnf(c_0_96, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_97, plain, (v4_ordinal1(X1)|X1!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_98, negated_conjecture, (r2_hidden(X1,k3_tarski(esk1_0))|v4_ordinal1(esk1_0)|~r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_99, negated_conjecture, (k3_tarski(esk1_0)=esk1_0|r2_hidden(k3_tarski(esk1_0),esk1_0)), inference(sr,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_100, plain, (~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_38, c_0_95])).
cnf(c_0_101, plain, (X1=k3_tarski(X1)|~v4_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_102, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|k3_tarski(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_103, negated_conjecture, (k3_tarski(esk1_0)=esk1_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_87])]), c_0_100]), c_0_101])).
cnf(c_0_104, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X3,X2)|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_64, c_0_66])).
cnf(c_0_105, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_106, negated_conjecture, (v3_ordinal1(esk2_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_107, negated_conjecture, (r2_hidden(esk2_0,X1)|~r2_hidden(esk1_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_108, negated_conjecture, (v3_ordinal1(esk2_0)|k3_tarski(esk1_0)!=esk1_0), inference(spm,[status(thm)],[c_0_106, c_0_97])).
cnf(c_0_109, negated_conjecture, (r2_hidden(esk2_0,esk3_1(X1))|~r2_hidden(esk1_0,esk3_1(X1))), inference(spm,[status(thm)],[c_0_107, c_0_32])).
cnf(c_0_110, negated_conjecture, (v3_ordinal1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_103])])).
cnf(c_0_111, plain, (~r2_hidden(esk3_1(k2_xboole_0(X1,k1_tarski(X1))),X1)), inference(spm,[status(thm)],[c_0_33, c_0_59])).
cnf(c_0_112, negated_conjecture, (r2_hidden(esk2_0,X1)|~r2_hidden(esk1_0,esk3_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_109]), c_0_110])])).
cnf(c_0_113, negated_conjecture, (esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))=esk1_0|r2_hidden(esk1_0,esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0))))), inference(spm,[status(thm)],[c_0_111, c_0_41])).
cnf(c_0_114, negated_conjecture, (esk3_1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))=esk1_0|r2_hidden(esk2_0,k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(spm,[status(thm)],[c_0_112, c_0_113])).
cnf(c_0_115, negated_conjecture, (r2_hidden(esk2_0,k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_114]), c_0_92])])).
fof(c_0_116, plain, ![X54]:(~v3_ordinal1(X54)|v3_ordinal1(k1_ordinal1(X54))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
cnf(c_0_117, negated_conjecture, (r2_hidden(esk2_0,X1)|~r2_hidden(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_104, c_0_115])).
cnf(c_0_118, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_119, negated_conjecture, (r2_hidden(esk2_0,k2_xboole_0(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),k1_tarski(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))))|~v3_ordinal1(k2_xboole_0(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),k1_tarski(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))))), inference(spm,[status(thm)],[c_0_117, c_0_92])).
cnf(c_0_120, plain, (v3_ordinal1(k2_xboole_0(X1,k1_tarski(X1)))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_118, c_0_52])).
cnf(c_0_121, negated_conjecture, (r2_hidden(esk2_0,k2_xboole_0(k2_xboole_0(esk1_0,k1_tarski(esk1_0)),k1_tarski(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))))|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(spm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_122, negated_conjecture, (r1_ordinal1(esk2_0,k2_xboole_0(esk1_0,k1_tarski(esk1_0)))|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_121]), c_0_110])])).
cnf(c_0_123, negated_conjecture, (r1_tarski(esk2_0,k2_xboole_0(esk1_0,k1_tarski(esk1_0)))|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_122]), c_0_110])])).
cnf(c_0_124, negated_conjecture, (esk3_1(esk1_0)=esk1_0|r2_hidden(esk1_0,esk3_1(esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_41])).
cnf(c_0_125, negated_conjecture, (r2_hidden(X1,k2_xboole_0(esk1_0,k1_tarski(esk1_0)))|~r2_hidden(X1,esk2_0)|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(spm,[status(thm)],[c_0_64, c_0_123])).
cnf(c_0_126, negated_conjecture, (esk3_1(esk1_0)=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_124]), c_0_26])]), c_0_100])).
cnf(c_0_127, negated_conjecture, (r2_hidden(X1,k2_xboole_0(esk1_0,k1_tarski(esk1_0)))|~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_120]), c_0_26])])).
cnf(c_0_128, negated_conjecture, (~r2_hidden(k1_ordinal1(esk2_0),esk1_0)|~v4_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_129, negated_conjecture, (r1_tarski(esk1_0,X1)|r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_39, c_0_126])).
cnf(c_0_130, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk2_0)|~v3_ordinal1(k2_xboole_0(esk1_0,k1_tarski(esk1_0)))), inference(spm,[status(thm)],[c_0_60, c_0_127])).
cnf(c_0_131, plain, (r2_hidden(X1,X2)|X1=X2|~r2_hidden(X1,k1_ordinal1(X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_132, negated_conjecture, (~v4_ordinal1(esk1_0)|~r2_hidden(k2_xboole_0(esk2_0,k1_tarski(esk2_0)),esk1_0)), inference(rw,[status(thm)],[c_0_128, c_0_52])).
cnf(c_0_133, negated_conjecture, (k2_xboole_0(X1,k1_tarski(X1))=esk1_0|r2_hidden(esk1_0,k2_xboole_0(X1,k1_tarski(X1)))|r2_hidden(k2_xboole_0(X1,k1_tarski(X1)),esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_120])).
cnf(c_0_134, negated_conjecture, (r2_hidden(X1,esk1_0)|r2_hidden(X2,X1)|~r2_hidden(X2,esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_64, c_0_129])).
cnf(c_0_135, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_105]), c_0_103])).
cnf(c_0_136, negated_conjecture, (v3_ordinal1(X1)|~r2_hidden(X1,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_120]), c_0_26])])).
cnf(c_0_137, plain, (X1=X2|r2_hidden(X1,X2)|~r2_hidden(X1,k2_xboole_0(X2,k1_tarski(X2)))), inference(rw,[status(thm)],[c_0_131, c_0_52])).
cnf(c_0_138, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|r2_hidden(esk1_0,k2_xboole_0(esk2_0,k1_tarski(esk2_0)))|~v4_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_110])])).
cnf(c_0_139, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_105]), c_0_110])])).
cnf(c_0_140, negated_conjecture, (r2_hidden(esk2_0,X1)|r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_134, c_0_105])).
cnf(c_0_141, negated_conjecture, (v3_ordinal1(k3_tarski(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_81]), c_0_80])).
cnf(c_0_142, negated_conjecture, (~r2_hidden(esk7_3(X1,k3_tarski(X1),X1),esk2_0)|~r2_hidden(X1,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_82, c_0_136])).
cnf(c_0_143, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|esk2_0=esk1_0|~v4_ordinal1(esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_139])).
cnf(c_0_144, negated_conjecture, (r2_hidden(k3_tarski(esk2_0),esk1_0)|r2_hidden(esk2_0,k3_tarski(esk2_0))), inference(spm,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_145, negated_conjecture, (~r2_hidden(esk2_0,k3_tarski(esk2_0))), inference(spm,[status(thm)],[c_0_142, c_0_71])).
cnf(c_0_146, plain, (~r2_hidden(esk7_3(X1,k3_tarski(X1),X2),X2)|~r2_hidden(X2,k3_tarski(X1))|~v3_ordinal1(X2)), inference(spm,[status(thm)],[c_0_70, c_0_76])).
cnf(c_0_147, negated_conjecture, (k2_xboole_0(esk2_0,k1_tarski(esk2_0))=esk1_0|esk2_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_97]), c_0_103])])).
cnf(c_0_148, negated_conjecture, (r2_hidden(k3_tarski(esk2_0),esk1_0)), inference(sr,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_149, negated_conjecture, (~r2_hidden(esk7_3(esk1_0,esk1_0,X1),X1)|~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_103]), c_0_74])).
cnf(c_0_150, negated_conjecture, (esk2_0=esk1_0|X1=esk2_0|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_137, c_0_147])).
cnf(c_0_151, negated_conjecture, (r2_hidden(k3_tarski(esk2_0),X1)|r2_hidden(X1,esk1_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_134, c_0_148])).
cnf(c_0_152, negated_conjecture, (esk7_3(esk1_0,esk1_0,esk2_0)=esk2_0|esk2_0=esk1_0|~r2_hidden(esk7_3(esk1_0,esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150]), c_0_105])])).
cnf(c_0_153, negated_conjecture, (r2_hidden(esk7_3(esk1_0,esk1_0,X1),esk1_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_71, c_0_103])).
cnf(c_0_154, negated_conjecture, (r2_hidden(esk3_1(X1),esk1_0)|r2_hidden(k3_tarski(esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_151]), c_0_141]), c_0_32])])).
cnf(c_0_155, negated_conjecture, (r2_hidden(X1,esk7_3(esk1_0,esk1_0,X1))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_103])).
cnf(c_0_156, negated_conjecture, (esk7_3(esk1_0,esk1_0,esk2_0)=esk2_0|esk2_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_105])])).
cnf(c_0_157, negated_conjecture, (r2_hidden(k3_tarski(esk2_0),esk4_1(esk1_0))), inference(spm,[status(thm)],[c_0_40, c_0_154])).
cnf(c_0_158, negated_conjecture, (esk2_0=esk1_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156]), c_0_105])]), c_0_100])).
cnf(c_0_159, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_157, c_0_158]), c_0_103]), c_0_45]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 160
# Proof object clause steps            : 123
# Proof object formula steps           : 37
# Proof object conjectures             : 74
# Proof object clause conjectures      : 71
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 74
# Proof object simplifying inferences  : 75
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 43
# Processed clauses                    : 1955
# ...of these trivial                  : 19
# ...subsumed                          : 969
# ...remaining for further processing  : 967
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 33
# Backward-rewritten                   : 415
# Generated clauses                    : 8649
# ...of the previous two non-trivial   : 8155
# Contextual simplify-reflections      : 52
# Paramodulations                      : 8639
# Factorizations                       : 1
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 469
#    Positive orientable unit clauses  : 26
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 97
#    Non-unit-clauses                  : 346
# Current number of unprocessed clauses: 6174
# ...number of literals in the above   : 18755
# Current number of archived formulas  : 0
# Current number of archived clauses   : 493
# Clause-clause subsumption calls (NU) : 59554
# Rec. Clause-clause subsumption calls : 24168
# Non-unit clause-clause subsumptions  : 306
# Unit Clause-clause subsumption calls : 20381
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 94
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 148228

# -------------------------------------------------
# User time                : 0.109 s
# System time              : 0.007 s
# Total time               : 0.115 s
# Maximum resident set size: 28756 pages
