# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t10_mcart_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', d1_xboole_0)).
fof(t69_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X6))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k5_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t69_mcart_1)).
fof(t34_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_mcart_1(X3,X4,X5,X6))))))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t34_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', d3_zfmisc_1)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', d4_zfmisc_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t2_subset)).
fof(t55_mcart_1, axiom, ![X1, X2, X3, X4]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(X4=k1_xboole_0))))<=>~(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0)), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t55_mcart_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', d1_enumset1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t2_xboole_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', d3_mcart_1)).
fof(t23_mcart_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1)))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t23_mcart_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', fc6_relat_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t1_subset)).
fof(dt_k7_mcart_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>m1_subset_1(k7_mcart_1(X1,X2,X3,X4),X3)), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', dt_k7_mcart_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/b0c3b9c86eeb4a6d91311e13dfd83f4f.E_input', t50_mcart_1)).
fof(c_0_17, plain, ![X9, X10, X11]:((r2_hidden(k1_mcart_1(X9),X10)|~r2_hidden(X9,k2_zfmisc_1(X10,X11)))&(r2_hidden(k2_mcart_1(X9),X11)|~r2_hidden(X9,k2_zfmisc_1(X10,X11)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
fof(c_0_18, plain, ![X12, X13, X14]:((~v1_xboole_0(X12)|~r2_hidden(X13,X12))&(r2_hidden(esk1_1(X14),X14)|v1_xboole_0(X14))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_19, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X6))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k5_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t69_mcart_1])).
fof(c_0_20, plain, ![X24, X26, X27, X28, X29]:((r2_hidden(esk7_1(X24),X24)|X24=k1_xboole_0)&((~r2_hidden(X26,X24)|esk7_1(X24)!=k4_mcart_1(X26,X27,X28,X29)|X24=k1_xboole_0)&(~r2_hidden(X27,X24)|esk7_1(X24)!=k4_mcart_1(X26,X27,X28,X29)|X24=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_mcart_1])])])])])])])).
cnf(c_0_21, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (r2_hidden(esk1_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, negated_conjecture, ![X21, X22, X23]:(m1_subset_1(esk6_0,k3_zfmisc_1(esk2_0,esk3_0,esk4_0))&((~m1_subset_1(X21,esk2_0)|(~m1_subset_1(X22,esk3_0)|(~m1_subset_1(X23,esk4_0)|(esk6_0!=k3_mcart_1(X21,X22,X23)|esk5_0=X21))))&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&esk5_0!=k5_mcart_1(esk2_0,esk3_0,esk4_0,esk6_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])])).
fof(c_0_24, plain, ![X30, X31, X32]:k3_zfmisc_1(X30,X31,X32)=k2_zfmisc_1(k2_zfmisc_1(X30,X31),X32), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_25, plain, ![X33, X34, X35, X36]:k4_zfmisc_1(X33,X34,X35,X36)=k2_zfmisc_1(k3_zfmisc_1(X33,X34,X35),X36), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
cnf(c_0_26, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (r2_hidden(esk7_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|r2_hidden(k1_mcart_1(esk1_1(k2_zfmisc_1(X1,X2))),X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_29, plain, ![X37, X38]:(~m1_subset_1(X37,X38)|(v1_xboole_0(X38)|r2_hidden(X37,X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk6_0,k3_zfmisc_1(esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_32, plain, ![X39, X40, X41, X42]:((X39=k1_xboole_0|(X40=k1_xboole_0|(X41=k1_xboole_0|X42=k1_xboole_0))|k4_zfmisc_1(X39,X40,X41,X42)!=k1_xboole_0)&((X39!=k1_xboole_0|k4_zfmisc_1(X39,X40,X41,X42)=k1_xboole_0)&((X40!=k1_xboole_0|k4_zfmisc_1(X39,X40,X41,X42)=k1_xboole_0)&((X41!=k1_xboole_0|k4_zfmisc_1(X39,X40,X41,X42)=k1_xboole_0)&(X42!=k1_xboole_0|k4_zfmisc_1(X39,X40,X41,X42)=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_mcart_1])])])])).
cnf(c_0_33, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_28])).
cnf(c_0_36, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_38, plain, ![X43, X44, X45, X46, X47, X48, X49, X50, X51, X52]:(((~r2_hidden(X47,X46)|(X47=X43|(X47=X44|X47=X45))|X46!=k1_enumset1(X43,X44,X45))&((X48!=X43|r2_hidden(X48,X46)|X46!=k1_enumset1(X43,X44,X45))&((X48!=X44|r2_hidden(X48,X46)|X46!=k1_enumset1(X43,X44,X45))&(X48!=X45|r2_hidden(X48,X46)|X46!=k1_enumset1(X43,X44,X45)))))&(((esk8_4(X49,X50,X51,X52)!=X49|~r2_hidden(esk8_4(X49,X50,X51,X52),X52)|X52=k1_enumset1(X49,X50,X51))&((esk8_4(X49,X50,X51,X52)!=X50|~r2_hidden(esk8_4(X49,X50,X51,X52),X52)|X52=k1_enumset1(X49,X50,X51))&(esk8_4(X49,X50,X51,X52)!=X51|~r2_hidden(esk8_4(X49,X50,X51,X52),X52)|X52=k1_enumset1(X49,X50,X51))))&(r2_hidden(esk8_4(X49,X50,X51,X52),X52)|(esk8_4(X49,X50,X51,X52)=X49|(esk8_4(X49,X50,X51,X52)=X50|esk8_4(X49,X50,X51,X52)=X51))|X52=k1_enumset1(X49,X50,X51)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_39, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k4_zfmisc_1(X1,X2,X3,X4)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_33, c_0_31])).
cnf(c_0_41, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))|r2_hidden(esk6_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_43, plain, ![X54, X55]:(~r2_hidden(X54,X55)|~r1_tarski(X55,X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_44, plain, ![X56]:r1_tarski(k1_xboole_0,X56), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_45, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_47, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0),X1)=k1_xboole_0|r2_hidden(esk6_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_49, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_50, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_51, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_53, plain, (r2_hidden(X1,k1_enumset1(X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_45])])).
cnf(c_0_54, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk6_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_55, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
fof(c_0_56, plain, ![X57, X58, X59]:k3_mcart_1(X57,X58,X59)=k4_tarski(k4_tarski(X57,X58),X59), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_57, plain, ![X60, X61]:(~v1_relat_1(X61)|(~r2_hidden(X60,X61)|X60=k4_tarski(k1_mcart_1(X60),k2_mcart_1(X60)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_mcart_1])])).
cnf(c_0_58, negated_conjecture, (r2_hidden(esk6_0,k2_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0),esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])).
fof(c_0_59, plain, ![X62, X63]:v1_relat_1(k2_zfmisc_1(X62,X63)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_60, negated_conjecture, (esk5_0=X1|~m1_subset_1(X1,esk2_0)|~m1_subset_1(X2,esk3_0)|~m1_subset_1(X3,esk4_0)|esk6_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_61, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_62, plain, (X2=k4_tarski(k1_mcart_1(X2),k2_mcart_1(X2))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_63, negated_conjecture, (r2_hidden(k1_mcart_1(esk6_0),k2_zfmisc_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_21, c_0_58])).
cnf(c_0_64, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65, negated_conjecture, (esk5_0=X1|esk6_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk4_0)|~m1_subset_1(X2,esk3_0)|~m1_subset_1(X1,esk2_0)), inference(rw,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_66, negated_conjecture, (k4_tarski(k1_mcart_1(k1_mcart_1(esk6_0)),k2_mcart_1(k1_mcart_1(esk6_0)))=k1_mcart_1(esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])])).
fof(c_0_67, plain, ![X64, X65]:(~r2_hidden(X64,X65)|m1_subset_1(X64,X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_68, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_69, plain, ![X66, X67, X68, X69]:(~m1_subset_1(X69,k3_zfmisc_1(X66,X67,X68))|m1_subset_1(k7_mcart_1(X66,X67,X68,X69),X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_mcart_1])])).
fof(c_0_70, plain, ![X70, X71, X72, X73]:((k5_mcart_1(X70,X71,X72,X73)=k1_mcart_1(k1_mcart_1(X73))|~m1_subset_1(X73,k3_zfmisc_1(X70,X71,X72))|X72=k1_xboole_0|X71=k1_xboole_0|X70=k1_xboole_0)&((k6_mcart_1(X70,X71,X72,X73)=k2_mcart_1(k1_mcart_1(X73))|~m1_subset_1(X73,k3_zfmisc_1(X70,X71,X72))|X72=k1_xboole_0|X71=k1_xboole_0|X70=k1_xboole_0)&(k7_mcart_1(X70,X71,X72,X73)=k2_mcart_1(X73)|~m1_subset_1(X73,k3_zfmisc_1(X70,X71,X72))|X72=k1_xboole_0|X71=k1_xboole_0|X70=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_71, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk6_0))=esk5_0|k4_tarski(k1_mcart_1(esk6_0),X1)!=esk6_0|~m1_subset_1(k2_mcart_1(k1_mcart_1(esk6_0)),esk3_0)|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk6_0)),esk2_0)|~m1_subset_1(X1,esk4_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_72, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_73, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk6_0)),esk3_0)), inference(spm,[status(thm)],[c_0_68, c_0_63])).
cnf(c_0_74, plain, (m1_subset_1(k7_mcart_1(X2,X3,X4,X1),X4)|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_75, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_76, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk6_0))=esk5_0|k4_tarski(k1_mcart_1(esk6_0),X1)!=esk6_0|~m1_subset_1(k1_mcart_1(k1_mcart_1(esk6_0)),esk2_0)|~m1_subset_1(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])])).
cnf(c_0_77, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk6_0)),esk2_0)), inference(spm,[status(thm)],[c_0_21, c_0_63])).
cnf(c_0_78, plain, (m1_subset_1(k7_mcart_1(X2,X3,X4,X1),X4)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[c_0_74, c_0_31])).
cnf(c_0_79, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_75, c_0_31])).
cnf(c_0_80, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk6_0))=esk5_0|k4_tarski(k1_mcart_1(esk6_0),X1)!=esk6_0|~m1_subset_1(X1,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_72]), c_0_77])])).
cnf(c_0_81, negated_conjecture, (k4_tarski(k1_mcart_1(esk6_0),k2_mcart_1(esk6_0))=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_58]), c_0_64])])).
cnf(c_0_82, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|m1_subset_1(k2_mcart_1(X4),X3)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_83, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_84, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk6_0))=esk5_0|~m1_subset_1(k2_mcart_1(esk6_0),esk4_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(k2_mcart_1(esk6_0),esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_37]), c_0_50]), c_0_49]), c_0_48])).
cnf(c_0_86, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_83, c_0_31])).
cnf(c_0_87, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk6_0))=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_84, c_0_85])])).
cnf(c_0_88, negated_conjecture, (esk5_0!=k5_mcart_1(esk2_0,esk3_0,esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_89, negated_conjecture, (k5_mcart_1(X1,X2,X3,esk6_0)=esk5_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_90, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_37])]), c_0_50]), c_0_49]), c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 91
# Proof object clause steps            : 56
# Proof object formula steps           : 35
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 36
# Processed clauses                    : 284
# ...of these trivial                  : 9
# ...subsumed                          : 78
# ...remaining for further processing  : 197
# Other redundant clauses eliminated   : 27
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 11
# Generated clauses                    : 2356
# ...of the previous two non-trivial   : 744
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2290
# Factorizations                       : 42
# Equation resolutions                 : 27
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 119
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 89
# Current number of unprocessed clauses: 269
# ...number of literals in the above   : 1167
# Current number of archived formulas  : 0
# Current number of archived clauses   : 73
# Clause-clause subsumption calls (NU) : 2039
# Rec. Clause-clause subsumption calls : 1504
# Non-unit clause-clause subsumptions  : 65
# Unit Clause-clause subsumption calls : 82
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 26
# BW rewrite match successes           : 13
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 26827

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.003 s
# Total time               : 0.022 s
# Maximum resident set size: 28684 pages
