# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', t69_enumset1)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', fc2_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', d1_xboole_0)).
fof(d1_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X3,X5)&r1_orders_2(X1,X4,X5))))))))))))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', d1_waybel_0)).
fof(t35_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', t35_orders_2)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', cc1_subset_1)).
fof(d2_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v2_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r2_hidden(X5,X2)&(r1_orders_2(X1,X5,X3)&r1_orders_2(X1,X5,X4))))))))))))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', d2_waybel_0)).
fof(t5_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', t5_waybel_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', redefinition_k6_domain_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', t1_subset)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', d2_tarski)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', t2_subset)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', t4_subset)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/b123ef3533aa4273913b36c153725de6.E_input', t24_orders_2)).
fof(c_0_15, plain, ![X6, X7, X8, X9, X10, X11]:(((~r2_hidden(X8,X7)|X8=X6|X7!=k1_tarski(X6))&(X9!=X6|r2_hidden(X9,X7)|X7!=k1_tarski(X6)))&((~r2_hidden(esk1_2(X10,X11),X11)|esk1_2(X10,X11)!=X10|X11=k1_tarski(X10))&(r2_hidden(esk1_2(X10,X11),X11)|esk1_2(X10,X11)=X10|X11=k1_tarski(X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_16, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_17, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, ![X14]:~v1_xboole_0(k1_tarski(X14)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
fof(c_0_20, plain, ![X15, X16, X17]:((~v1_xboole_0(X15)|~r2_hidden(X16,X15))&(r2_hidden(esk2_1(X17),X17)|v1_xboole_0(X17))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_21, plain, ![X31, X32, X33, X34, X38]:(((m1_subset_1(esk6_4(X31,X32,X33,X34),u1_struct_0(X31))|~r2_hidden(X34,X32)|~r2_hidden(X33,X32)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31))&((r2_hidden(esk6_4(X31,X32,X33,X34),X32)|~r2_hidden(X34,X32)|~r2_hidden(X33,X32)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31))&((r1_orders_2(X31,X33,esk6_4(X31,X32,X33,X34))|~r2_hidden(X34,X32)|~r2_hidden(X33,X32)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31))&(r1_orders_2(X31,X34,esk6_4(X31,X32,X33,X34))|~r2_hidden(X34,X32)|~r2_hidden(X33,X32)|~m1_subset_1(X34,u1_struct_0(X31))|~m1_subset_1(X33,u1_struct_0(X31))|~v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31)))))&((m1_subset_1(esk7_2(X31,X32),u1_struct_0(X31))|v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31))&((m1_subset_1(esk8_2(X31,X32),u1_struct_0(X31))|v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31))&((r2_hidden(esk7_2(X31,X32),X32)|v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31))&((r2_hidden(esk8_2(X31,X32),X32)|v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31))&(~m1_subset_1(X38,u1_struct_0(X31))|(~r2_hidden(X38,X32)|(~r1_orders_2(X31,esk7_2(X31,X32),X38)|~r1_orders_2(X31,esk8_2(X31,X32),X38)))|v1_waybel_0(X32,X31)|~m1_subset_1(X32,k1_zfmisc_1(u1_struct_0(X31)))|~l1_orders_2(X31))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_waybel_0])])])])])).
cnf(c_0_22, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, plain, (r2_hidden(esk7_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_26, plain, ![X27, X28]:((v6_orders_2(k6_domain_1(u1_struct_0(X27),X28),X27)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v3_orders_2(X27)|~l1_orders_2(X27))))&(m1_subset_1(k6_domain_1(u1_struct_0(X27),X28),k1_zfmisc_1(u1_struct_0(X27)))|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v3_orders_2(X27)|~l1_orders_2(X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_orders_2])])])])])).
fof(c_0_27, plain, ![X29, X30]:(~v1_xboole_0(X29)|(~m1_subset_1(X30,k1_zfmisc_1(X29))|v1_xboole_0(X30))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_28, plain, ![X19, X20, X21, X22, X26]:(((m1_subset_1(esk3_4(X19,X20,X21,X22),u1_struct_0(X19))|~r2_hidden(X22,X20)|~r2_hidden(X21,X20)|~m1_subset_1(X22,u1_struct_0(X19))|~m1_subset_1(X21,u1_struct_0(X19))|~v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&((r2_hidden(esk3_4(X19,X20,X21,X22),X20)|~r2_hidden(X22,X20)|~r2_hidden(X21,X20)|~m1_subset_1(X22,u1_struct_0(X19))|~m1_subset_1(X21,u1_struct_0(X19))|~v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&((r1_orders_2(X19,esk3_4(X19,X20,X21,X22),X21)|~r2_hidden(X22,X20)|~r2_hidden(X21,X20)|~m1_subset_1(X22,u1_struct_0(X19))|~m1_subset_1(X21,u1_struct_0(X19))|~v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&(r1_orders_2(X19,esk3_4(X19,X20,X21,X22),X22)|~r2_hidden(X22,X20)|~r2_hidden(X21,X20)|~m1_subset_1(X22,u1_struct_0(X19))|~m1_subset_1(X21,u1_struct_0(X19))|~v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19)))))&((m1_subset_1(esk4_2(X19,X20),u1_struct_0(X19))|v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&((m1_subset_1(esk5_2(X19,X20),u1_struct_0(X19))|v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&((r2_hidden(esk4_2(X19,X20),X20)|v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&((r2_hidden(esk5_2(X19,X20),X20)|v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))&(~m1_subset_1(X26,u1_struct_0(X19))|(~r2_hidden(X26,X20)|(~r1_orders_2(X19,X26,esk4_2(X19,X20))|~r1_orders_2(X19,X26,esk5_2(X19,X20))))|v2_waybel_0(X20,X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X19)))|~l1_orders_2(X19))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_waybel_0])])])])])).
cnf(c_0_29, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (r2_hidden(esk2_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (~v1_xboole_0(k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_23, c_0_18])).
cnf(c_0_32, plain, (r2_hidden(esk1_2(X1,X2),X2)|esk1_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_33, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)&v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1))))), inference(assume_negation,[status(cth)],[t5_waybel_0])).
cnf(c_0_34, plain, (v1_waybel_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_35, plain, (m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (r2_hidden(esk4_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (esk2_1(k2_tarski(X1,X1))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])).
cnf(c_0_39, plain, (X2=k2_tarski(X1,X1)|esk1_2(X1,X2)=X1|r2_hidden(esk1_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_32, c_0_18])).
fof(c_0_40, plain, ![X41, X42]:(v1_xboole_0(X41)|~m1_subset_1(X42,X41)|k6_domain_1(X41,X42)=k1_tarski(X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_41, negated_conjecture, ((~v2_struct_0(esk9_0)&(v3_orders_2(esk9_0)&l1_orders_2(esk9_0)))&(m1_subset_1(esk10_0,u1_struct_0(esk9_0))&(~v1_waybel_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk9_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk9_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_33])])])])).
cnf(c_0_42, plain, (v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, plain, (v2_struct_0(X1)|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_36, c_0_35])).
cnf(c_0_44, plain, (v2_waybel_0(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_24, c_0_37])).
cnf(c_0_45, plain, (esk1_2(X1,X2)=X1|esk2_1(X2)=X1|r2_hidden(esk1_2(X1,X2),X2)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, negated_conjecture, (~v1_waybel_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk9_0)|~v2_waybel_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk9_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, plain, (v1_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, negated_conjecture, (v3_orders_2(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk10_0,u1_struct_0(esk9_0))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_51, negated_conjecture, (l1_orders_2(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, negated_conjecture, (~v2_struct_0(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (v2_struct_0(X1)|v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))), inference(spm,[status(thm)],[c_0_44, c_0_35])).
fof(c_0_54, plain, ![X43, X44]:(~r2_hidden(X43,X44)|m1_subset_1(X43,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_55, plain, (esk1_2(X1,k2_tarski(X2,X2))=X1|esk1_2(X1,k2_tarski(X2,X2))=X2|X2=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_45]), c_0_38])).
cnf(c_0_56, plain, (k6_domain_1(X1,X2)=k2_tarski(X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_46, c_0_18])).
cnf(c_0_57, negated_conjecture, (~v2_waybel_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk9_0)|~v1_xboole_0(u1_struct_0(esk9_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])).
cnf(c_0_58, plain, (v2_struct_0(X1)|v2_waybel_0(k6_domain_1(u1_struct_0(X1),X2),X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_53, c_0_43])).
cnf(c_0_59, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_60, plain, (esk1_2(X1,k6_domain_1(X2,X3))=X3|esk1_2(X1,k6_domain_1(X2,X3))=X1|X3=X1|v1_xboole_0(X2)|~m1_subset_1(X3,X2)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_61, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk9_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_49]), c_0_50]), c_0_51])]), c_0_52])).
cnf(c_0_62, plain, (esk1_2(X1,X2)=X1|esk2_1(X2)=X1|m1_subset_1(esk1_2(X1,X2),X2)), inference(spm,[status(thm)],[c_0_59, c_0_45])).
cnf(c_0_63, negated_conjecture, (esk1_2(X1,k6_domain_1(u1_struct_0(esk9_0),esk10_0))=esk10_0|esk1_2(X1,k6_domain_1(u1_struct_0(esk9_0),esk10_0))=X1|esk10_0=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_50]), c_0_61])).
cnf(c_0_64, plain, (X2=k1_tarski(X1)|~r2_hidden(esk1_2(X1,X2),X2)|esk1_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_65, plain, (esk2_1(k6_domain_1(X1,X2))=X2|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_56])).
cnf(c_0_66, negated_conjecture, (esk1_2(X1,k6_domain_1(u1_struct_0(esk9_0),esk10_0))=X1|esk2_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0))=X1|esk10_0=X1|m1_subset_1(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_67, plain, (X2=k2_tarski(X1,X1)|esk1_2(X1,X2)!=X1|~r2_hidden(esk1_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_64, c_0_18])).
cnf(c_0_68, negated_conjecture, (esk1_2(X1,k6_domain_1(u1_struct_0(esk9_0),esk10_0))=X1|X1=esk10_0|m1_subset_1(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_50])]), c_0_61])).
fof(c_0_69, plain, ![X45, X46, X47, X48, X49, X50, X51, X52]:(((~r2_hidden(X48,X47)|(X48=X45|X48=X46)|X47!=k2_tarski(X45,X46))&((X49!=X45|r2_hidden(X49,X47)|X47!=k2_tarski(X45,X46))&(X49!=X46|r2_hidden(X49,X47)|X47!=k2_tarski(X45,X46))))&(((esk11_3(X50,X51,X52)!=X50|~r2_hidden(esk11_3(X50,X51,X52),X52)|X52=k2_tarski(X50,X51))&(esk11_3(X50,X51,X52)!=X51|~r2_hidden(esk11_3(X50,X51,X52),X52)|X52=k2_tarski(X50,X51)))&(r2_hidden(esk11_3(X50,X51,X52),X52)|(esk11_3(X50,X51,X52)=X50|esk11_3(X50,X51,X52)=X51)|X52=k2_tarski(X50,X51)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_70, negated_conjecture, (k6_domain_1(u1_struct_0(esk9_0),esk10_0)=k2_tarski(X1,X1)|X1=esk10_0|m1_subset_1(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))|~r2_hidden(X1,k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_71, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_72, negated_conjecture, (k2_tarski(esk2_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0)),esk2_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0)))=k6_domain_1(u1_struct_0(esk9_0),esk10_0)|esk2_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0))=esk10_0|m1_subset_1(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))|v1_xboole_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(spm,[status(thm)],[c_0_70, c_0_30])).
cnf(c_0_73, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_71])])).
cnf(c_0_74, negated_conjecture, (esk2_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0))=esk10_0|X1=esk2_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0))|m1_subset_1(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))|~r2_hidden(X1,k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_72]), c_0_24])).
cnf(c_0_75, plain, (v1_xboole_0(X1)|r2_hidden(X2,k6_domain_1(X1,X2))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_73, c_0_56])).
cnf(c_0_76, plain, (m1_subset_1(esk2_1(X1),X1)|v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_59, c_0_30])).
cnf(c_0_77, negated_conjecture, (esk2_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0))=esk10_0|m1_subset_1(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_50])]), c_0_61])).
fof(c_0_78, plain, ![X54, X55]:(~m1_subset_1(X54,X55)|(v1_xboole_0(X55)|r2_hidden(X54,X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_79, plain, (v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~v1_xboole_0(k6_domain_1(X1,X2))), inference(spm,[status(thm)],[c_0_31, c_0_56])).
cnf(c_0_80, negated_conjecture, (m1_subset_1(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))|v1_xboole_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_81, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_82, negated_conjecture, (m1_subset_1(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_50])]), c_0_61])).
cnf(c_0_83, plain, (esk1_2(X1,X2)=X1|k6_domain_1(X3,X1)=X2|v1_xboole_0(X3)|r2_hidden(esk1_2(X1,X2),X2)|~m1_subset_1(X1,X3)), inference(spm,[status(thm)],[c_0_56, c_0_39])).
cnf(c_0_84, negated_conjecture, (v1_xboole_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0))|r2_hidden(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_85, plain, (X1=X2|v1_xboole_0(X3)|~m1_subset_1(X2,X3)|~r2_hidden(X1,k6_domain_1(X3,X2))), inference(spm,[status(thm)],[c_0_29, c_0_56])).
cnf(c_0_86, negated_conjecture, (k6_domain_1(u1_struct_0(esk9_0),esk10_0)=X1|esk1_2(esk10_0,X1)=esk10_0|r2_hidden(esk1_2(esk10_0,X1),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_50]), c_0_61])).
cnf(c_0_87, negated_conjecture, (r2_hidden(esk10_0,k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_84]), c_0_50])]), c_0_61])).
fof(c_0_88, plain, ![X56, X57, X58]:(~r2_hidden(X56,X57)|~m1_subset_1(X57,k1_zfmisc_1(X58))|m1_subset_1(X56,X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_89, plain, ![X59, X60]:(v2_struct_0(X59)|(~v3_orders_2(X59)|~l1_orders_2(X59))|(~m1_subset_1(X60,u1_struct_0(X59))|r1_orders_2(X59,X60,X60))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_90, negated_conjecture, (k6_domain_1(u1_struct_0(esk9_0),esk10_0)=k6_domain_1(X1,X2)|esk1_2(esk10_0,k6_domain_1(X1,X2))=esk10_0|esk1_2(esk10_0,k6_domain_1(X1,X2))=X2|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_91, negated_conjecture, (~v1_xboole_0(k6_domain_1(u1_struct_0(esk9_0),esk10_0))), inference(spm,[status(thm)],[c_0_24, c_0_87])).
cnf(c_0_92, plain, (v1_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,esk7_2(X2,X3),X1)|~r1_orders_2(X2,esk8_2(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_93, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_94, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_95, plain, (m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_96, negated_conjecture, (k6_domain_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk10_0)=k6_domain_1(u1_struct_0(esk9_0),esk10_0)|esk1_2(esk10_0,k6_domain_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk10_0))=esk10_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_82]), c_0_91])).
cnf(c_0_97, plain, (v1_waybel_0(X1,X2)|~r1_orders_2(X2,esk7_2(X2,X1),X3)|~r1_orders_2(X2,esk8_2(X2,X1),X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,X1)), inference(csr,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_98, plain, (v1_waybel_0(X1,X2)|v2_struct_0(X2)|r1_orders_2(X2,esk8_2(X2,X1),esk8_2(X2,X1))|~v3_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_99, plain, (r2_hidden(esk8_2(X1,X2),X2)|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_100, plain, (m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|v1_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_101, plain, (v2_waybel_0(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r1_orders_2(X2,X1,esk4_2(X2,X3))|~r1_orders_2(X2,X1,esk5_2(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_102, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_103, negated_conjecture, (k6_domain_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk10_0)=k6_domain_1(u1_struct_0(esk9_0),esk10_0)|k6_domain_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk10_0)=k2_tarski(esk10_0,esk10_0)|~r2_hidden(esk10_0,k6_domain_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk10_0))), inference(spm,[status(thm)],[c_0_67, c_0_96])).
cnf(c_0_104, plain, (v1_waybel_0(X1,X2)|v2_struct_0(X2)|~v3_orders_2(X2)|~r1_orders_2(X2,esk7_2(X2,X1),esk8_2(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])).
cnf(c_0_105, plain, (esk8_2(X1,k2_tarski(X2,X2))=X2|v1_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_99])).
cnf(c_0_106, plain, (v1_waybel_0(X1,X2)|v2_struct_0(X2)|r1_orders_2(X2,esk7_2(X2,X1),esk7_2(X2,X1))|~v3_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_94, c_0_100])).
cnf(c_0_107, plain, (esk7_2(X1,k2_tarski(X2,X2))=X2|v1_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_25])).
cnf(c_0_108, plain, (v2_waybel_0(X1,X2)|~r1_orders_2(X2,X3,esk4_2(X2,X1))|~r1_orders_2(X2,X3,esk5_2(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~r2_hidden(X3,X1)), inference(csr,[status(thm)],[c_0_101, c_0_93])).
cnf(c_0_109, plain, (v2_struct_0(X1)|r1_orders_2(X1,esk5_2(X1,X2),esk5_2(X1,X2))|v2_waybel_0(X2,X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_94, c_0_102])).
cnf(c_0_110, plain, (r2_hidden(esk5_2(X1,X2),X2)|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_111, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v2_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_112, plain, (k6_domain_1(X1,X2)=k6_domain_1(X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X2,X3)), inference(spm,[status(thm)],[c_0_56, c_0_56])).
cnf(c_0_113, negated_conjecture, (k6_domain_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk10_0)=k6_domain_1(u1_struct_0(esk9_0),esk10_0)|k6_domain_1(k6_domain_1(u1_struct_0(esk9_0),esk10_0),esk10_0)=k2_tarski(esk10_0,esk10_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_75]), c_0_82])]), c_0_91])).
cnf(c_0_114, plain, (v1_waybel_0(k2_tarski(X1,X1),X2)|v2_struct_0(X2)|~v3_orders_2(X2)|~r1_orders_2(X2,esk7_2(X2,k2_tarski(X1,X1)),X1)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_115, plain, (v1_waybel_0(k2_tarski(X1,X1),X2)|v2_struct_0(X2)|r1_orders_2(X2,X1,X1)|~v3_orders_2(X2)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_116, plain, (v2_struct_0(X1)|v2_waybel_0(X2,X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk5_2(X1,X2),esk4_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])).
cnf(c_0_117, plain, (esk4_2(X1,k2_tarski(X2,X2))=X2|v2_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_37])).
cnf(c_0_118, plain, (v2_struct_0(X1)|r1_orders_2(X1,esk4_2(X1,X2),esk4_2(X1,X2))|v2_waybel_0(X2,X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_94, c_0_111])).
cnf(c_0_119, negated_conjecture, (v1_xboole_0(u1_struct_0(esk9_0))|v1_xboole_0(X1)|~v1_waybel_0(k6_domain_1(X1,esk10_0),esk9_0)|~v2_waybel_0(k6_domain_1(X1,esk10_0),esk9_0)|~m1_subset_1(esk10_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_112]), c_0_50])])).
cnf(c_0_120, negated_conjecture, (k6_domain_1(k2_tarski(esk10_0,esk10_0),esk10_0)=k2_tarski(esk10_0,esk10_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_56]), c_0_50])]), c_0_61])).
cnf(c_0_121, plain, (m1_subset_1(X1,k2_tarski(X1,X2))), inference(spm,[status(thm)],[c_0_59, c_0_73])).
cnf(c_0_122, plain, (v1_waybel_0(k2_tarski(X1,X1),X2)|v2_struct_0(X2)|~v3_orders_2(X2)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_107]), c_0_115])).
cnf(c_0_123, plain, (v2_struct_0(X1)|m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_35, c_0_56])).
cnf(c_0_124, plain, (v2_struct_0(X1)|v2_waybel_0(k2_tarski(X2,X2),X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk5_2(X1,k2_tarski(X2,X2)),X2)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_125, plain, (esk5_2(X1,k2_tarski(X2,X2))=X2|v2_waybel_0(k2_tarski(X2,X2),X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_110])).
cnf(c_0_126, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|v2_waybel_0(k2_tarski(X2,X2),X1)|~v3_orders_2(X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_118, c_0_117])).
cnf(c_0_127, negated_conjecture, (~v1_waybel_0(k2_tarski(esk10_0,esk10_0),esk9_0)|~v2_waybel_0(k2_tarski(esk10_0,esk10_0),esk9_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_121])]), c_0_61]), c_0_31])).
cnf(c_0_128, plain, (v1_waybel_0(k2_tarski(X1,X1),X2)|v2_struct_0(X2)|v1_xboole_0(u1_struct_0(X2))|~v3_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_129, plain, (v2_struct_0(X1)|v2_waybel_0(k2_tarski(X2,X2),X1)|~v3_orders_2(X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126])).
cnf(c_0_130, negated_conjecture, (~v2_waybel_0(k2_tarski(esk10_0,esk10_0),esk9_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_49]), c_0_50]), c_0_51])]), c_0_52]), c_0_61])).
cnf(c_0_131, plain, (v2_struct_0(X1)|v2_waybel_0(k2_tarski(X2,X2),X1)|v1_xboole_0(u1_struct_0(X1))|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_129, c_0_123])).
cnf(c_0_132, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_49]), c_0_50]), c_0_51])]), c_0_52]), c_0_61]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 133
# Proof object clause steps            : 102
# Proof object formula steps           : 31
# Proof object conjectures             : 32
# Proof object clause conjectures      : 29
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 66
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 45
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 44
# Processed clauses                    : 485
# ...of these trivial                  : 3
# ...subsumed                          : 203
# ...remaining for further processing  : 279
# Other redundant clauses eliminated   : 29
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 17
# Backward-rewritten                   : 29
# Generated clauses                    : 2427
# ...of the previous two non-trivial   : 2176
# Contextual simplify-reflections      : 27
# Paramodulations                      : 2383
# Factorizations                       : 17
# Equation resolutions                 : 29
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 184
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 164
# Current number of unprocessed clauses: 1671
# ...number of literals in the above   : 10367
# Current number of archived formulas  : 0
# Current number of archived clauses   : 91
# Clause-clause subsumption calls (NU) : 9767
# Rec. Clause-clause subsumption calls : 1415
# Non-unit clause-clause subsumptions  : 210
# Unit Clause-clause subsumption calls : 92
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 53244

# -------------------------------------------------
# User time                : 0.039 s
# System time              : 0.006 s
# Total time               : 0.045 s
# Maximum resident set size: 28708 pages
