# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t58_setfam_1, conjecture, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(r2_hidden(X2,X1)=>(r2_hidden(X2,k8_setfam_1(X1,X3))<=>![X4]:(r2_hidden(X4,X3)=>r2_hidden(X2,X4))))), file('data/MPTP/proofs/b1c0ada41e5f45d08bb68c4ef0357c36.E_input', t58_setfam_1)).
fof(d1_setfam_1, axiom, ![X1, X2]:((~(X1=k1_xboole_0)=>(X2=k1_setfam_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>![X4]:(r2_hidden(X4,X1)=>r2_hidden(X3,X4)))))&(X1=k1_xboole_0=>(X2=k1_setfam_1(X1)<=>X2=k1_xboole_0))), file('data/MPTP/proofs/b1c0ada41e5f45d08bb68c4ef0357c36.E_input', d1_setfam_1)).
fof(d10_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>((~(X2=k1_xboole_0)=>k8_setfam_1(X1,X2)=k6_setfam_1(X1,X2))&(X2=k1_xboole_0=>k8_setfam_1(X1,X2)=X1))), file('data/MPTP/proofs/b1c0ada41e5f45d08bb68c4ef0357c36.E_input', d10_setfam_1)).
fof(redefinition_k6_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k6_setfam_1(X1,X2)=k1_setfam_1(X2)), file('data/MPTP/proofs/b1c0ada41e5f45d08bb68c4ef0357c36.E_input', redefinition_k6_setfam_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/b1c0ada41e5f45d08bb68c4ef0357c36.E_input', t113_zfmisc_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/b1c0ada41e5f45d08bb68c4ef0357c36.E_input', t4_subset_1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/b1c0ada41e5f45d08bb68c4ef0357c36.E_input', t152_zfmisc_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(r2_hidden(X2,X1)=>(r2_hidden(X2,k8_setfam_1(X1,X3))<=>![X4]:(r2_hidden(X4,X3)=>r2_hidden(X2,X4)))))), inference(assume_negation,[status(cth)],[t58_setfam_1])).
fof(c_0_8, plain, ![X19, X20, X21, X22, X23, X25, X28, X29, X30, X31]:((((~r2_hidden(X21,X20)|(~r2_hidden(X22,X19)|r2_hidden(X21,X22))|X20!=k1_setfam_1(X19)|X19=k1_xboole_0)&((r2_hidden(esk5_3(X19,X20,X23),X19)|r2_hidden(X23,X20)|X20!=k1_setfam_1(X19)|X19=k1_xboole_0)&(~r2_hidden(X23,esk5_3(X19,X20,X23))|r2_hidden(X23,X20)|X20!=k1_setfam_1(X19)|X19=k1_xboole_0)))&(((r2_hidden(esk7_2(X19,X25),X19)|~r2_hidden(esk6_2(X19,X25),X25)|X25=k1_setfam_1(X19)|X19=k1_xboole_0)&(~r2_hidden(esk6_2(X19,X25),esk7_2(X19,X25))|~r2_hidden(esk6_2(X19,X25),X25)|X25=k1_setfam_1(X19)|X19=k1_xboole_0))&(r2_hidden(esk6_2(X19,X25),X25)|(~r2_hidden(X28,X19)|r2_hidden(esk6_2(X19,X25),X28))|X25=k1_setfam_1(X19)|X19=k1_xboole_0)))&((X30!=k1_setfam_1(X29)|X30=k1_xboole_0|X29!=k1_xboole_0)&(X31!=k1_xboole_0|X31=k1_setfam_1(X29)|X29!=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_setfam_1])])])])])])])).
fof(c_0_9, negated_conjecture, ![X9]:(m1_subset_1(esk3_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))&(r2_hidden(esk2_0,esk1_0)&(((r2_hidden(esk4_0,esk3_0)|~r2_hidden(esk2_0,k8_setfam_1(esk1_0,esk3_0)))&(~r2_hidden(esk2_0,esk4_0)|~r2_hidden(esk2_0,k8_setfam_1(esk1_0,esk3_0))))&(r2_hidden(esk2_0,k8_setfam_1(esk1_0,esk3_0))|(~r2_hidden(X9,esk3_0)|r2_hidden(esk2_0,X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])])).
cnf(c_0_10, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_hidden(X3,X2)|X1=k1_xboole_0|X2!=k1_setfam_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, plain, ![X16, X17]:((X17=k1_xboole_0|k8_setfam_1(X16,X17)=k6_setfam_1(X16,X17)|~m1_subset_1(X17,k1_zfmisc_1(k1_zfmisc_1(X16))))&(X17!=k1_xboole_0|k8_setfam_1(X16,X17)=X16|~m1_subset_1(X17,k1_zfmisc_1(k1_zfmisc_1(X16))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_setfam_1])])])])).
fof(c_0_12, plain, ![X14, X15]:(~m1_subset_1(X15,k1_zfmisc_1(k1_zfmisc_1(X14)))|k6_setfam_1(X14,X15)=k1_setfam_1(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k6_setfam_1])])).
cnf(c_0_13, negated_conjecture, (r2_hidden(esk2_0,k8_setfam_1(esk1_0,esk3_0))|r2_hidden(esk2_0,X1)|~r2_hidden(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (X1=k1_xboole_0|r2_hidden(esk5_3(X1,k1_setfam_1(X1),X2),X1)|r2_hidden(X2,k1_setfam_1(X1))), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (X1=k1_xboole_0|k8_setfam_1(X2,X1)=k6_setfam_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (k6_setfam_1(X2,X1)=k1_setfam_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (r2_hidden(X1,X3)|X2=k1_xboole_0|~r2_hidden(X1,esk5_3(X2,X3,X1))|X3!=k1_setfam_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk2_0,esk5_3(esk3_0,k1_setfam_1(esk3_0),X1))|r2_hidden(esk2_0,k8_setfam_1(esk1_0,esk3_0))|r2_hidden(X1,k1_setfam_1(esk3_0))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, negated_conjecture, (k6_setfam_1(esk1_0,esk3_0)=k8_setfam_1(esk1_0,esk3_0)|esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21, negated_conjecture, (k6_setfam_1(esk1_0,esk3_0)=k1_setfam_1(esk3_0)), inference(spm,[status(thm)],[c_0_17, c_0_16])).
cnf(c_0_22, plain, (r2_hidden(X1,X3)|X4=k1_xboole_0|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X2!=k1_setfam_1(X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_23, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk2_0,k8_setfam_1(esk1_0,esk3_0))|r2_hidden(esk2_0,k1_setfam_1(esk3_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, negated_conjecture, (k8_setfam_1(esk1_0,esk3_0)=k1_setfam_1(esk3_0)|esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~r2_hidden(esk2_0,k8_setfam_1(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26, negated_conjecture, (~r2_hidden(esk2_0,esk4_0)|~r2_hidden(esk2_0,k8_setfam_1(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_27, plain, ![X12, X13]:((k2_zfmisc_1(X12,X13)!=k1_xboole_0|(X12=k1_xboole_0|X13=k1_xboole_0))&((X12!=k1_xboole_0|k2_zfmisc_1(X12,X13)=k1_xboole_0)&(X13!=k1_xboole_0|k2_zfmisc_1(X12,X13)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_28, plain, (X1=k1_xboole_0|r2_hidden(X2,X3)|~r2_hidden(X2,k1_setfam_1(X1))|~r2_hidden(X3,X1)), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_29, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk2_0,k1_setfam_1(esk3_0))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_30, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk4_0,esk3_0)|~r2_hidden(esk2_0,k1_setfam_1(esk3_0))), inference(spm,[status(thm)],[c_0_25, c_0_24])).
cnf(c_0_31, negated_conjecture, (esk3_0=k1_xboole_0|~r2_hidden(esk2_0,k1_setfam_1(esk3_0))|~r2_hidden(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_26, c_0_24])).
fof(c_0_32, plain, ![X18]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X18)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_33, plain, ![X10, X11]:~r2_hidden(X10,k2_zfmisc_1(X10,X11)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_34, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk2_0,X1)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk4_0,esk3_0)), inference(spm,[status(thm)],[c_0_30, c_0_29])).
cnf(c_0_37, negated_conjecture, (esk3_0=k1_xboole_0|~r2_hidden(esk2_0,esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_23]), c_0_31])).
cnf(c_0_38, plain, (k8_setfam_1(X2,X1)=X2|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_39, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_42, negated_conjecture, (esk3_0=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_43, plain, (k8_setfam_1(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_45, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_42]), c_0_42]), c_0_43]), c_0_44])]), c_0_45]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 32
# Proof object formula steps           : 15
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 44
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 14
# Generated clauses                    : 43
# ...of the previous two non-trivial   : 43
# Contextual simplify-reflections      : 2
# Paramodulations                      : 34
# Factorizations                       : 0
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 27
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 16
# ...number of literals in the above   : 78
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 93
# Rec. Clause-clause subsumption calls : 50
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 28
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1955

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28684 pages
