# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d15_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((~(X2=k1_xboole_0)=>(m1_orders_2(X3,X1,X2)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&(r2_hidden(X4,X2)&X3=k3_orders_2(X1,X2,X4)))))&(X2=k1_xboole_0=>(m1_orders_2(X3,X1,X2)<=>X3=k1_xboole_0)))))), file('data/MPTP/proofs/b23e3e757f9e4a8db21bf3e221772ebe.E_input', d15_orders_2)).
fof(dt_m1_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>![X3]:(m1_orders_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/b23e3e757f9e4a8db21bf3e221772ebe.E_input', dt_m1_orders_2)).
fof(t57_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X2,k3_orders_2(X1,X4,X3))<=>(r2_orders_2(X1,X2,X3)&r2_hidden(X2,X4))))))), file('data/MPTP/proofs/b23e3e757f9e4a8db21bf3e221772ebe.E_input', t57_orders_2)).
fof(t67_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_orders_2(X3,X1,X2)=>r1_tarski(X3,X2)))), file('data/MPTP/proofs/b23e3e757f9e4a8db21bf3e221772ebe.E_input', t67_orders_2)).
fof(t7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/b23e3e757f9e4a8db21bf3e221772ebe.E_input', t7_subset_1)).
fof(c_0_5, plain, ![X5, X6, X7, X9]:((((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~m1_orders_2(X7,X5,X6)|X6=k1_xboole_0|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&((r2_hidden(esk1_3(X5,X6,X7),X6)|~m1_orders_2(X7,X5,X6)|X6=k1_xboole_0|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(X7=k3_orders_2(X5,X6,esk1_3(X5,X6,X7))|~m1_orders_2(X7,X5,X6)|X6=k1_xboole_0|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))))&(~m1_subset_1(X9,u1_struct_0(X5))|(~r2_hidden(X9,X6)|X7!=k3_orders_2(X5,X6,X9))|m1_orders_2(X7,X5,X6)|X6=k1_xboole_0|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))&((~m1_orders_2(X7,X5,X6)|X7=k1_xboole_0|X6!=k1_xboole_0|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5))))))&(X7!=k1_xboole_0|m1_orders_2(X7,X5,X6)|X6!=k1_xboole_0|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X5)))|(v2_struct_0(X5)|(~v3_orders_2(X5)|(~v4_orders_2(X5)|(~v5_orders_2(X5)|~l1_orders_2(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d15_orders_2])])])])])])).
fof(c_0_6, plain, ![X10, X11, X12]:(v2_struct_0(X10)|(~v3_orders_2(X10)|(~v4_orders_2(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10))))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(~m1_orders_2(X12,X10,X11)|m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_orders_2])])])])).
fof(c_0_7, plain, ![X13, X14, X15, X16]:(((r2_orders_2(X13,X14,X15)|~r2_hidden(X14,k3_orders_2(X13,X16,X15))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v3_orders_2(X13)|(~v4_orders_2(X13)|(~v5_orders_2(X13)|~l1_orders_2(X13))))))&(r2_hidden(X14,X16)|~r2_hidden(X14,k3_orders_2(X13,X16,X15))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v3_orders_2(X13)|(~v4_orders_2(X13)|(~v5_orders_2(X13)|~l1_orders_2(X13)))))))&(~r2_orders_2(X13,X14,X15)|~r2_hidden(X14,X16)|r2_hidden(X14,k3_orders_2(X13,X16,X15))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X13)))|~m1_subset_1(X15,u1_struct_0(X13))|~m1_subset_1(X14,u1_struct_0(X13))|(v2_struct_0(X13)|(~v3_orders_2(X13)|(~v4_orders_2(X13)|(~v5_orders_2(X13)|~l1_orders_2(X13))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t57_orders_2])])])])])).
cnf(c_0_8, plain, (X1=k3_orders_2(X2,X3,esk1_3(X2,X3,X1))|X3=k1_xboole_0|v2_struct_0(X2)|~m1_orders_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_orders_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|X2=k1_xboole_0|v2_struct_0(X1)|~m1_orders_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_11, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_orders_2(X3,X1,X2)=>r1_tarski(X3,X2))))), inference(assume_negation,[status(cth)],[t67_orders_2])).
cnf(c_0_12, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~r2_hidden(X1,k3_orders_2(X3,X2,X4))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~v3_orders_2(X3)|~v4_orders_2(X3)|~v5_orders_2(X3)|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (k3_orders_2(X1,X2,esk1_3(X1,X2,X3))=X3|X2=k1_xboole_0|v2_struct_0(X1)|~m1_orders_2(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14, plain, (X1=k1_xboole_0|m1_subset_1(esk1_3(X2,X1,X3),u1_struct_0(X2))|v2_struct_0(X2)|~m1_orders_2(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_10, c_0_9])).
fof(c_0_15, plain, ![X20, X21, X22]:((m1_subset_1(esk5_3(X20,X21,X22),X20)|r1_tarski(X21,X22)|~m1_subset_1(X22,k1_zfmisc_1(X20))|~m1_subset_1(X21,k1_zfmisc_1(X20)))&((r2_hidden(esk5_3(X20,X21,X22),X21)|r1_tarski(X21,X22)|~m1_subset_1(X22,k1_zfmisc_1(X20))|~m1_subset_1(X21,k1_zfmisc_1(X20)))&(~r2_hidden(esk5_3(X20,X21,X22),X22)|r1_tarski(X21,X22)|~m1_subset_1(X22,k1_zfmisc_1(X20))|~m1_subset_1(X21,k1_zfmisc_1(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_subset_1])])])])])).
fof(c_0_16, negated_conjecture, ((~v2_struct_0(esk2_0)&(v3_orders_2(esk2_0)&(v4_orders_2(esk2_0)&(v5_orders_2(esk2_0)&l1_orders_2(esk2_0)))))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(m1_orders_2(esk4_0,esk2_0,esk3_0)&~r1_tarski(esk4_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_11])])])])).
cnf(c_0_17, plain, (X1=k1_xboole_0|r2_hidden(X2,X1)|v2_struct_0(X3)|~r2_hidden(X2,X4)|~m1_orders_2(X4,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,u1_struct_0(X3))|~l1_orders_2(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
cnf(c_0_18, plain, (m1_subset_1(esk5_3(X1,X2,X3),X1)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (r1_tarski(X2,X3)|~r2_hidden(esk5_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (r2_hidden(esk5_3(X1,X2,X3),X2)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (m1_orders_2(esk4_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (v5_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, negated_conjecture, (v4_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, negated_conjecture, (v3_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, plain, (X1=k1_xboole_0|v2_struct_0(X2)|~m1_orders_2(X1,X2,X3)|X3!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29, plain, (X1=k1_xboole_0|r1_tarski(X2,X3)|r2_hidden(esk5_3(u1_struct_0(X4),X2,X3),X1)|v2_struct_0(X4)|~r2_hidden(esk5_3(u1_struct_0(X4),X2,X3),X5)|~m1_orders_2(X5,X4,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~l1_orders_2(X4)|~v5_orders_2(X4)|~v4_orders_2(X4)|~v3_orders_2(X4)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_30, plain, (r1_tarski(X1,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_9, c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_32, plain, (X1=k1_xboole_0|v2_struct_0(X2)|X3!=k1_xboole_0|~m1_orders_2(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_28, c_0_9])).
cnf(c_0_33, plain, (X1=k1_xboole_0|r1_tarski(X2,X3)|r2_hidden(esk5_3(u1_struct_0(X4),X2,X3),X1)|v2_struct_0(X4)|~m1_orders_2(X2,X4,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))|~l1_orders_2(X4)|~v5_orders_2(X4)|~v4_orders_2(X4)|~v3_orders_2(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_20]), c_0_9])).
cnf(c_0_34, negated_conjecture, (r1_tarski(esk4_0,esk4_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, negated_conjecture, (esk4_0=k1_xboole_0|esk3_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_36, negated_conjecture, (esk3_0=k1_xboole_0|r1_tarski(esk4_0,X1)|r2_hidden(esk5_3(u1_struct_0(esk2_0),esk4_0,X1),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25]), c_0_26])]), c_0_27])).
cnf(c_0_37, negated_conjecture, (~r1_tarski(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_38, negated_conjecture, (r1_tarski(k1_xboole_0,k1_xboole_0)|esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_39, negated_conjecture, (esk3_0=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_36]), c_0_22]), c_0_31])]), c_0_37])).
cnf(c_0_40, negated_conjecture, (esk3_0!=k1_xboole_0|~r1_tarski(k1_xboole_0,esk3_0)), inference(spm,[status(thm)],[c_0_37, c_0_35])).
cnf(c_0_41, negated_conjecture, (r1_tarski(k1_xboole_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39])])).
cnf(c_0_42, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_39]), c_0_39])]), c_0_41])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 32
# Proof object formula steps           : 11
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 21
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 21
# Processed clauses                    : 63
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 61
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 12
# Generated clauses                    : 36
# ...of the previous two non-trivial   : 36
# Contextual simplify-reflections      : 7
# Paramodulations                      : 34
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 28
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 15
# ...number of literals in the above   : 112
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 454
# Rec. Clause-clause subsumption calls : 33
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 27
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4001

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 29020 pages
