# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t117_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>k11_relat_1(X2,X1)=k1_tarski(k1_funct_1(X2,X1)))), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', t117_funct_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', t71_enumset1)).
fof(l44_zfmisc_1, axiom, ![X1, X2]:~((~(X1=k1_tarski(X2))&(~(X1=k1_xboole_0)&![X3]:~((r2_hidden(X3,X1)&~(X3=X2)))))), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', l44_zfmisc_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', t3_xboole_0)).
fof(t151_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(k9_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k1_relat_1(X2),X1))), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', t151_relat_1)).
fof(t204_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k11_relat_1(X3,X1)))), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', t204_relat_1)).
fof(d16_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', d16_relat_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', d1_enumset1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/b251c88c6a82495f91d851313b0e0eea.E_input', t8_funct_1)).
fof(c_0_11, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>k11_relat_1(X2,X1)=k1_tarski(k1_funct_1(X2,X1))))), inference(assume_negation,[status(cth)],[t117_funct_1])).
fof(c_0_12, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&(r2_hidden(esk1_0,k1_relat_1(esk2_0))&k11_relat_1(esk2_0,esk1_0)!=k1_tarski(k1_funct_1(esk2_0,esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
fof(c_0_13, plain, ![X43]:k2_tarski(X43,X43)=k1_tarski(X43), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_14, plain, ![X17, X18]:k1_enumset1(X17,X17,X18)=k2_tarski(X17,X18), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_15, plain, ![X8, X9, X10]:k2_enumset1(X8,X8,X9,X10)=k1_enumset1(X8,X9,X10), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_16, plain, ![X38, X39]:((r2_hidden(esk5_2(X38,X39),X38)|X38=k1_xboole_0|X38=k1_tarski(X39))&(esk5_2(X38,X39)!=X39|X38=k1_xboole_0|X38=k1_tarski(X39))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l44_zfmisc_1])])])])])).
cnf(c_0_17, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)!=k1_tarski(k1_funct_1(esk2_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (r2_hidden(esk5_2(X1,X2),X1)|X1=k1_xboole_0|X1=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_22, plain, ![X32, X33, X35, X36, X37]:(((r2_hidden(esk4_2(X32,X33),X32)|r1_xboole_0(X32,X33))&(r2_hidden(esk4_2(X32,X33),X33)|r1_xboole_0(X32,X33)))&(~r2_hidden(X37,X35)|~r2_hidden(X37,X36)|~r1_xboole_0(X35,X36))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_23, plain, ![X19, X20]:((k9_relat_1(X20,X19)!=k1_xboole_0|r1_xboole_0(k1_relat_1(X20),X19)|~v1_relat_1(X20))&(~r1_xboole_0(k1_relat_1(X20),X19)|k9_relat_1(X20,X19)=k1_xboole_0|~v1_relat_1(X20))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t151_relat_1])])])).
fof(c_0_24, plain, ![X14, X15, X16]:((~r2_hidden(k4_tarski(X14,X15),X16)|r2_hidden(X15,k11_relat_1(X16,X14))|~v1_relat_1(X16))&(~r2_hidden(X15,k11_relat_1(X16,X14))|r2_hidden(k4_tarski(X14,X15),X16)|~v1_relat_1(X16))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t204_relat_1])])])).
cnf(c_0_25, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)!=k2_enumset1(k1_funct_1(esk2_0,esk1_0),k1_funct_1(esk2_0,esk1_0),k1_funct_1(esk2_0,esk1_0),k1_funct_1(esk2_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_26, plain, (X1=k1_xboole_0|X1=k2_enumset1(X2,X2,X2,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_27, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r1_xboole_0(k1_relat_1(X1),X2)|k9_relat_1(X1,X2)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_29, plain, ![X41, X42]:(~v1_relat_1(X41)|k11_relat_1(X41,X42)=k9_relat_1(X41,k1_tarski(X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_1])])])).
fof(c_0_30, plain, ![X21, X22, X23, X24, X25, X26, X27, X28, X29, X30]:(((~r2_hidden(X25,X24)|(X25=X21|(X25=X22|X25=X23))|X24!=k1_enumset1(X21,X22,X23))&((X26!=X21|r2_hidden(X26,X24)|X24!=k1_enumset1(X21,X22,X23))&((X26!=X22|r2_hidden(X26,X24)|X24!=k1_enumset1(X21,X22,X23))&(X26!=X23|r2_hidden(X26,X24)|X24!=k1_enumset1(X21,X22,X23)))))&(((esk3_4(X27,X28,X29,X30)!=X27|~r2_hidden(esk3_4(X27,X28,X29,X30),X30)|X30=k1_enumset1(X27,X28,X29))&((esk3_4(X27,X28,X29,X30)!=X28|~r2_hidden(esk3_4(X27,X28,X29,X30),X30)|X30=k1_enumset1(X27,X28,X29))&(esk3_4(X27,X28,X29,X30)!=X29|~r2_hidden(esk3_4(X27,X28,X29,X30),X30)|X30=k1_enumset1(X27,X28,X29))))&(r2_hidden(esk3_4(X27,X28,X29,X30),X30)|(esk3_4(X27,X28,X29,X30)=X27|(esk3_4(X27,X28,X29,X30)=X28|esk3_4(X27,X28,X29,X30)=X29))|X30=k1_enumset1(X27,X28,X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
fof(c_0_31, plain, ![X11, X12, X13]:(((r2_hidden(X11,k1_relat_1(X13))|~r2_hidden(k4_tarski(X11,X12),X13)|(~v1_relat_1(X13)|~v1_funct_1(X13)))&(X12=k1_funct_1(X13,X11)|~r2_hidden(k4_tarski(X11,X12),X13)|(~v1_relat_1(X13)|~v1_funct_1(X13))))&(~r2_hidden(X11,k1_relat_1(X13))|X12!=k1_funct_1(X13,X11)|r2_hidden(k4_tarski(X11,X12),X13)|(~v1_relat_1(X13)|~v1_funct_1(X13)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,k11_relat_1(X2,X3))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)=k1_xboole_0|r2_hidden(esk5_2(k11_relat_1(esk2_0,esk1_0),k1_funct_1(esk2_0,esk1_0)),k11_relat_1(esk2_0,esk1_0))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26])])).
cnf(c_0_34, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_35, plain, (k9_relat_1(X1,X2)!=k1_xboole_0|~r2_hidden(X3,k1_relat_1(X1))|~r2_hidden(X3,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X2,X5)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|esk5_2(X1,X2)!=X2), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_40, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)=k1_xboole_0|r2_hidden(k4_tarski(esk1_0,esk5_2(k11_relat_1(esk2_0,esk1_0),k1_funct_1(esk2_0,esk1_0))),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])])).
cnf(c_0_42, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_43, negated_conjecture, (k9_relat_1(esk2_0,X1)!=k1_xboole_0|~r2_hidden(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_34])])).
cnf(c_0_44, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k2_enumset1(X2,X2,X2,X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_45, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X4,X4,X2,X5)), inference(rw,[status(thm)],[c_0_38, c_0_20])).
cnf(c_0_46, plain, (X1=k1_xboole_0|X1=k2_enumset1(X2,X2,X2,X2)|esk5_2(X1,X2)!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_47, negated_conjecture, (esk5_2(k11_relat_1(esk2_0,esk1_0),k1_funct_1(esk2_0,esk1_0))=k1_funct_1(esk2_0,esk1_0)|k11_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_34])])).
cnf(c_0_48, negated_conjecture, (k11_relat_1(esk2_0,X1)!=k1_xboole_0|~r2_hidden(esk1_0,k2_enumset1(X1,X1,X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_34])])).
cnf(c_0_49, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X1,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_45])])).
cnf(c_0_50, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_25])).
cnf(c_0_51, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 29
# Proof object formula steps           : 23
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 25
# Processed clauses                    : 65
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 64
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 4
# Generated clauses                    : 160
# ...of the previous two non-trivial   : 141
# Contextual simplify-reflections      : 0
# Paramodulations                      : 145
# Factorizations                       : 1
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 120
# ...number of literals in the above   : 591
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 110
# Rec. Clause-clause subsumption calls : 62
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3889

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28632 pages
