# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/b2b24afcd21044e79878e432b7629a0e.E_input', d4_zfmisc_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/b2b24afcd21044e79878e432b7629a0e.E_input', d3_zfmisc_1)).
fof(d4_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4), file('data/MPTP/proofs/b2b24afcd21044e79878e432b7629a0e.E_input', d4_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/b2b24afcd21044e79878e432b7629a0e.E_input', d3_mcart_1)).
fof(t81_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X9)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6))))))), file('data/MPTP/proofs/b2b24afcd21044e79878e432b7629a0e.E_input', t81_mcart_1)).
fof(t59_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))&?[X6, X7, X8, X9]:(X5=k4_mcart_1(X6,X7,X8,X9)&~((k8_mcart_1(X1,X2,X3,X4,X5)=X6&(k9_mcart_1(X1,X2,X3,X4,X5)=X7&(k10_mcart_1(X1,X2,X3,X4,X5)=X8&k11_mcart_1(X1,X2,X3,X4,X5)=X9))))))))))), file('data/MPTP/proofs/b2b24afcd21044e79878e432b7629a0e.E_input', t59_mcart_1)).
fof(l68_mcart_1, axiom, ![X1, X2, X3, X4]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&(~(X4=k1_xboole_0)&?[X5]:(m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))&![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>![X9]:(m1_subset_1(X9,X4)=>~(X5=k4_mcart_1(X6,X7,X8,X9)))))))))))), file('data/MPTP/proofs/b2b24afcd21044e79878e432b7629a0e.E_input', l68_mcart_1)).
fof(c_0_7, plain, ![X11, X12, X13, X14]:k4_zfmisc_1(X11,X12,X13,X14)=k2_zfmisc_1(k3_zfmisc_1(X11,X12,X13),X14), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
fof(c_0_8, plain, ![X15, X16, X17]:k3_zfmisc_1(X15,X16,X17)=k2_zfmisc_1(k2_zfmisc_1(X15,X16),X17), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_9, plain, ![X18, X19, X20, X21]:k4_mcart_1(X18,X19,X20,X21)=k4_tarski(k3_mcart_1(X18,X19,X20),X21), inference(variable_rename,[status(thm)],[d4_mcart_1])).
fof(c_0_10, plain, ![X22, X23, X24]:k3_mcart_1(X22,X23,X24)=k4_tarski(k4_tarski(X22,X23),X24), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_11, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(m1_subset_1(X6,k4_zfmisc_1(X1,X2,X3,X4))=>(![X7]:(m1_subset_1(X7,X1)=>![X8]:(m1_subset_1(X8,X2)=>![X9]:(m1_subset_1(X9,X3)=>![X10]:(m1_subset_1(X10,X4)=>(X6=k4_mcart_1(X7,X8,X9,X10)=>X5=X9)))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|(X4=k1_xboole_0|X5=k10_mcart_1(X1,X2,X3,X4,X6)))))))), inference(assume_negation,[status(cth)],[t81_mcart_1])).
fof(c_0_12, plain, ![X35, X36, X37, X38, X39, X40, X41, X42, X43]:((k8_mcart_1(X35,X36,X37,X38,X39)=X40|X39!=k4_mcart_1(X40,X41,X42,X43)|~m1_subset_1(X39,k4_zfmisc_1(X35,X36,X37,X38))|X38=k1_xboole_0|X37=k1_xboole_0|X36=k1_xboole_0|X35=k1_xboole_0)&((k9_mcart_1(X35,X36,X37,X38,X39)=X41|X39!=k4_mcart_1(X40,X41,X42,X43)|~m1_subset_1(X39,k4_zfmisc_1(X35,X36,X37,X38))|X38=k1_xboole_0|X37=k1_xboole_0|X36=k1_xboole_0|X35=k1_xboole_0)&((k10_mcart_1(X35,X36,X37,X38,X39)=X42|X39!=k4_mcart_1(X40,X41,X42,X43)|~m1_subset_1(X39,k4_zfmisc_1(X35,X36,X37,X38))|X38=k1_xboole_0|X37=k1_xboole_0|X36=k1_xboole_0|X35=k1_xboole_0)&(k11_mcart_1(X35,X36,X37,X38,X39)=X43|X39!=k4_mcart_1(X40,X41,X42,X43)|~m1_subset_1(X39,k4_zfmisc_1(X35,X36,X37,X38))|X38=k1_xboole_0|X37=k1_xboole_0|X36=k1_xboole_0|X35=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t59_mcart_1])])])])])).
cnf(c_0_13, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k3_mcart_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, negated_conjecture, ![X31, X32, X33, X34]:(m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))&((~m1_subset_1(X31,esk1_0)|(~m1_subset_1(X32,esk2_0)|(~m1_subset_1(X33,esk3_0)|(~m1_subset_1(X34,esk4_0)|(esk6_0!=k4_mcart_1(X31,X32,X33,X34)|esk5_0=X33)))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&esk5_0!=k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
fof(c_0_18, plain, ![X44, X45, X46, X47, X48]:((m1_subset_1(esk7_5(X44,X45,X46,X47,X48),X44)|~m1_subset_1(X48,k4_zfmisc_1(X44,X45,X46,X47))|X47=k1_xboole_0|X46=k1_xboole_0|X45=k1_xboole_0|X44=k1_xboole_0)&((m1_subset_1(esk8_5(X44,X45,X46,X47,X48),X45)|~m1_subset_1(X48,k4_zfmisc_1(X44,X45,X46,X47))|X47=k1_xboole_0|X46=k1_xboole_0|X45=k1_xboole_0|X44=k1_xboole_0)&((m1_subset_1(esk9_5(X44,X45,X46,X47,X48),X46)|~m1_subset_1(X48,k4_zfmisc_1(X44,X45,X46,X47))|X47=k1_xboole_0|X46=k1_xboole_0|X45=k1_xboole_0|X44=k1_xboole_0)&((m1_subset_1(esk10_5(X44,X45,X46,X47,X48),X47)|~m1_subset_1(X48,k4_zfmisc_1(X44,X45,X46,X47))|X47=k1_xboole_0|X46=k1_xboole_0|X45=k1_xboole_0|X44=k1_xboole_0)&(X48=k4_mcart_1(esk7_5(X44,X45,X46,X47,X48),esk8_5(X44,X45,X46,X47,X48),esk9_5(X44,X45,X46,X47,X48),esk10_5(X44,X45,X46,X47,X48))|~m1_subset_1(X48,k4_zfmisc_1(X44,X45,X46,X47))|X47=k1_xboole_0|X46=k1_xboole_0|X45=k1_xboole_0|X44=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l68_mcart_1])])])])])])])).
cnf(c_0_19, plain, (k10_mcart_1(X1,X2,X3,X4,X5)=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X5!=k4_mcart_1(X7,X8,X6,X9)|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_21, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22, negated_conjecture, (esk5_0=X3|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X4,esk4_0)|esk6_0!=k4_mcart_1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (X1=k4_mcart_1(esk7_5(X2,X3,X4,X5,X1),esk8_5(X2,X3,X4,X5,X1),esk9_5(X2,X3,X4,X5,X1),esk10_5(X2,X3,X4,X5,X1))|X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X1,k4_zfmisc_1(X2,X3,X4,X5))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k10_mcart_1(X1,X2,X3,X4,X5)=X6|X5!=k4_tarski(k4_tarski(k4_tarski(X7,X8),X6),X9)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_25, negated_conjecture, (esk5_0=X3|esk6_0!=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)|~m1_subset_1(X4,esk4_0)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_22, c_0_21])).
cnf(c_0_26, plain, (X5=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k4_tarski(k4_tarski(k4_tarski(esk7_5(X2,X3,X4,X5,X1),esk8_5(X2,X3,X4,X5,X1)),esk9_5(X2,X3,X4,X5,X1)),esk10_5(X2,X3,X4,X5,X1))|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4),X5))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_20]), c_0_21])).
cnf(c_0_27, plain, (m1_subset_1(esk10_5(X1,X2,X3,X4,X5),X4)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k10_mcart_1(X1,X2,X3,X4,k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8))=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(k4_tarski(X5,X6),X7),X8),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk6_0,k4_zfmisc_1(esk1_0,esk2_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, negated_conjecture, (esk9_5(X1,X2,X3,X4,esk6_0)=esk5_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))|~m1_subset_1(esk10_5(X1,X2,X3,X4,esk6_0),esk4_0)|~m1_subset_1(esk9_5(X1,X2,X3,X4,esk6_0),esk3_0)|~m1_subset_1(esk8_5(X1,X2,X3,X4,esk6_0),esk2_0)|~m1_subset_1(esk7_5(X1,X2,X3,X4,esk6_0),esk1_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26])])).
cnf(c_0_31, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk10_5(X1,X2,X3,X4,X5),X4)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_27, c_0_20])).
cnf(c_0_32, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_33, plain, (m1_subset_1(esk9_5(X1,X2,X3,X4,X5),X3)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_34, plain, (k10_mcart_1(X1,X2,X3,X4,X5)=esk9_5(X6,X7,X8,X9,X5)|X6=k1_xboole_0|X7=k1_xboole_0|X8=k1_xboole_0|X9=k1_xboole_0|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X6,X7),X8),X9))), inference(spm,[status(thm)],[c_0_28, c_0_26])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0),esk4_0))), inference(rw,[status(thm)],[c_0_29, c_0_20])).
cnf(c_0_36, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_38, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_39, negated_conjecture, (esk9_5(X1,X2,X3,esk4_0,esk6_0)=esk5_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),esk4_0))|~m1_subset_1(esk9_5(X1,X2,X3,esk4_0,esk6_0),esk3_0)|~m1_subset_1(esk8_5(X1,X2,X3,esk4_0,esk6_0),esk2_0)|~m1_subset_1(esk7_5(X1,X2,X3,esk4_0,esk6_0),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_40, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk9_5(X1,X2,X3,X4,X5),X3)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_33, c_0_20])).
cnf(c_0_41, plain, (m1_subset_1(esk8_5(X1,X2,X3,X4,X5),X2)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_42, negated_conjecture, (k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)=esk9_5(X1,X2,X3,X4,esk6_0)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_32]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_43, negated_conjecture, (esk9_5(X1,X2,esk3_0,esk4_0,esk6_0)=esk5_0|X1=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),esk3_0),esk4_0))|~m1_subset_1(esk8_5(X1,X2,esk3_0,esk4_0,esk6_0),esk2_0)|~m1_subset_1(esk7_5(X1,X2,esk3_0,esk4_0,esk6_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_36]), c_0_32])).
cnf(c_0_44, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk8_5(X1,X2,X3,X4,X5),X2)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_41, c_0_20])).
cnf(c_0_45, plain, (m1_subset_1(esk7_5(X1,X2,X3,X4,X5),X1)|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X5,k4_zfmisc_1(X1,X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_46, negated_conjecture, (esk5_0!=k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_47, negated_conjecture, (k10_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)=esk9_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_35]), c_0_32]), c_0_36]), c_0_37]), c_0_38])).
cnf(c_0_48, negated_conjecture, (esk9_5(X1,esk2_0,esk3_0,esk4_0,esk6_0)=esk5_0|X1=k1_xboole_0|~m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,esk2_0),esk3_0),esk4_0))|~m1_subset_1(esk7_5(X1,esk2_0,esk3_0,esk4_0,esk6_0),esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_37]), c_0_32]), c_0_36])).
cnf(c_0_49, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk7_5(X1,X2,X3,X4,X5),X1)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4))), inference(rw,[status(thm)],[c_0_45, c_0_20])).
cnf(c_0_50, negated_conjecture, (esk9_5(esk1_0,esk2_0,esk3_0,esk4_0,esk6_0)!=esk5_0), inference(rw,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_35])]), c_0_50]), c_0_38]), c_0_32]), c_0_36]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 37
# Proof object formula steps           : 15
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 16
# Processed clauses                    : 53
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 41
# ...of the previous two non-trivial   : 42
# Contextual simplify-reflections      : 0
# Paramodulations                      : 36
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 30
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 21
# ...number of literals in the above   : 288
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 112
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 11
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3586

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 28644 pages
