# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', d10_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t3_boole)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t100_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t28_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', d5_xboole_0)).
fof(t65_zfmisc_1, conjecture, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t65_zfmisc_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t48_xboole_1)).
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t1_xboole_0)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t72_enumset1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', d4_xboole_0)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', t83_xboole_1)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', l27_zfmisc_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', l32_xboole_1)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/b2eda2dc50fa4ecd989f5f9273dfbc3b.E_input', d1_enumset1)).
fof(c_0_17, plain, ![X19, X20]:(((r1_tarski(X19,X20)|X19!=X20)&(r1_tarski(X20,X19)|X19!=X20))&(~r1_tarski(X19,X20)|~r1_tarski(X20,X19)|X19=X20)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_18, plain, ![X68]:k4_xboole_0(X68,k1_xboole_0)=X68, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_19, plain, ![X51, X52]:k4_xboole_0(X51,X52)=k5_xboole_0(X51,k3_xboole_0(X51,X52)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_20, plain, ![X8, X9]:(~r1_tarski(X8,X9)|k3_xboole_0(X8,X9)=X8), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_21, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X39, X40, X41, X42, X43, X44, X45, X46]:((((r2_hidden(X42,X39)|~r2_hidden(X42,X41)|X41!=k4_xboole_0(X39,X40))&(~r2_hidden(X42,X40)|~r2_hidden(X42,X41)|X41!=k4_xboole_0(X39,X40)))&(~r2_hidden(X43,X39)|r2_hidden(X43,X40)|r2_hidden(X43,X41)|X41!=k4_xboole_0(X39,X40)))&((~r2_hidden(esk5_3(X44,X45,X46),X46)|(~r2_hidden(esk5_3(X44,X45,X46),X44)|r2_hidden(esk5_3(X44,X45,X46),X45))|X46=k4_xboole_0(X44,X45))&((r2_hidden(esk5_3(X44,X45,X46),X44)|r2_hidden(esk5_3(X44,X45,X46),X46)|X46=k4_xboole_0(X44,X45))&(~r2_hidden(esk5_3(X44,X45,X46),X45)|r2_hidden(esk5_3(X44,X45,X46),X46)|X46=k4_xboole_0(X44,X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_23, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_21])).
fof(c_0_27, negated_conjecture, ~(![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1)))), inference(assume_negation,[status(cth)],[t65_zfmisc_1])).
fof(c_0_28, plain, ![X49, X50]:k4_xboole_0(X49,k4_xboole_0(X49,X50))=k3_xboole_0(X49,X50), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_29, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_30, plain, ![X53, X54, X55]:(((~r2_hidden(X53,X54)|~r2_hidden(X53,X55)|~r2_hidden(X53,k5_xboole_0(X54,X55)))&(r2_hidden(X53,X54)|r2_hidden(X53,X55)|~r2_hidden(X53,k5_xboole_0(X54,X55))))&((~r2_hidden(X53,X54)|r2_hidden(X53,X55)|r2_hidden(X53,k5_xboole_0(X54,X55)))&(~r2_hidden(X53,X55)|r2_hidden(X53,X54)|r2_hidden(X53,k5_xboole_0(X54,X55))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
cnf(c_0_31, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_33, negated_conjecture, ((k4_xboole_0(esk1_0,k1_tarski(esk2_0))!=esk1_0|r2_hidden(esk2_0,esk1_0))&(k4_xboole_0(esk1_0,k1_tarski(esk2_0))=esk1_0|~r2_hidden(esk2_0,esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_27])])])])).
fof(c_0_34, plain, ![X48]:k2_tarski(X48,X48)=k1_tarski(X48), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_35, plain, ![X21, X22]:k1_enumset1(X21,X21,X22)=k2_tarski(X21,X22), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_36, plain, ![X36, X37, X38]:k2_enumset1(X36,X36,X37,X38)=k1_enumset1(X36,X37,X38), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_37, plain, ![X64, X65, X66, X67]:k3_enumset1(X64,X64,X65,X66,X67)=k2_enumset1(X64,X65,X66,X67), inference(variable_rename,[status(thm)],[t72_enumset1])).
cnf(c_0_38, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_29, c_0_24])).
fof(c_0_40, plain, ![X10, X11, X12, X13, X14, X15, X16, X17]:((((r2_hidden(X13,X10)|~r2_hidden(X13,X12)|X12!=k3_xboole_0(X10,X11))&(r2_hidden(X13,X11)|~r2_hidden(X13,X12)|X12!=k3_xboole_0(X10,X11)))&(~r2_hidden(X14,X10)|~r2_hidden(X14,X11)|r2_hidden(X14,X12)|X12!=k3_xboole_0(X10,X11)))&((~r2_hidden(esk3_3(X15,X16,X17),X17)|(~r2_hidden(esk3_3(X15,X16,X17),X15)|~r2_hidden(esk3_3(X15,X16,X17),X16))|X17=k3_xboole_0(X15,X16))&((r2_hidden(esk3_3(X15,X16,X17),X15)|r2_hidden(esk3_3(X15,X16,X17),X17)|X17=k3_xboole_0(X15,X16))&(r2_hidden(esk3_3(X15,X16,X17),X16)|r2_hidden(esk3_3(X15,X16,X17),X17)|X17=k3_xboole_0(X15,X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (k5_xboole_0(k1_xboole_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_43, plain, ![X34, X35]:((~r1_xboole_0(X34,X35)|k4_xboole_0(X34,X35)=X34)&(k4_xboole_0(X34,X35)!=X34|r1_xboole_0(X34,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_44, plain, ![X56, X57]:(r2_hidden(X56,X57)|r1_xboole_0(k1_tarski(X56),X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
fof(c_0_45, plain, ![X58, X59]:((k4_xboole_0(X58,X59)!=k1_xboole_0|r1_tarski(X58,X59))&(~r1_tarski(X58,X59)|k4_xboole_0(X58,X59)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_46, negated_conjecture, (k4_xboole_0(esk1_0,k1_tarski(esk2_0))=esk1_0|~r2_hidden(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_47, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_49, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_24]), c_0_24])).
fof(c_0_52, plain, ![X23, X24, X25, X26, X27, X28, X29, X30, X31, X32]:(((~r2_hidden(X27,X26)|(X27=X23|(X27=X24|X27=X25))|X26!=k1_enumset1(X23,X24,X25))&((X28!=X23|r2_hidden(X28,X26)|X26!=k1_enumset1(X23,X24,X25))&((X28!=X24|r2_hidden(X28,X26)|X26!=k1_enumset1(X23,X24,X25))&(X28!=X25|r2_hidden(X28,X26)|X26!=k1_enumset1(X23,X24,X25)))))&(((esk4_4(X29,X30,X31,X32)!=X29|~r2_hidden(esk4_4(X29,X30,X31,X32),X32)|X32=k1_enumset1(X29,X30,X31))&((esk4_4(X29,X30,X31,X32)!=X30|~r2_hidden(esk4_4(X29,X30,X31,X32),X32)|X32=k1_enumset1(X29,X30,X31))&(esk4_4(X29,X30,X31,X32)!=X31|~r2_hidden(esk4_4(X29,X30,X31,X32),X32)|X32=k1_enumset1(X29,X30,X31))))&(r2_hidden(esk4_4(X29,X30,X31,X32),X32)|(esk4_4(X29,X30,X31,X32)=X29|(esk4_4(X29,X30,X31,X32)=X30|esk4_4(X29,X30,X31,X32)=X31))|X32=k1_enumset1(X29,X30,X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_53, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_54, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_55, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_56, plain, (r2_hidden(esk3_3(X1,X2,X3),X1)|r2_hidden(esk3_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_57, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_58, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_59, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_60, negated_conjecture, (k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))=esk1_0|~r2_hidden(esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_24]), c_0_49]), c_0_50])).
cnf(c_0_61, plain, (k5_xboole_0(X1,X1)=k3_xboole_0(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_31]), c_0_32])).
cnf(c_0_62, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X5,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63, plain, (X1=k3_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X3,X4)))|r2_hidden(esk3_3(X2,k5_xboole_0(X3,k3_xboole_0(X3,X4)),X1),X1)|~r2_hidden(esk3_3(X2,k5_xboole_0(X3,k3_xboole_0(X3,X4)),X1),X4)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_64, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk3_3(X1,X2,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_65, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=X1|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_57, c_0_24])).
cnf(c_0_66, plain, (r2_hidden(X1,X2)|r1_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_47]), c_0_48]), c_0_49]), c_0_50])).
cnf(c_0_67, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k1_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_59, c_0_24])).
cnf(c_0_68, negated_conjecture, (k3_xboole_0(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))=k3_xboole_0(esk1_0,k1_xboole_0)|~r2_hidden(esk2_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_60]), c_0_32]), c_0_61])).
cnf(c_0_69, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k3_enumset1(X4,X4,X4,X5,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_49]), c_0_50])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|k4_xboole_0(esk1_0,k1_tarski(esk2_0))!=esk1_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_71, plain, (k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_55])).
cnf(c_0_72, plain, (k5_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),k3_xboole_0(k3_enumset1(X1,X1,X1,X1,X1),X2))=k3_enumset1(X1,X1,X1,X1,X1)|r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_73, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_26]), c_0_32]), c_0_61])).
cnf(c_0_74, negated_conjecture, (~r2_hidden(X1,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|~r2_hidden(esk2_0,esk1_0)|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_68]), c_0_31])).
cnf(c_0_75, plain, (r2_hidden(X1,k3_enumset1(X2,X2,X2,X3,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_69])])).
cnf(c_0_76, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))!=esk1_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_47]), c_0_48]), c_0_24]), c_0_49]), c_0_50])).
cnf(c_0_77, plain, (k3_xboole_0(X1,k3_enumset1(X2,X2,X2,X2,X2))=k1_xboole_0|r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_78, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_31, c_0_73])).
cnf(c_0_79, negated_conjecture, (~r2_hidden(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])]), c_0_79]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 46
# Proof object formula steps           : 35
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 39
# Processed clauses                    : 7156
# ...of these trivial                  : 233
# ...subsumed                          : 6101
# ...remaining for further processing  : 822
# Other redundant clauses eliminated   : 2273
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 32
# Generated clauses                    : 125202
# ...of the previous two non-trivial   : 108161
# Contextual simplify-reflections      : 7
# Paramodulations                      : 122662
# Factorizations                       : 270
# Equation resolutions                 : 2273
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 722
#    Positive orientable unit clauses  : 56
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 662
# Current number of unprocessed clauses: 100792
# ...number of literals in the above   : 590488
# Current number of archived formulas  : 0
# Current number of archived clauses   : 91
# Clause-clause subsumption calls (NU) : 191030
# Rec. Clause-clause subsumption calls : 81899
# Non-unit clause-clause subsumptions  : 2895
# Unit Clause-clause subsumption calls : 2827
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 129
# BW rewrite match successes           : 38
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2707874

# -------------------------------------------------
# User time                : 2.161 s
# System time              : 0.030 s
# Total time               : 2.190 s
# Maximum resident set size: 91064 pages
