# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t72_zfmisc_1, axiom, ![X1, X2, X3]:(k4_xboole_0(k2_tarski(X1,X2),X3)=k2_tarski(X1,X2)<=>(~(r2_hidden(X1,X3))&~(r2_hidden(X2,X3)))), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t72_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t70_enumset1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', d2_tarski)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t65_zfmisc_1)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t9_mcart_1)).
fof(t51_mcart_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4)))))))))), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t51_mcart_1)).
fof(t29_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k3_mcart_1(X3,X4,X5))))))), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t29_mcart_1)).
fof(t48_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4)))))))), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t48_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', d3_zfmisc_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', d3_mcart_1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t152_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/b3150b71af32429f8e1b5ded25981462.E_input', t113_zfmisc_1)).
fof(c_0_12, plain, ![X6, X7, X8]:(((~r2_hidden(X6,X8)|k4_xboole_0(k2_tarski(X6,X7),X8)!=k2_tarski(X6,X7))&(~r2_hidden(X7,X8)|k4_xboole_0(k2_tarski(X6,X7),X8)!=k2_tarski(X6,X7)))&(r2_hidden(X6,X8)|r2_hidden(X7,X8)|k4_xboole_0(k2_tarski(X6,X7),X8)=k2_tarski(X6,X7))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t72_zfmisc_1])])])])).
fof(c_0_13, plain, ![X9, X10]:k1_enumset1(X9,X9,X10)=k2_tarski(X9,X10), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_14, plain, ![X11, X12, X13, X14, X15, X16, X17, X18]:(((~r2_hidden(X14,X13)|(X14=X11|X14=X12)|X13!=k2_tarski(X11,X12))&((X15!=X11|r2_hidden(X15,X13)|X13!=k2_tarski(X11,X12))&(X15!=X12|r2_hidden(X15,X13)|X13!=k2_tarski(X11,X12))))&(((esk1_3(X16,X17,X18)!=X16|~r2_hidden(esk1_3(X16,X17,X18),X18)|X18=k2_tarski(X16,X17))&(esk1_3(X16,X17,X18)!=X17|~r2_hidden(esk1_3(X16,X17,X18),X18)|X18=k2_tarski(X16,X17)))&(r2_hidden(esk1_3(X16,X17,X18),X18)|(esk1_3(X16,X17,X18)=X16|esk1_3(X16,X17,X18)=X17)|X18=k2_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_15, plain, (~r2_hidden(X1,X2)|k4_xboole_0(k2_tarski(X3,X1),X2)!=k2_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, ![X20, X21]:((k4_xboole_0(X20,k1_tarski(X21))!=X20|~r2_hidden(X21,X20))&(r2_hidden(X21,X20)|k4_xboole_0(X20,k1_tarski(X21))=X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
cnf(c_0_18, plain, (~r2_hidden(X1,X2)|k4_xboole_0(k2_tarski(X1,X3),X2)!=k2_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k4_xboole_0(k1_enumset1(X3,X3,X1),X2)!=k1_enumset1(X3,X3,X1)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16]), c_0_16])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X22, X24, X25]:((r2_hidden(esk2_1(X22),X22)|X22=k1_xboole_0)&((~r2_hidden(X24,X22)|esk2_1(X22)!=k4_tarski(X24,X25)|X22=k1_xboole_0)&(~r2_hidden(X25,X22)|esk2_1(X22)!=k4_tarski(X24,X25)|X22=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
fof(c_0_23, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(~(X4=k5_mcart_1(X1,X2,X3,X4))&(~(X4=k6_mcart_1(X1,X2,X3,X4))&~(X4=k7_mcart_1(X1,X2,X3,X4))))))))))), inference(assume_negation,[status(cth)],[t51_mcart_1])).
cnf(c_0_24, plain, (k4_xboole_0(k1_enumset1(X1,X1,X3),X2)!=k1_enumset1(X1,X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_16]), c_0_16])).
fof(c_0_25, plain, ![X30, X32, X33, X34]:((r2_hidden(esk7_1(X30),X30)|X30=k1_xboole_0)&((~r2_hidden(X32,X30)|esk7_1(X30)!=k3_mcart_1(X32,X33,X34)|X30=k1_xboole_0)&(~r2_hidden(X33,X30)|esk7_1(X30)!=k3_mcart_1(X32,X33,X34)|X30=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_mcart_1])])])])])])])).
cnf(c_0_26, plain, (X1=X4|X1=X3|X2!=k1_enumset1(X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_19, c_0_16])).
cnf(c_0_27, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X3))|~r2_hidden(X3,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_28, plain, (r2_hidden(esk2_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X35, X36, X37, X38]:(X35=k1_xboole_0|(X36=k1_xboole_0|(X37=k1_xboole_0|(~m1_subset_1(X38,k3_zfmisc_1(X35,X36,X37))|X38=k3_mcart_1(k5_mcart_1(X35,X36,X37,X38),k6_mcart_1(X35,X36,X37,X38),k7_mcart_1(X35,X36,X37,X38)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_mcart_1])])])])).
fof(c_0_30, plain, ![X39, X40, X41]:k3_zfmisc_1(X39,X40,X41)=k2_zfmisc_1(k2_zfmisc_1(X39,X40),X41), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_31, plain, ![X42, X43, X44]:k3_mcart_1(X42,X43,X44)=k4_tarski(k4_tarski(X42,X43),X44), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_32, negated_conjecture, (esk3_0!=k1_xboole_0&(esk4_0!=k1_xboole_0&(esk5_0!=k1_xboole_0&(m1_subset_1(esk6_0,k3_zfmisc_1(esk3_0,esk4_0,esk5_0))&(esk6_0=k5_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)|(esk6_0=k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)|esk6_0=k7_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
cnf(c_0_33, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X3))|~r2_hidden(X2,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_24, c_0_21])).
cnf(c_0_34, plain, (r2_hidden(esk7_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (X1=X2|X1=X3|~r2_hidden(X1,k1_enumset1(X2,X2,X3))), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (k1_tarski(X1)=k1_xboole_0|r2_hidden(X1,k1_enumset1(X2,X2,esk2_1(k1_tarski(X1))))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_37, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_39, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk6_0,k3_zfmisc_1(esk3_0,esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (k1_tarski(X1)=k1_xboole_0|r2_hidden(X1,k1_enumset1(esk7_1(k1_tarski(X1)),esk7_1(k1_tarski(X1)),X2))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_42, plain, (esk2_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0|X1=X2), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk7_1(X2)!=k3_mcart_1(X3,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_44, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4=k4_tarski(k4_tarski(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4)),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk6_0,k2_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0),esk5_0))), inference(rw,[status(thm)],[c_0_40, c_0_38])).
cnf(c_0_46, negated_conjecture, (esk5_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_47, negated_conjecture, (esk4_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_48, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_49, plain, (esk7_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0|X1=X2), inference(spm,[status(thm)],[c_0_35, c_0_41])).
cnf(c_0_50, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk2_1(X2)!=k4_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_51, plain, (esk2_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0), inference(ef,[status(thm)],[c_0_42])).
cnf(c_0_52, plain, (X2=k1_xboole_0|esk7_1(X2)!=k4_tarski(k4_tarski(X3,X1),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_43, c_0_39])).
cnf(c_0_53, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0),k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)),k7_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0))=esk6_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_47]), c_0_48])).
cnf(c_0_54, plain, (esk7_1(k1_tarski(X1))=X1|k1_tarski(X1)=k1_xboole_0), inference(ef,[status(thm)],[c_0_49])).
cnf(c_0_55, plain, (k1_tarski(X1)=k1_xboole_0|X1!=k4_tarski(X2,X3)|~r2_hidden(X3,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_56, negated_conjecture, (X1=k1_xboole_0|esk7_1(X1)!=esk6_0|~r2_hidden(k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0),X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_57, plain, (k1_tarski(X1)=k1_xboole_0|r2_hidden(X1,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_34, c_0_54])).
cnf(c_0_58, plain, (k1_tarski(k4_tarski(X1,X2))=k1_xboole_0|~r2_hidden(X2,k1_tarski(k4_tarski(X1,X2)))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_59, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk7_1(X2)!=k3_mcart_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_60, negated_conjecture, (k1_tarski(k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0))=k1_xboole_0|esk7_1(k1_tarski(k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)))!=esk6_0), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_61, negated_conjecture, (k1_tarski(esk6_0)=k1_xboole_0|~r2_hidden(k7_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0),k1_tarski(esk6_0))), inference(spm,[status(thm)],[c_0_58, c_0_53])).
cnf(c_0_62, negated_conjecture, (esk6_0=k5_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)|esk6_0=k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)|esk6_0=k7_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_63, plain, (X2=k1_xboole_0|esk7_1(X2)!=k4_tarski(k4_tarski(X1,X3),X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_59, c_0_39])).
cnf(c_0_64, negated_conjecture, (k1_tarski(k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0))=k1_xboole_0|k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)!=esk6_0), inference(spm,[status(thm)],[c_0_60, c_0_54])).
cnf(c_0_65, negated_conjecture, (k6_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)=esk6_0|k5_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)=esk6_0|k1_tarski(esk6_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_57])).
cnf(c_0_66, negated_conjecture, (X1=k1_xboole_0|esk7_1(X1)!=esk6_0|~r2_hidden(k5_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0),X1)), inference(spm,[status(thm)],[c_0_63, c_0_53])).
cnf(c_0_67, negated_conjecture, (k5_mcart_1(esk3_0,esk4_0,esk5_0,esk6_0)=esk6_0|k1_tarski(esk6_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_68, negated_conjecture, (k1_tarski(esk6_0)=k1_xboole_0|X1=k1_xboole_0|esk7_1(X1)!=esk6_0|~r2_hidden(esk6_0,X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
fof(c_0_69, plain, ![X45, X46]:~r2_hidden(X45,k2_zfmisc_1(X45,X46)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
fof(c_0_70, plain, ![X47, X48]:((k2_zfmisc_1(X47,X48)!=k1_xboole_0|(X47=k1_xboole_0|X48=k1_xboole_0))&((X47!=k1_xboole_0|k2_zfmisc_1(X47,X48)=k1_xboole_0)&(X48!=k1_xboole_0|k2_zfmisc_1(X47,X48)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_71, negated_conjecture, (k1_tarski(esk6_0)=k1_xboole_0|k1_tarski(X1)=k1_xboole_0|X1!=esk6_0|~r2_hidden(esk6_0,k1_tarski(X1))), inference(spm,[status(thm)],[c_0_68, c_0_54])).
cnf(c_0_72, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_73, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_74, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_75, plain, (k4_xboole_0(X1,k1_tarski(X2))!=X1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_76, negated_conjecture, (k1_tarski(esk6_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_71]), c_0_57])).
cnf(c_0_77, plain, (r2_hidden(X1,X2)|r2_hidden(X3,X2)|k4_xboole_0(k2_tarski(X1,X3),X2)=k2_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_78, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_79, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_74, c_0_16])).
cnf(c_0_80, negated_conjecture, (k4_xboole_0(X1,k1_xboole_0)!=X1|~r2_hidden(esk6_0,X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_81, plain, (k4_xboole_0(k1_enumset1(X1,X1,X3),X2)=k1_enumset1(X1,X1,X3)|r2_hidden(X3,X2)|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_16]), c_0_16])).
cnf(c_0_82, plain, (~r2_hidden(X1,k1_xboole_0)), inference(er,[status(thm)],[c_0_78])).
cnf(c_0_83, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X3,X1)), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_84, negated_conjecture, (~r2_hidden(esk6_0,k1_enumset1(X1,X1,X2))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_82]), c_0_82])).
cnf(c_0_85, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[c_0_83])).
cnf(c_0_86, negated_conjecture, ($false), inference(spm,[status(thm)],[c_0_84, c_0_85]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 87
# Proof object clause steps            : 62
# Proof object formula steps           : 25
# Proof object conjectures             : 23
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 21
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 27
# Processed clauses                    : 129
# ...of these trivial                  : 1
# ...subsumed                          : 12
# ...remaining for further processing  : 116
# Other redundant clauses eliminated   : 43
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 16
# Backward-rewritten                   : 4
# Generated clauses                    : 641
# ...of the previous two non-trivial   : 564
# Contextual simplify-reflections      : 2
# Paramodulations                      : 572
# Factorizations                       : 12
# Equation resolutions                 : 57
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 67
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 56
# Current number of unprocessed clauses: 419
# ...number of literals in the above   : 1867
# Current number of archived formulas  : 0
# Current number of archived clauses   : 50
# Clause-clause subsumption calls (NU) : 528
# Rec. Clause-clause subsumption calls : 427
# Non-unit clause-clause subsumptions  : 25
# Unit Clause-clause subsumption calls : 27
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9189

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28684 pages
