# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', d10_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', t3_subset)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', t5_subset)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', dt_u1_pre_topc)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', d2_subset_1)).
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', t24_yellow_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', d1_pre_topc)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', d4_tarski)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', d3_tarski)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', t99_zfmisc_1)).
fof(redefinition_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2)), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', redefinition_k5_setfam_1)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/b3c69eefaa984ff8b9745d88d9048bac.E_input', t14_yellow_1)).
fof(c_0_12, plain, ![X5, X6]:(((r1_tarski(X5,X6)|X5!=X6)&(r1_tarski(X6,X5)|X5!=X6))&(~r1_tarski(X5,X6)|~r1_tarski(X6,X5)|X5=X6)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_13, plain, ![X7, X8]:((~m1_subset_1(X7,k1_zfmisc_1(X8))|r1_tarski(X7,X8))&(~r1_tarski(X7,X8)|m1_subset_1(X7,k1_zfmisc_1(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_14, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_15, plain, ![X9, X10, X11]:(~r2_hidden(X9,X10)|(~m1_subset_1(X10,k1_zfmisc_1(X11))|~v1_xboole_0(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_16, plain, ![X23]:(~l1_pre_topc(X23)|m1_subset_1(u1_pre_topc(X23),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X23))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
fof(c_0_17, plain, ![X24, X25]:(((~m1_subset_1(X25,X24)|r2_hidden(X25,X24)|v1_xboole_0(X24))&(~r2_hidden(X25,X24)|m1_subset_1(X25,X24)|v1_xboole_0(X24)))&((~m1_subset_1(X25,X24)|v1_xboole_0(X25)|~v1_xboole_0(X24))&(~v1_xboole_0(X25)|m1_subset_1(X25,X24)|~v1_xboole_0(X24)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_18, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_14])).
fof(c_0_20, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
cnf(c_0_21, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_25, plain, ![X28, X29, X30, X31]:(((r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|~v2_pre_topc(X28)|~l1_pre_topc(X28))&((~m1_subset_1(X29,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X28))))|(~r1_tarski(X29,u1_pre_topc(X28))|r2_hidden(k5_setfam_1(u1_struct_0(X28),X29),u1_pre_topc(X28)))|~v2_pre_topc(X28)|~l1_pre_topc(X28))&(~m1_subset_1(X30,k1_zfmisc_1(u1_struct_0(X28)))|(~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(X28)))|(~r2_hidden(X30,u1_pre_topc(X28))|~r2_hidden(X31,u1_pre_topc(X28))|r2_hidden(k9_subset_1(u1_struct_0(X28),X30,X31),u1_pre_topc(X28))))|~v2_pre_topc(X28)|~l1_pre_topc(X28))))&(((m1_subset_1(esk6_1(X28),k1_zfmisc_1(u1_struct_0(X28)))|m1_subset_1(esk5_1(X28),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X28))))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&((m1_subset_1(esk7_1(X28),k1_zfmisc_1(u1_struct_0(X28)))|m1_subset_1(esk5_1(X28),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X28))))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&(((r2_hidden(esk6_1(X28),u1_pre_topc(X28))|m1_subset_1(esk5_1(X28),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X28))))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&(r2_hidden(esk7_1(X28),u1_pre_topc(X28))|m1_subset_1(esk5_1(X28),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X28))))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28)))&(~r2_hidden(k9_subset_1(u1_struct_0(X28),esk6_1(X28),esk7_1(X28)),u1_pre_topc(X28))|m1_subset_1(esk5_1(X28),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X28))))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28)))))&(((m1_subset_1(esk6_1(X28),k1_zfmisc_1(u1_struct_0(X28)))|r1_tarski(esk5_1(X28),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&((m1_subset_1(esk7_1(X28),k1_zfmisc_1(u1_struct_0(X28)))|r1_tarski(esk5_1(X28),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&(((r2_hidden(esk6_1(X28),u1_pre_topc(X28))|r1_tarski(esk5_1(X28),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&(r2_hidden(esk7_1(X28),u1_pre_topc(X28))|r1_tarski(esk5_1(X28),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28)))&(~r2_hidden(k9_subset_1(u1_struct_0(X28),esk6_1(X28),esk7_1(X28)),u1_pre_topc(X28))|r1_tarski(esk5_1(X28),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28)))))&((m1_subset_1(esk6_1(X28),k1_zfmisc_1(u1_struct_0(X28)))|~r2_hidden(k5_setfam_1(u1_struct_0(X28),esk5_1(X28)),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&((m1_subset_1(esk7_1(X28),k1_zfmisc_1(u1_struct_0(X28)))|~r2_hidden(k5_setfam_1(u1_struct_0(X28),esk5_1(X28)),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&(((r2_hidden(esk6_1(X28),u1_pre_topc(X28))|~r2_hidden(k5_setfam_1(u1_struct_0(X28),esk5_1(X28)),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28))&(r2_hidden(esk7_1(X28),u1_pre_topc(X28))|~r2_hidden(k5_setfam_1(u1_struct_0(X28),esk5_1(X28)),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28)))&(~r2_hidden(k9_subset_1(u1_struct_0(X28),esk6_1(X28),esk7_1(X28)),u1_pre_topc(X28))|~r2_hidden(k5_setfam_1(u1_struct_0(X28),esk5_1(X28)),u1_pre_topc(X28))|~r2_hidden(u1_struct_0(X28),u1_pre_topc(X28))|v2_pre_topc(X28)|~l1_pre_topc(X28)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_26, negated_conjecture, ((~v2_struct_0(esk4_0)&(v2_pre_topc(esk4_0)&l1_pre_topc(esk4_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk4_0)))!=u1_struct_0(esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
cnf(c_0_27, plain, (~l1_pre_topc(X1)|~v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|~r2_hidden(X2,u1_pre_topc(X1))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_29, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (v2_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, negated_conjecture, (l1_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (r2_hidden(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~r2_hidden(X2,u1_pre_topc(X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, negated_conjecture, (r2_hidden(u1_struct_0(esk4_0),u1_pre_topc(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
cnf(c_0_34, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_35, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_21, c_0_24])).
cnf(c_0_36, negated_conjecture, (r2_hidden(u1_struct_0(esk4_0),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_31])])).
cnf(c_0_37, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_34, c_0_24])).
fof(c_0_38, plain, ![X12, X13, X14, X16, X17, X18, X19, X21]:((((r2_hidden(X14,esk1_3(X12,X13,X14))|~r2_hidden(X14,X13)|X13!=k3_tarski(X12))&(r2_hidden(esk1_3(X12,X13,X14),X12)|~r2_hidden(X14,X13)|X13!=k3_tarski(X12)))&(~r2_hidden(X16,X17)|~r2_hidden(X17,X12)|r2_hidden(X16,X13)|X13!=k3_tarski(X12)))&((~r2_hidden(esk2_2(X18,X19),X19)|(~r2_hidden(esk2_2(X18,X19),X21)|~r2_hidden(X21,X18))|X19=k3_tarski(X18))&((r2_hidden(esk2_2(X18,X19),esk3_2(X18,X19))|r2_hidden(esk2_2(X18,X19),X19)|X19=k3_tarski(X18))&(r2_hidden(esk3_2(X18,X19),X18)|r2_hidden(esk2_2(X18,X19),X19)|X19=k3_tarski(X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
cnf(c_0_39, negated_conjecture, (~v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, plain, (v1_xboole_0(X1)|r2_hidden(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_37, c_0_28])).
cnf(c_0_41, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|X4!=k3_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_42, negated_conjecture, (r2_hidden(k1_zfmisc_1(u1_struct_0(esk4_0)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk4_0))))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_43, plain, ![X35, X36, X37, X38, X39]:((~r1_tarski(X35,X36)|(~r2_hidden(X37,X35)|r2_hidden(X37,X36)))&((r2_hidden(esk8_2(X38,X39),X38)|r1_tarski(X38,X39))&(~r2_hidden(esk8_2(X38,X39),X39)|r1_tarski(X38,X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_44, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_45, plain, (v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|r2_hidden(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_23, c_0_22])).
cnf(c_0_46, negated_conjecture, (~v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk4_0))))), inference(spm,[status(thm)],[c_0_35, c_0_42])).
fof(c_0_47, plain, ![X27]:k3_tarski(k1_zfmisc_1(X27))=X27, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
cnf(c_0_48, plain, (r1_tarski(X1,X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, negated_conjecture, (r2_hidden(X1,k3_tarski(u1_pre_topc(esk4_0)))|~r2_hidden(X1,u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_44, c_0_33])).
cnf(c_0_50, negated_conjecture, (r2_hidden(u1_pre_topc(esk4_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk4_0))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_31]), c_0_46])).
cnf(c_0_51, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_52, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_54, negated_conjecture, (r1_tarski(X1,k3_tarski(u1_pre_topc(esk4_0)))|~r2_hidden(esk8_2(X1,k3_tarski(u1_pre_topc(esk4_0))),u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_55, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_56, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk4_0)))|~r2_hidden(X1,u1_pre_topc(esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_50]), c_0_51])).
cnf(c_0_57, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_58, negated_conjecture, (r1_tarski(u1_struct_0(esk4_0),k3_tarski(u1_pre_topc(esk4_0)))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_60, plain, ![X41, X42]:(~m1_subset_1(X42,k1_zfmisc_1(k1_zfmisc_1(X41)))|k5_setfam_1(X41,X42)=k3_tarski(X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
cnf(c_0_61, negated_conjecture, (r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk4_0)))|~r2_hidden(esk8_2(X1,k1_zfmisc_1(u1_struct_0(esk4_0))),u1_pre_topc(esk4_0))), inference(spm,[status(thm)],[c_0_48, c_0_56])).
cnf(c_0_62, negated_conjecture, (k3_tarski(u1_pre_topc(esk4_0))=u1_struct_0(esk4_0)|~m1_subset_1(k3_tarski(u1_pre_topc(esk4_0)),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_63, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_59, c_0_35])).
cnf(c_0_64, plain, (r2_hidden(k5_setfam_1(u1_struct_0(X2),X1),u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~r1_tarski(X1,u1_pre_topc(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_65, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66, negated_conjecture, (r1_tarski(u1_pre_topc(esk4_0),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(spm,[status(thm)],[c_0_61, c_0_55])).
fof(c_0_67, plain, ![X43]:(v1_xboole_0(X43)|(~r2_hidden(k3_tarski(X43),X43)|k4_yellow_0(k2_yellow_1(X43))=k3_tarski(X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
cnf(c_0_68, negated_conjecture, (k3_tarski(u1_pre_topc(esk4_0))=u1_struct_0(esk4_0)|~r2_hidden(k3_tarski(u1_pre_topc(esk4_0)),k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_69, plain, (r2_hidden(k3_tarski(X1),u1_pre_topc(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~r1_tarski(X1,u1_pre_topc(X2))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(u1_pre_topc(esk4_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk4_0))))), inference(spm,[status(thm)],[c_0_18, c_0_66])).
cnf(c_0_71, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_72, negated_conjecture, (k3_tarski(u1_pre_topc(esk4_0))=u1_struct_0(esk4_0)|~r2_hidden(k3_tarski(u1_pre_topc(esk4_0)),u1_pre_topc(esk4_0))), inference(spm,[status(thm)],[c_0_68, c_0_56])).
cnf(c_0_73, negated_conjecture, (r2_hidden(k3_tarski(u1_pre_topc(esk4_0)),u1_pre_topc(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_30]), c_0_31]), c_0_19])])).
cnf(c_0_74, plain, (k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(csr,[status(thm)],[c_0_71, c_0_35])).
cnf(c_0_75, negated_conjecture, (k3_tarski(u1_pre_topc(esk4_0))=u1_struct_0(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_73])])).
cnf(c_0_76, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk4_0)))!=u1_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_77, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_33])]), c_0_76]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 78
# Proof object clause steps            : 53
# Proof object formula steps           : 25
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 45
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 45
# Processed clauses                    : 1704
# ...of these trivial                  : 5
# ...subsumed                          : 961
# ...remaining for further processing  : 738
# Other redundant clauses eliminated   : 76
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 138
# Generated clauses                    : 23340
# ...of the previous two non-trivial   : 22511
# Contextual simplify-reflections      : 9
# Paramodulations                      : 23256
# Factorizations                       : 8
# Equation resolutions                 : 76
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 541
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 510
# Current number of unprocessed clauses: 20872
# ...number of literals in the above   : 86115
# Current number of archived formulas  : 0
# Current number of archived clauses   : 192
# Clause-clause subsumption calls (NU) : 65854
# Rec. Clause-clause subsumption calls : 45538
# Non-unit clause-clause subsumptions  : 854
# Unit Clause-clause subsumption calls : 143
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 34
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 400165

# -------------------------------------------------
# User time                : 0.219 s
# System time              : 0.007 s
# Total time               : 0.226 s
# Maximum resident set size: 28676 pages
