# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.007 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/b3f512ae643f4e29994ddd63c72525ad.E_input', d3_tarski)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/b3f512ae643f4e29994ddd63c72525ad.E_input', l49_zfmisc_1)).
fof(t23_mcart_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1)))), file('data/MPTP/proofs/b3f512ae643f4e29994ddd63c72525ad.E_input', t23_mcart_1)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/b3f512ae643f4e29994ddd63c72525ad.E_input', t99_zfmisc_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/b3f512ae643f4e29994ddd63c72525ad.E_input', d1_zfmisc_1)).
fof(t21_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))), file('data/MPTP/proofs/b3f512ae643f4e29994ddd63c72525ad.E_input', t21_relat_1)).
fof(l139_zfmisc_1, axiom, ![X1, X2, X3]:~((r2_hidden(X1,k2_zfmisc_1(X2,X3))&![X4, X5]:~(k4_tarski(X4,X5)=X1))), file('data/MPTP/proofs/b3f512ae643f4e29994ddd63c72525ad.E_input', l139_zfmisc_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/b3f512ae643f4e29994ddd63c72525ad.E_input', t7_mcart_1)).
fof(c_0_8, plain, ![X9, X10, X11, X12, X13]:((~r1_tarski(X9,X10)|(~r2_hidden(X11,X9)|r2_hidden(X11,X10)))&((r2_hidden(esk3_2(X12,X13),X12)|r1_tarski(X12,X13))&(~r2_hidden(esk3_2(X12,X13),X13)|r1_tarski(X12,X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_9, plain, ![X30, X31]:(~r2_hidden(X30,X31)|r1_tarski(X30,k3_tarski(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
fof(c_0_10, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,X2)=>X1=k4_tarski(k1_mcart_1(X1),k2_mcart_1(X1))))), inference(assume_negation,[status(cth)],[t23_mcart_1])).
cnf(c_0_11, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, negated_conjecture, (v1_relat_1(esk2_0)&(r2_hidden(esk1_0,esk2_0)&esk1_0!=k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_14, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X1,X3)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_16, plain, ![X15]:k3_tarski(k1_zfmisc_1(X15))=X15, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
fof(c_0_17, plain, ![X16, X17, X18, X19, X20, X21]:(((~r2_hidden(X18,X17)|r1_tarski(X18,X16)|X17!=k1_zfmisc_1(X16))&(~r1_tarski(X19,X16)|r2_hidden(X19,X17)|X17!=k1_zfmisc_1(X16)))&((~r2_hidden(esk4_2(X20,X21),X21)|~r1_tarski(esk4_2(X20,X21),X20)|X21=k1_zfmisc_1(X20))&(r2_hidden(esk4_2(X20,X21),X21)|r1_tarski(esk4_2(X20,X21),X20)|X21=k1_zfmisc_1(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_18, negated_conjecture, (r2_hidden(esk1_0,k3_tarski(X1))|~r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk1_0,X1)|~r2_hidden(esk2_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_20])).
fof(c_0_23, plain, ![X8]:(~v1_relat_1(X8)|r1_tarski(X8,k2_zfmisc_1(k1_relat_1(X8),k2_relat_1(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_relat_1])])).
fof(c_0_24, plain, ![X23, X24, X25]:(~r2_hidden(X23,k2_zfmisc_1(X24,X25))|k4_tarski(esk5_1(X23),esk6_1(X23))=X23), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l139_zfmisc_1])])])])])])).
cnf(c_0_25, negated_conjecture, (r2_hidden(esk1_0,X1)|~r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_28, plain, ![X28, X29]:(k1_mcart_1(k4_tarski(X28,X29))=X28&k2_mcart_1(k4_tarski(X28,X29))=X29), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_29, plain, (k4_tarski(esk5_1(X1),esk6_1(X1))=X1|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(k1_relat_1(esk2_0),k2_relat_1(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])])).
cnf(c_0_31, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, negated_conjecture, (k4_tarski(esk5_1(esk1_0),esk6_1(esk1_0))=esk1_0), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, negated_conjecture, (esk5_1(esk1_0)=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_34, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),esk6_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36, negated_conjecture, (esk6_1(esk1_0)=k2_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_37, negated_conjecture, (esk1_0!=k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_38, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36]), c_0_37]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 22
# Proof object formula steps           : 17
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 49
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 46
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 35
# ...of the previous two non-trivial   : 32
# Contextual simplify-reflections      : 0
# Paramodulations                      : 33
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 14
# ...number of literals in the above   : 40
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 64
# Rec. Clause-clause subsumption calls : 49
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1347

# -------------------------------------------------
# User time                : 0.004 s
# System time              : 0.003 s
# Total time               : 0.008 s
# Maximum resident set size: 28644 pages
