# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t70_enumset1)).
fof(t34_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2)))), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t34_ordinal1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', d3_xboole_0)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t75_enumset1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t69_enumset1)).
fof(connectedness_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1))), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', connectedness_r1_ordinal1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', d2_tarski)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t7_ordinal1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', redefinition_r1_ordinal1)).
fof(t24_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1))))))), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', t24_ordinal1)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/b439d9cd46fb4c808a14d2b3c460c7a3.E_input', d6_enumset1)).
fof(c_0_17, plain, ![X61, X62]:k3_tarski(k2_tarski(X61,X62))=k2_xboole_0(X61,X62), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_18, plain, ![X13, X14]:k1_enumset1(X13,X13,X14)=k2_tarski(X13,X14), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_19, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,k1_ordinal1(X2))<=>r1_ordinal1(X1,X2))))), inference(assume_negation,[status(cth)],[t34_ordinal1])).
fof(c_0_20, plain, ![X77, X78, X79, X80, X81, X82, X83, X84]:(((~r2_hidden(X80,X79)|(r2_hidden(X80,X77)|r2_hidden(X80,X78))|X79!=k2_xboole_0(X77,X78))&((~r2_hidden(X81,X77)|r2_hidden(X81,X79)|X79!=k2_xboole_0(X77,X78))&(~r2_hidden(X81,X78)|r2_hidden(X81,X79)|X79!=k2_xboole_0(X77,X78))))&(((~r2_hidden(esk5_3(X82,X83,X84),X82)|~r2_hidden(esk5_3(X82,X83,X84),X84)|X84=k2_xboole_0(X82,X83))&(~r2_hidden(esk5_3(X82,X83,X84),X83)|~r2_hidden(esk5_3(X82,X83,X84),X84)|X84=k2_xboole_0(X82,X83)))&(r2_hidden(esk5_3(X82,X83,X84),X84)|(r2_hidden(esk5_3(X82,X83,X84),X82)|r2_hidden(esk5_3(X82,X83,X84),X83))|X84=k2_xboole_0(X82,X83)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_21, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X67, X68, X69]:k2_enumset1(X67,X67,X68,X69)=k1_enumset1(X67,X68,X69), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_24, plain, ![X73, X74, X75, X76]:k3_enumset1(X73,X73,X74,X75,X76)=k2_enumset1(X73,X74,X75,X76), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_25, plain, ![X86, X87, X88, X89, X90]:k4_enumset1(X86,X86,X87,X88,X89,X90)=k3_enumset1(X86,X87,X88,X89,X90), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_26, plain, ![X17, X18, X19, X20, X21, X22]:k5_enumset1(X17,X17,X18,X19,X20,X21,X22)=k4_enumset1(X17,X18,X19,X20,X21,X22), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_27, plain, ![X54, X55, X56, X57, X58, X59, X60]:k6_enumset1(X54,X54,X55,X56,X57,X58,X59,X60)=k5_enumset1(X54,X55,X56,X57,X58,X59,X60), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_28, plain, ![X72]:k1_ordinal1(X72)=k2_xboole_0(X72,k1_tarski(X72)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_29, plain, ![X53]:k2_tarski(X53,X53)=k1_tarski(X53), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_30, plain, ![X65, X66]:(~v3_ordinal1(X65)|~v3_ordinal1(X66)|(r1_ordinal1(X65,X66)|r1_ordinal1(X66,X65))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_ordinal1])])).
fof(c_0_31, negated_conjecture, (v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&((~r2_hidden(esk1_0,k1_ordinal1(esk2_0))|~r1_ordinal1(esk1_0,esk2_0))&(r2_hidden(esk1_0,k1_ordinal1(esk2_0))|r1_ordinal1(esk1_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
fof(c_0_32, plain, ![X23, X24, X25, X26, X27, X28, X29, X30]:(((~r2_hidden(X26,X25)|(X26=X23|X26=X24)|X25!=k2_tarski(X23,X24))&((X27!=X23|r2_hidden(X27,X25)|X25!=k2_tarski(X23,X24))&(X27!=X24|r2_hidden(X27,X25)|X25!=k2_tarski(X23,X24))))&(((esk3_3(X28,X29,X30)!=X28|~r2_hidden(esk3_3(X28,X29,X30),X30)|X30=k2_tarski(X28,X29))&(esk3_3(X28,X29,X30)!=X29|~r2_hidden(esk3_3(X28,X29,X30),X30)|X30=k2_tarski(X28,X29)))&(r2_hidden(esk3_3(X28,X29,X30),X30)|(esk3_3(X28,X29,X30)=X28|esk3_3(X28,X29,X30)=X29)|X30=k2_tarski(X28,X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_33, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_34, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_35, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_38, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_42, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_43, plain, ![X15, X16]:(~r2_hidden(X15,X16)|~r1_tarski(X16,X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_44, plain, ![X63, X64]:((~r1_ordinal1(X63,X64)|r1_tarski(X63,X64)|(~v3_ordinal1(X63)|~v3_ordinal1(X64)))&(~r1_tarski(X63,X64)|r1_ordinal1(X63,X64)|(~v3_ordinal1(X63)|~v3_ordinal1(X64)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_45, plain, (r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_46, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_47, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_48, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk1_0,k1_ordinal1(esk2_0))|r1_ordinal1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_50, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51, negated_conjecture, (~r2_hidden(esk1_0,k1_ordinal1(esk2_0))|~r1_ordinal1(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_52, plain, (r2_hidden(X1,X3)|X3!=k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_53, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, negated_conjecture, (r1_ordinal1(X1,esk2_0)|r1_ordinal1(esk2_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_57, plain, (X1=X4|X1=X3|X2!=k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_22]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_58, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2)))), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_59, negated_conjecture, (r1_ordinal1(esk1_0,esk2_0)|r2_hidden(esk1_0,k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50]), c_0_22]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38]), c_0_39]), c_0_39])).
cnf(c_0_60, negated_conjecture, (~r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,k3_tarski(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_50]), c_0_22]), c_0_34]), c_0_35]), c_0_35]), c_0_36]), c_0_36]), c_0_37]), c_0_37]), c_0_38]), c_0_38]), c_0_39]), c_0_39])).
cnf(c_0_61, plain, (r2_hidden(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (~r1_ordinal1(X1,X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_63, negated_conjecture, (r1_ordinal1(esk2_0,esk1_0)|r1_ordinal1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_64, plain, (X1=X2|X1=X3|~r2_hidden(X1,k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_65, negated_conjecture, (r1_ordinal1(esk1_0,esk2_0)|r2_hidden(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_66, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_67, plain, ![X70, X71]:(~v3_ordinal1(X70)|(~v3_ordinal1(X71)|(r2_hidden(X70,X71)|(X70=X71|r2_hidden(X71,X70))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])).
cnf(c_0_68, negated_conjecture, (~r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_69, negated_conjecture, (r1_ordinal1(esk2_0,esk1_0)|~r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_46]), c_0_56])])).
cnf(c_0_70, negated_conjecture, (esk1_0=esk2_0|r1_ordinal1(esk1_0,esk2_0)|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_71, plain, (r2_hidden(X1,X3)|X3!=k3_tarski(k6_enumset1(X4,X4,X4,X4,X4,X4,X4,X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_34]), c_0_35]), c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_72, plain, (r2_hidden(X1,X2)|X1=X2|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_73, negated_conjecture, (r1_ordinal1(esk2_0,esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_68, c_0_63])).
cnf(c_0_74, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)|~r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_69]), c_0_56]), c_0_46])])).
cnf(c_0_75, negated_conjecture, (esk1_0=esk2_0|r2_hidden(esk1_0,esk2_0)|~r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_70]), c_0_46]), c_0_56])])).
fof(c_0_76, plain, ![X32, X33, X34, X35, X36, X37, X38, X39, X40, X41, X42, X43, X44, X45, X46, X47, X48, X49, X50, X51]:(((~r2_hidden(X41,X40)|(X41=X32|(X41=X33|(X41=X34|(X41=X35|(X41=X36|(X41=X37|(X41=X38|X41=X39)))))))|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X32|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X33|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X34|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X35|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X36|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X37|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&((X42!=X38|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))&(X42!=X39|r2_hidden(X42,X40)|X40!=k6_enumset1(X32,X33,X34,X35,X36,X37,X38,X39))))))))))&(((esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X43|~r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X44|~r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X45|~r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X46|~r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X47|~r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X48|~r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&((esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X49|~r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50))&(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)!=X50|~r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50)))))))))&(r2_hidden(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51),X51)|(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X43|(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X44|(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X45|(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X46|(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X47|(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X48|(esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X49|esk4_9(X43,X44,X45,X46,X47,X48,X49,X50,X51)=X50)))))))|X51=k6_enumset1(X43,X44,X45,X46,X47,X48,X49,X50)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
cnf(c_0_77, plain, (r2_hidden(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_78, negated_conjecture, (X1=esk2_0|r2_hidden(X1,esk2_0)|r2_hidden(esk2_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_72, c_0_46])).
cnf(c_0_79, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_73]), c_0_56]), c_0_46])])).
cnf(c_0_80, negated_conjecture, (esk1_0=esk2_0|~r2_hidden(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_81, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X2,X4,X5,X6,X7,X8,X9,X10)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_82, negated_conjecture, (~r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_60, c_0_77])).
cnf(c_0_83, negated_conjecture, (esk1_0=esk2_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_56]), c_0_79]), c_0_80])).
cnf(c_0_84, negated_conjecture, (r1_ordinal1(esk2_0,esk2_0)), inference(spm,[status(thm)],[c_0_55, c_0_46])).
cnf(c_0_85, plain, (r2_hidden(X1,k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_81])])).
cnf(c_0_86, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_83]), c_0_84]), c_0_83]), c_0_85])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 87
# Proof object clause steps            : 52
# Proof object formula steps           : 35
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 77
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 48
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 39
# Processed clauses                    : 104
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 100
# Other redundant clauses eliminated   : 29
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 16
# Generated clauses                    : 108
# ...of the previous two non-trivial   : 96
# Contextual simplify-reflections      : 1
# Paramodulations                      : 87
# Factorizations                       : 0
# Equation resolutions                 : 29
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 65
# ...number of literals in the above   : 296
# Current number of archived formulas  : 0
# Current number of archived clauses   : 65
# Clause-clause subsumption calls (NU) : 200
# Rec. Clause-clause subsumption calls : 136
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 40
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 69
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4201

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28796 pages
