# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/b455a87e2a374dc2853c4d941d396678.E_input', d1_xboole_0)).
fof(d1_relat_1, axiom, ![X1]:(v1_relat_1(X1)<=>![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(X2=k4_tarski(X3,X4))))), file('data/MPTP/proofs/b455a87e2a374dc2853c4d941d396678.E_input', d1_relat_1)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/b455a87e2a374dc2853c4d941d396678.E_input', d8_relat_1)).
fof(t62_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0)), file('data/MPTP/proofs/b455a87e2a374dc2853c4d941d396678.E_input', t62_relat_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/b455a87e2a374dc2853c4d941d396678.E_input', fc1_xboole_0)).
fof(c_0_5, plain, ![X7, X8, X9]:((~v1_xboole_0(X7)|~r2_hidden(X8,X7))&(r2_hidden(esk1_1(X9),X9)|v1_xboole_0(X9))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_6, plain, ![X11, X12, X15, X17, X18]:((~v1_relat_1(X11)|(~r2_hidden(X12,X11)|X12=k4_tarski(esk2_2(X11,X12),esk3_2(X11,X12))))&((r2_hidden(esk4_1(X15),X15)|v1_relat_1(X15))&(esk4_1(X15)!=k4_tarski(X17,X18)|v1_relat_1(X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_relat_1])])])])])])])).
fof(c_0_7, plain, ![X19, X20, X21, X22, X23, X25, X26, X27, X30]:((((r2_hidden(k4_tarski(X22,esk5_5(X19,X20,X21,X22,X23)),X19)|~r2_hidden(k4_tarski(X22,X23),X21)|X21!=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19))&(r2_hidden(k4_tarski(esk5_5(X19,X20,X21,X22,X23),X23),X20)|~r2_hidden(k4_tarski(X22,X23),X21)|X21!=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19)))&(~r2_hidden(k4_tarski(X25,X27),X19)|~r2_hidden(k4_tarski(X27,X26),X20)|r2_hidden(k4_tarski(X25,X26),X21)|X21!=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19)))&((~r2_hidden(k4_tarski(esk6_3(X19,X20,X21),esk7_3(X19,X20,X21)),X21)|(~r2_hidden(k4_tarski(esk6_3(X19,X20,X21),X30),X19)|~r2_hidden(k4_tarski(X30,esk7_3(X19,X20,X21)),X20))|X21=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19))&((r2_hidden(k4_tarski(esk6_3(X19,X20,X21),esk8_3(X19,X20,X21)),X19)|r2_hidden(k4_tarski(esk6_3(X19,X20,X21),esk7_3(X19,X20,X21)),X21)|X21=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19))&(r2_hidden(k4_tarski(esk8_3(X19,X20,X21),esk7_3(X19,X20,X21)),X20)|r2_hidden(k4_tarski(esk6_3(X19,X20,X21),esk7_3(X19,X20,X21)),X21)|X21=k5_relat_1(X19,X20)|~v1_relat_1(X21)|~v1_relat_1(X20)|~v1_relat_1(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
cnf(c_0_8, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (r2_hidden(esk4_1(X1),X1)|v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (r2_hidden(k4_tarski(esk8_3(X1,X2,X3),esk7_3(X1,X2,X3)),X2)|r2_hidden(k4_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),X3)|X3=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_8, c_0_9])).
fof(c_0_12, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0))), inference(assume_negation,[status(cth)],[t62_relat_1])).
cnf(c_0_13, plain, (X1=k5_relat_1(X2,X3)|r2_hidden(k4_tarski(esk8_3(X2,X3,X1),esk7_3(X2,X3,X1)),X3)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_8])).
fof(c_0_14, negated_conjecture, (v1_relat_1(esk9_0)&(k5_relat_1(k1_xboole_0,esk9_0)!=k1_xboole_0|k5_relat_1(esk9_0,k1_xboole_0)!=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_15, plain, (r2_hidden(k4_tarski(esk6_3(X1,X2,X3),esk8_3(X1,X2,X3)),X1)|r2_hidden(k4_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),X3)|X3=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, plain, (X1=k5_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_11]), c_0_8])).
cnf(c_0_17, negated_conjecture, (v1_relat_1(esk9_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (X1=k5_relat_1(X2,X3)|r2_hidden(k4_tarski(esk6_3(X2,X3,X1),esk8_3(X2,X3,X1)),X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_11]), c_0_8])).
cnf(c_0_19, negated_conjecture, (X1=k5_relat_1(esk9_0,X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_21, negated_conjecture, (X1=k5_relat_1(X2,esk9_0)|r2_hidden(k4_tarski(esk6_3(X2,esk9_0,X1),esk8_3(X2,esk9_0,X1)),X2)|~v1_relat_1(X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_22, negated_conjecture, (k5_relat_1(k1_xboole_0,esk9_0)!=k1_xboole_0|k5_relat_1(esk9_0,k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, negated_conjecture, (k5_relat_1(esk9_0,X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_24, negated_conjecture, (X1=k5_relat_1(X2,esk9_0)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_11]), c_0_8])).
cnf(c_0_25, negated_conjecture, (k5_relat_1(k1_xboole_0,esk9_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_20])])).
cnf(c_0_26, negated_conjecture, (k5_relat_1(X1,esk9_0)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_24, c_0_20])).
cnf(c_0_27, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_20])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 18
# Proof object formula steps           : 10
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 14
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 14
# Processed clauses                    : 105
# ...of these trivial                  : 0
# ...subsumed                          : 15
# ...remaining for further processing  : 90
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 3
# Generated clauses                    : 162
# ...of the previous two non-trivial   : 156
# Contextual simplify-reflections      : 13
# Paramodulations                      : 154
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 67
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 57
# Current number of unprocessed clauses: 79
# ...number of literals in the above   : 418
# Current number of archived formulas  : 0
# Current number of archived clauses   : 21
# Clause-clause subsumption calls (NU) : 1385
# Rec. Clause-clause subsumption calls : 301
# Non-unit clause-clause subsumptions  : 25
# Unit Clause-clause subsumption calls : 82
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4796

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28632 pages
