# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(fraenkel_a_2_9_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_9_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattice3(X2,X4,X3))))), file('data/MPTP/proofs/b473d972fe4848c4bf33e3c192220f80.E_input', fraenkel_a_2_9_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/b473d972fe4848c4bf33e3c192220f80.E_input', d17_lattice3)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/b473d972fe4848c4bf33e3c192220f80.E_input', redefinition_r3_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/b473d972fe4848c4bf33e3c192220f80.E_input', cc1_lattices)).
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/b473d972fe4848c4bf33e3c192220f80.E_input', d16_lattice3)).
fof(t34_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2))))))), file('data/MPTP/proofs/b473d972fe4848c4bf33e3c192220f80.E_input', t34_lattice3)).
fof(d18_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>(v4_lattice3(X1)<=>![X2]:?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4))))))), file('data/MPTP/proofs/b473d972fe4848c4bf33e3c192220f80.E_input', d18_lattice3)).
fof(t40_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)=>r3_lattices(X1,X2,k16_lattice3(X1,X3))))), file('data/MPTP/proofs/b473d972fe4848c4bf33e3c192220f80.E_input', t40_lattice3)).
fof(c_0_8, plain, ![X5, X6, X7, X9, X10]:(((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X6))|~r2_hidden(X5,a_2_9_lattice3(X6,X7))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&((X5=esk1_3(X5,X6,X7)|~r2_hidden(X5,a_2_9_lattice3(X6,X7))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))&(r3_lattice3(X6,esk1_3(X5,X6,X7),X7)|~r2_hidden(X5,a_2_9_lattice3(X6,X7))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))))&(~m1_subset_1(X10,u1_struct_0(X6))|(X5!=X10|~r3_lattice3(X6,X10,X9))|r2_hidden(X5,a_2_9_lattice3(X6,X9))|(v2_struct_0(X6)|(~v10_lattices(X6)|(~v4_lattice3(X6)|~l3_lattices(X6)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_9_lattice3])])])])])])])).
fof(c_0_9, plain, ![X11, X12, X13, X14, X15]:((~r4_lattice3(X11,X12,X13)|(~m1_subset_1(X14,u1_struct_0(X11))|(~r2_hidden(X14,X13)|r1_lattices(X11,X14,X12)))|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|~l3_lattices(X11)))&((m1_subset_1(esk2_3(X11,X12,X15),u1_struct_0(X11))|r4_lattice3(X11,X12,X15)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|~l3_lattices(X11)))&((r2_hidden(esk2_3(X11,X12,X15),X15)|r4_lattice3(X11,X12,X15)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|~l3_lattices(X11)))&(~r1_lattices(X11,esk2_3(X11,X12,X15),X12)|r4_lattice3(X11,X12,X15)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|~l3_lattices(X11)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_10, plain, (X1=esk1_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_9_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (r2_hidden(esk2_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_9_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_13, plain, ![X17, X18, X19]:((~r3_lattices(X17,X18,X19)|r1_lattices(X17,X18,X19)|(v2_struct_0(X17)|(~v6_lattices(X17)|(~v8_lattices(X17)|(~v9_lattices(X17)|~l3_lattices(X17))))|(~m1_subset_1(X18,u1_struct_0(X17))|~m1_subset_1(X19,u1_struct_0(X17)))))&(~r1_lattices(X17,X18,X19)|r3_lattices(X17,X18,X19)|(v2_struct_0(X17)|(~v6_lattices(X17)|(~v8_lattices(X17)|(~v9_lattices(X17)|~l3_lattices(X17))))|(~m1_subset_1(X18,u1_struct_0(X17))|~m1_subset_1(X19,u1_struct_0(X17)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
fof(c_0_14, plain, ![X20]:((~v2_struct_0(X20)|(v2_struct_0(X20)|~v10_lattices(X20))|~l3_lattices(X20))&((v4_lattices(X20)|(v2_struct_0(X20)|~v10_lattices(X20))|~l3_lattices(X20))&((v5_lattices(X20)|(v2_struct_0(X20)|~v10_lattices(X20))|~l3_lattices(X20))&((v6_lattices(X20)|(v2_struct_0(X20)|~v10_lattices(X20))|~l3_lattices(X20))&((v7_lattices(X20)|(v2_struct_0(X20)|~v10_lattices(X20))|~l3_lattices(X20))&((v8_lattices(X20)|(v2_struct_0(X20)|~v10_lattices(X20))|~l3_lattices(X20))&(v9_lattices(X20)|(v2_struct_0(X20)|~v10_lattices(X20))|~l3_lattices(X20)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_15, plain, ![X21, X22, X23, X24, X25]:((~r3_lattice3(X21,X22,X23)|(~m1_subset_1(X24,u1_struct_0(X21))|(~r2_hidden(X24,X23)|r1_lattices(X21,X22,X24)))|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|~l3_lattices(X21)))&((m1_subset_1(esk3_3(X21,X22,X25),u1_struct_0(X21))|r3_lattice3(X21,X22,X25)|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|~l3_lattices(X21)))&((r2_hidden(esk3_3(X21,X22,X25),X25)|r3_lattice3(X21,X22,X25)|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|~l3_lattices(X21)))&(~r1_lattices(X21,X22,esk3_3(X21,X22,X25))|r3_lattice3(X21,X22,X25)|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|~l3_lattices(X21)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
cnf(c_0_16, plain, (r3_lattice3(X1,esk1_3(X2,X1,X3),X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_9_lattice3(X1,X3))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, plain, (esk1_3(esk2_3(X1,X2,a_2_9_lattice3(X3,X4)),X3,X4)=esk2_3(X1,X2,a_2_9_lattice3(X3,X4))|r4_lattice3(X1,X2,a_2_9_lattice3(X3,X4))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)|~v4_lattice3(X3)|~v10_lattices(X3)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_18, plain, (r4_lattice3(X1,X2,a_2_9_lattice3(X3,X4))|m1_subset_1(esk1_3(esk2_3(X1,X2,a_2_9_lattice3(X3,X4)),X3,X4),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)|~v4_lattice3(X3)|~v10_lattices(X3)), inference(spm,[status(thm)],[c_0_12, c_0_11])).
fof(c_0_19, plain, ![X27, X28, X29, X30, X31]:(((r3_lattice3(X27,X28,X29)|X28!=k16_lattice3(X27,X29)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v10_lattices(X27)|(~v4_lattice3(X27)|~l3_lattices(X27)))))&(~m1_subset_1(X30,u1_struct_0(X27))|(~r3_lattice3(X27,X30,X29)|r3_lattices(X27,X30,X28))|X28!=k16_lattice3(X27,X29)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v10_lattices(X27)|(~v4_lattice3(X27)|~l3_lattices(X27))))))&((m1_subset_1(esk4_3(X27,X28,X31),u1_struct_0(X27))|~r3_lattice3(X27,X28,X31)|X28=k16_lattice3(X27,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v10_lattices(X27)|(~v4_lattice3(X27)|~l3_lattices(X27)))))&((r3_lattice3(X27,esk4_3(X27,X28,X31),X31)|~r3_lattice3(X27,X28,X31)|X28=k16_lattice3(X27,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v10_lattices(X27)|(~v4_lattice3(X27)|~l3_lattices(X27)))))&(~r3_lattices(X27,esk4_3(X27,X28,X31),X28)|~r3_lattice3(X27,X28,X31)|X28=k16_lattice3(X27,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|(~v10_lattices(X27)|(~v4_lattice3(X27)|~l3_lattices(X27)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])).
cnf(c_0_20, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_24, plain, ![X33, X34, X36, X38]:(((m1_subset_1(esk5_2(X33,X34),u1_struct_0(X33))|~v4_lattice3(X33)|(v2_struct_0(X33)|~l3_lattices(X33)))&((r4_lattice3(X33,esk5_2(X33,X34),X34)|~v4_lattice3(X33)|(v2_struct_0(X33)|~l3_lattices(X33)))&(~m1_subset_1(X36,u1_struct_0(X33))|(~r4_lattice3(X33,X36,X34)|r1_lattices(X33,esk5_2(X33,X34),X36))|~v4_lattice3(X33)|(v2_struct_0(X33)|~l3_lattices(X33)))))&((m1_subset_1(esk7_2(X33,X38),u1_struct_0(X33))|~r4_lattice3(X33,X38,esk6_1(X33))|~m1_subset_1(X38,u1_struct_0(X33))|v4_lattice3(X33)|(v2_struct_0(X33)|~l3_lattices(X33)))&((r4_lattice3(X33,esk7_2(X33,X38),esk6_1(X33))|~r4_lattice3(X33,X38,esk6_1(X33))|~m1_subset_1(X38,u1_struct_0(X33))|v4_lattice3(X33)|(v2_struct_0(X33)|~l3_lattices(X33)))&(~r1_lattices(X33,X38,esk7_2(X33,X38))|~r4_lattice3(X33,X38,esk6_1(X33))|~m1_subset_1(X38,u1_struct_0(X33))|v4_lattice3(X33)|(v2_struct_0(X33)|~l3_lattices(X33)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_lattice3])])])])])])).
cnf(c_0_25, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26, plain, (r4_lattice3(X1,X2,a_2_9_lattice3(X3,X4))|r3_lattice3(X3,esk2_3(X1,X2,a_2_9_lattice3(X3,X4)),X4)|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)|~v4_lattice3(X3)|~v10_lattices(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_11])).
cnf(c_0_27, plain, (r4_lattice3(X1,X2,a_2_9_lattice3(X3,X4))|m1_subset_1(esk2_3(X1,X2,a_2_9_lattice3(X3,X4)),u1_struct_0(X3))|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)|~v4_lattice3(X3)|~v10_lattices(X3)), inference(spm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_28, plain, (X2=k16_lattice3(X1,X3)|v2_struct_0(X1)|~r3_lattices(X1,esk4_3(X1,X2,X3),X2)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_30, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|X2=k16_lattice3(X1,X3)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (r1_lattices(X1,X4,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, plain, (r4_lattice3(X1,esk5_2(X1,X2),X2)|v2_struct_0(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,esk3_3(X1,X2,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_35, plain, (r1_lattices(X2,esk5_2(X2,X3),X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r4_lattice3(X2,X1,X3)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,esk2_3(X1,X2,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37, plain, (r1_lattices(X1,esk2_3(X2,X3,a_2_9_lattice3(X1,X4)),X5)|r4_lattice3(X2,X3,a_2_9_lattice3(X1,X4))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X5,X4)|~l3_lattices(X1)|~l3_lattices(X2)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
cnf(c_0_38, plain, (X1=k16_lattice3(X2,X3)|v2_struct_0(X2)|~r1_lattices(X2,esk4_3(X2,X1,X3),X1)|~r3_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
cnf(c_0_39, plain, (r1_lattices(X1,X2,esk5_2(X1,X3))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,X3)|~l3_lattices(X1)|~v4_lattice3(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_40, plain, (r3_lattice3(X1,esk5_2(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,esk3_3(X1,esk5_2(X1,X2),X3),X2)|~m1_subset_1(esk3_3(X1,esk5_2(X1,X2),X3),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_33])).
cnf(c_0_41, plain, (r4_lattice3(X1,X2,a_2_9_lattice3(X1,X3))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r2_hidden(X2,X3)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42, plain, (esk5_2(X1,X2)=k16_lattice3(X1,X3)|v2_struct_0(X1)|~r3_lattice3(X1,esk5_2(X1,X2),X3)|~m1_subset_1(esk4_3(X1,esk5_2(X1,X2),X3),u1_struct_0(X1))|~r2_hidden(esk4_3(X1,esk5_2(X1,X2),X3),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_33])).
cnf(c_0_43, plain, (r2_hidden(X3,a_2_9_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattice3(X2,X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_44, plain, (r3_lattice3(X1,esk5_2(X1,a_2_9_lattice3(X1,X2)),X3)|v2_struct_0(X1)|~m1_subset_1(esk3_3(X1,esk5_2(X1,a_2_9_lattice3(X1,X2)),X3),u1_struct_0(X1))|~r2_hidden(esk3_3(X1,esk5_2(X1,a_2_9_lattice3(X1,X2)),X3),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_46, plain, (esk5_2(X1,X2)=k16_lattice3(X1,X3)|v2_struct_0(X1)|~r3_lattice3(X1,esk5_2(X1,X2),X3)|~r2_hidden(esk4_3(X1,esk5_2(X1,X2),X3),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_30]), c_0_33])).
cnf(c_0_47, plain, (r2_hidden(X1,a_2_9_lattice3(X2,X3))|v2_struct_0(X2)|~r3_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_48, plain, (r3_lattice3(X1,esk5_2(X1,a_2_9_lattice3(X1,X2)),X3)|v2_struct_0(X1)|~r2_hidden(esk3_3(X1,esk5_2(X1,a_2_9_lattice3(X1,X2)),X3),X2)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_33])).
cnf(c_0_49, plain, (r2_hidden(esk3_3(X1,X2,X3),X3)|r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_50, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)=>r3_lattices(X1,X2,k16_lattice3(X1,X3)))))), inference(assume_negation,[status(cth)],[t40_lattice3])).
cnf(c_0_51, plain, (esk5_2(X1,a_2_9_lattice3(X2,X3))=k16_lattice3(X1,X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_lattice3(X2,esk4_3(X1,esk5_2(X1,a_2_9_lattice3(X2,X3)),X4),X3)|~r3_lattice3(X1,esk5_2(X1,a_2_9_lattice3(X2,X3)),X4)|~m1_subset_1(esk4_3(X1,esk5_2(X1,a_2_9_lattice3(X2,X3)),X4),u1_struct_0(X2))|~l3_lattices(X1)|~l3_lattices(X2)|~v4_lattice3(X1)|~v4_lattice3(X2)|~v10_lattices(X1)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, plain, (r3_lattice3(X1,esk4_3(X1,X2,X3),X3)|X2=k16_lattice3(X1,X3)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_53, plain, (r3_lattice3(X1,esk5_2(X1,a_2_9_lattice3(X1,X2)),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_33])).
fof(c_0_54, negated_conjecture, ((~v2_struct_0(esk8_0)&(v10_lattices(esk8_0)&(v4_lattice3(esk8_0)&l3_lattices(esk8_0))))&(m1_subset_1(esk9_0,u1_struct_0(esk8_0))&(r3_lattice3(esk8_0,esk9_0,esk10_0)&~r3_lattices(esk8_0,esk9_0,k16_lattice3(esk8_0,esk10_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_50])])])])).
cnf(c_0_55, plain, (r3_lattices(X2,X1,X4)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattice3(X2,X1,X3)|X4!=k16_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_56, plain, (esk5_2(X1,a_2_9_lattice3(X1,X2))=k16_lattice3(X1,X2)|v2_struct_0(X1)|~m1_subset_1(esk4_3(X1,esk5_2(X1,a_2_9_lattice3(X1,X2)),X2),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_33]), c_0_53])).
cnf(c_0_57, negated_conjecture, (~r3_lattices(esk8_0,esk9_0,k16_lattice3(esk8_0,esk10_0))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_58, plain, (r3_lattices(X1,X2,k16_lattice3(X1,X3))|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(k16_lattice3(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_59, negated_conjecture, (r3_lattice3(esk8_0,esk9_0,esk10_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(esk9_0,u1_struct_0(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, negated_conjecture, (l3_lattices(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, negated_conjecture, (v4_lattice3(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_63, negated_conjecture, (v10_lattices(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64, negated_conjecture, (~v2_struct_0(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_65, plain, (esk5_2(X1,a_2_9_lattice3(X1,X2))=k16_lattice3(X1,X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_30]), c_0_33]), c_0_53])).
cnf(c_0_66, negated_conjecture, (~m1_subset_1(k16_lattice3(esk8_0,esk10_0),u1_struct_0(esk8_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_60]), c_0_61]), c_0_62]), c_0_63])]), c_0_64])).
cnf(c_0_67, plain, (m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_33, c_0_65])).
cnf(c_0_68, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_61]), c_0_62]), c_0_63])]), c_0_64]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 52
# Proof object formula steps           : 17
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 38
# Processed clauses                    : 170
# ...of these trivial                  : 0
# ...subsumed                          : 25
# ...remaining for further processing  : 145
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 7
# Backward-rewritten                   : 0
# Generated clauses                    : 157
# ...of the previous two non-trivial   : 133
# Contextual simplify-reflections      : 45
# Paramodulations                      : 154
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 97
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 88
# Current number of unprocessed clauses: 37
# ...number of literals in the above   : 399
# Current number of archived formulas  : 0
# Current number of archived clauses   : 45
# Clause-clause subsumption calls (NU) : 6479
# Rec. Clause-clause subsumption calls : 189
# Non-unit clause-clause subsumptions  : 67
# Unit Clause-clause subsumption calls : 65
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10765

# -------------------------------------------------
# User time                : 0.017 s
# System time              : 0.000 s
# Total time               : 0.017 s
# Maximum resident set size: 28576 pages
