# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t191_funct_2, conjecture, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>(![X5]:(m1_subset_1(X5,X2)=>r2_hidden(k3_funct_2(X2,X3,X4,X5),X1))=>r1_tarski(k2_relset_1(X2,X3,X4),X1))))), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', t191_funct_2)).
fof(t5_funct_2, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', t5_funct_2)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', redefinition_k3_funct_2)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', cc1_relset_1)).
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', d2_funct_2)).
fof(t11_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((X2=k1_xboole_0=>X1=k1_xboole_0)=>r2_hidden(X3,k1_funct_2(X1,X2)))), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', t11_funct_2)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', t1_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', d1_xboole_0)).
fof(t16_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(![X4]:~((r2_hidden(X4,X2)&![X5]:~((r2_hidden(X5,X1)&X4=k1_funct_1(X3,X5)))))=>k2_relset_1(X1,X2,X3)=X2)), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', t16_funct_2)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', redefinition_k2_relset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', d10_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/b5fc45c6aed24c08b783b7513098b09f.E_input', fc1_xboole_0)).
fof(c_0_12, negated_conjecture, ~(![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(~(v1_xboole_0(X3))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>(![X5]:(m1_subset_1(X5,X2)=>r2_hidden(k3_funct_2(X2,X3,X4,X5),X1))=>r1_tarski(k2_relset_1(X2,X3,X4),X1)))))), inference(assume_negation,[status(cth)],[t191_funct_2])).
fof(c_0_13, plain, ![X52, X53, X54]:(((v1_funct_1(X54)|(r2_hidden(esk8_3(X52,X53,X54),X52)|k1_relat_1(X54)!=X52)|(~v1_relat_1(X54)|~v1_funct_1(X54)))&((v1_funct_2(X54,X52,X53)|(r2_hidden(esk8_3(X52,X53,X54),X52)|k1_relat_1(X54)!=X52)|(~v1_relat_1(X54)|~v1_funct_1(X54)))&(m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X53)))|(r2_hidden(esk8_3(X52,X53,X54),X52)|k1_relat_1(X54)!=X52)|(~v1_relat_1(X54)|~v1_funct_1(X54)))))&((v1_funct_1(X54)|(~r2_hidden(k1_funct_1(X54,esk8_3(X52,X53,X54)),X53)|k1_relat_1(X54)!=X52)|(~v1_relat_1(X54)|~v1_funct_1(X54)))&((v1_funct_2(X54,X52,X53)|(~r2_hidden(k1_funct_1(X54,esk8_3(X52,X53,X54)),X53)|k1_relat_1(X54)!=X52)|(~v1_relat_1(X54)|~v1_funct_1(X54)))&(m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,X53)))|(~r2_hidden(k1_funct_1(X54,esk8_3(X52,X53,X54)),X53)|k1_relat_1(X54)!=X52)|(~v1_relat_1(X54)|~v1_funct_1(X54)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_funct_2])])])])).
fof(c_0_14, negated_conjecture, ![X10]:(~v1_xboole_0(esk2_0)&(~v1_xboole_0(esk3_0)&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk2_0,esk3_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))))&((~m1_subset_1(X10,esk2_0)|r2_hidden(k3_funct_2(esk2_0,esk3_0,esk4_0,X10),esk1_0))&~r1_tarski(k2_relset_1(esk2_0,esk3_0,esk4_0),esk1_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])).
fof(c_0_15, plain, ![X63, X64, X65, X66]:(v1_xboole_0(X63)|(~v1_funct_1(X65)|(~v1_funct_2(X65,X63,X64)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64))))|~m1_subset_1(X66,X63))|k3_funct_2(X63,X64,X65,X66)=k1_funct_1(X65,X66)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
fof(c_0_16, plain, ![X24, X25, X26]:(~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))|v1_relat_1(X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_17, plain, ![X78, X79, X80, X81, X83, X84, X85, X86, X87, X89]:(((((v1_relat_1(esk10_4(X78,X79,X80,X81))|~r2_hidden(X81,X80)|X80!=k1_funct_2(X78,X79))&(v1_funct_1(esk10_4(X78,X79,X80,X81))|~r2_hidden(X81,X80)|X80!=k1_funct_2(X78,X79)))&((X81=esk10_4(X78,X79,X80,X81)|~r2_hidden(X81,X80)|X80!=k1_funct_2(X78,X79))&((k1_relat_1(esk10_4(X78,X79,X80,X81))=X78|~r2_hidden(X81,X80)|X80!=k1_funct_2(X78,X79))&(r1_tarski(k2_relat_1(esk10_4(X78,X79,X80,X81)),X79)|~r2_hidden(X81,X80)|X80!=k1_funct_2(X78,X79)))))&(~v1_relat_1(X84)|~v1_funct_1(X84)|(X83!=X84|(k1_relat_1(X84)!=X78|~r1_tarski(k2_relat_1(X84),X79)))|r2_hidden(X83,X80)|X80!=k1_funct_2(X78,X79)))&((~r2_hidden(esk11_3(X85,X86,X87),X87)|(~v1_relat_1(X89)|~v1_funct_1(X89)|(esk11_3(X85,X86,X87)!=X89|(k1_relat_1(X89)!=X85|~r1_tarski(k2_relat_1(X89),X86))))|X87=k1_funct_2(X85,X86))&(((v1_relat_1(esk12_3(X85,X86,X87))|r2_hidden(esk11_3(X85,X86,X87),X87)|X87=k1_funct_2(X85,X86))&(v1_funct_1(esk12_3(X85,X86,X87))|r2_hidden(esk11_3(X85,X86,X87),X87)|X87=k1_funct_2(X85,X86)))&((esk11_3(X85,X86,X87)=esk12_3(X85,X86,X87)|r2_hidden(esk11_3(X85,X86,X87),X87)|X87=k1_funct_2(X85,X86))&((k1_relat_1(esk12_3(X85,X86,X87))=X85|r2_hidden(esk11_3(X85,X86,X87),X87)|X87=k1_funct_2(X85,X86))&(r1_tarski(k2_relat_1(esk12_3(X85,X86,X87)),X86)|r2_hidden(esk11_3(X85,X86,X87),X87)|X87=k1_funct_2(X85,X86))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
cnf(c_0_18, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(k1_funct_1(X1,esk8_3(X2,X3,X1)),X3)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, negated_conjecture, (r2_hidden(k3_funct_2(esk2_0,esk3_0,esk4_0,X1),esk1_0)|~m1_subset_1(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (v1_funct_2(esk4_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_24, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_25, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (k1_relat_1(esk10_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (X1=esk10_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_28, plain, ![X69, X70, X71]:((X70=k1_xboole_0|r2_hidden(X71,k1_funct_2(X69,X70))|(~v1_funct_1(X71)|(~v1_funct_2(X71,X69,X70)|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(X69,X70))))))&(X69!=k1_xboole_0|r2_hidden(X71,k1_funct_2(X69,X70))|(~v1_funct_1(X71)|(~v1_funct_2(X71,X69,X70)|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(X69,X70))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_funct_2])])])).
cnf(c_0_29, plain, (v1_funct_2(X1,X2,X3)|~r2_hidden(k1_funct_1(X1,esk8_3(X2,X3,X1)),X3)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),X2)))|~v1_relat_1(X1)|~r2_hidden(k1_funct_1(X1,esk8_3(k1_relat_1(X1),X2,X1)),X2)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_31, negated_conjecture, (r2_hidden(k1_funct_1(esk4_0,X1),esk1_0)|~m1_subset_1(X1,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_24])).
cnf(c_0_32, negated_conjecture, (v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_25, c_0_21])).
fof(c_0_33, plain, ![X61, X62]:(~r2_hidden(X61,X62)|m1_subset_1(X61,X62)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_34, plain, (k1_relat_1(esk10_4(X1,X2,k1_funct_2(X1,X2),X3))=X1|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (esk10_4(X1,X2,k1_funct_2(X1,X2),X3)=X3|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (X1=k1_xboole_0|r2_hidden(X2,k1_funct_2(X3,X1))|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (v1_funct_2(X1,k1_relat_1(X1),X2)|~v1_relat_1(X1)|~r2_hidden(k1_funct_1(X1,esk8_3(k1_relat_1(X1),X2,X1)),X2)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_29])).
fof(c_0_38, plain, ![X72, X73, X74]:((~v1_xboole_0(X72)|~r2_hidden(X73,X72))&(r2_hidden(esk9_1(X74),X74)|v1_xboole_0(X74))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_39, plain, ![X47, X48, X49, X51]:((r2_hidden(esk7_3(X47,X48,X49),X48)|k2_relset_1(X47,X48,X49)=X48|(~v1_funct_1(X49)|(~v1_funct_2(X49,X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48))))))&(~r2_hidden(X51,X47)|esk7_3(X47,X48,X49)!=k1_funct_1(X49,X51)|k2_relset_1(X47,X48,X49)=X48|(~v1_funct_1(X49)|(~v1_funct_2(X49,X47,X48)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_funct_2])])])])])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk4_0),esk1_0)))|~m1_subset_1(esk8_3(k1_relat_1(esk4_0),esk1_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_23])])).
cnf(c_0_41, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk4_0,k1_funct_2(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_21]), c_0_22]), c_0_23])])).
cnf(c_0_44, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|r2_hidden(esk8_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_45, negated_conjecture, (v1_funct_2(esk4_0,k1_relat_1(esk4_0),esk1_0)|~m1_subset_1(esk8_3(k1_relat_1(esk4_0),esk1_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_31]), c_0_32]), c_0_23])])).
cnf(c_0_46, plain, (v1_funct_2(X1,X2,X3)|r2_hidden(esk8_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_47, plain, ![X34, X35, X36]:(~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X34,X35)))|k2_relset_1(X34,X35,X36)=k2_relat_1(X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_48, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49, plain, (r2_hidden(esk7_3(X1,X2,X3),X2)|k2_relset_1(X1,X2,X3)=X2|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk4_0),esk1_0)))|~r2_hidden(esk8_3(k1_relat_1(esk4_0),esk1_0,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51, negated_conjecture, (k1_relat_1(esk4_0)=esk2_0|esk3_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_52, plain, (r2_hidden(esk8_3(k1_relat_1(X1),X2,X1),k1_relat_1(X1))|m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),X2)))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_53, negated_conjecture, (v1_funct_2(esk4_0,k1_relat_1(esk4_0),esk1_0)|~r2_hidden(esk8_3(k1_relat_1(esk4_0),esk1_0,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_41])).
cnf(c_0_54, plain, (r2_hidden(esk8_3(k1_relat_1(X1),X2,X1),k1_relat_1(X1))|v1_funct_2(X1,k1_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (k2_relset_1(X1,X2,X3)=X2|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_57, negated_conjecture, (esk3_0=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))|~r2_hidden(esk8_3(esk2_0,esk1_0,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk8_3(esk2_0,X1,esk4_0),esk2_0)|m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_51]), c_0_32]), c_0_23])])).
cnf(c_0_59, negated_conjecture, (esk3_0=k1_xboole_0|v1_funct_2(esk4_0,esk2_0,esk1_0)|~r2_hidden(esk8_3(esk2_0,esk1_0,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_53, c_0_51])).
cnf(c_0_60, negated_conjecture, (esk3_0=k1_xboole_0|r2_hidden(esk8_3(esk2_0,X1,esk4_0),esk2_0)|v1_funct_2(esk4_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_51]), c_0_32]), c_0_23])])).
fof(c_0_61, plain, ![X27, X28]:(((r1_tarski(X27,X28)|X27!=X28)&(r1_tarski(X28,X27)|X27!=X28))&(~r1_tarski(X27,X28)|~r1_tarski(X28,X27)|X27=X28)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_62, plain, (r1_tarski(k2_relat_1(esk10_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_63, negated_conjecture, (~r1_tarski(k2_relset_1(esk2_0,esk3_0,esk4_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_64, plain, (X1=k2_relat_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(X2,X3,X1)|~v1_funct_1(X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_65, negated_conjecture, (esk3_0=k1_xboole_0|m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk1_0)))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_66, negated_conjecture, (esk3_0=k1_xboole_0|v1_funct_2(esk4_0,esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_67, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68, plain, (r1_tarski(k2_relat_1(esk10_4(X1,X2,k1_funct_2(X1,X2),X3)),X2)|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_69, negated_conjecture, (~r1_tarski(k2_relat_1(esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_55]), c_0_21])])).
cnf(c_0_70, negated_conjecture, (k2_relat_1(esk4_0)=esk1_0|esk3_0=k1_xboole_0|~v1_xboole_0(esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_23])]), c_0_66])).
cnf(c_0_71, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_67])).
cnf(c_0_72, plain, (r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(spm,[status(thm)],[c_0_68, c_0_35])).
cnf(c_0_73, negated_conjecture, (esk3_0=k1_xboole_0|esk1_0=k1_xboole_0|r2_hidden(esk4_0,k1_funct_2(esk2_0,esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_65]), c_0_23])]), c_0_66])).
cnf(c_0_74, negated_conjecture, (esk3_0=k1_xboole_0|~v1_xboole_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])])).
cnf(c_0_75, negated_conjecture, (esk1_0=k1_xboole_0|esk3_0=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_69])).
cnf(c_0_76, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_77, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_78, negated_conjecture, (esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])])).
cnf(c_0_79, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_78]), c_0_76])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 80
# Proof object clause steps            : 56
# Proof object formula steps           : 24
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 24
# Proof object simplifying inferences  : 44
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 25
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 66
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 64
# Processed clauses                    : 11245
# ...of these trivial                  : 27
# ...subsumed                          : 9194
# ...remaining for further processing  : 2024
# Other redundant clauses eliminated   : 50
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 129
# Backward-rewritten                   : 848
# Generated clauses                    : 98842
# ...of the previous two non-trivial   : 94168
# Contextual simplify-reflections      : 61
# Paramodulations                      : 98757
# Factorizations                       : 38
# Equation resolutions                 : 50
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 962
#    Positive orientable unit clauses  : 36
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 25
#    Non-unit-clauses                  : 901
# Current number of unprocessed clauses: 81911
# ...number of literals in the above   : 415516
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1040
# Clause-clause subsumption calls (NU) : 376594
# Rec. Clause-clause subsumption calls : 146704
# Non-unit clause-clause subsumptions  : 6219
# Unit Clause-clause subsumption calls : 4372
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 174
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1977163

# -------------------------------------------------
# User time                : 1.463 s
# System time              : 0.023 s
# Total time               : 1.486 s
# Maximum resident set size: 73668 pages
