# No SInE strategy applied
# Auto-Mode selected heuristic H_____042_B03_F1_AE_Q4_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t49_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k2_pre_topc(X1,X2),k2_pre_topc(X1,X3)))))), file('data/MPTP/proofs/b6027864fece4836b393625f30525f19.E_input', t49_pre_topc)).
fof(dt_k2_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/b6027864fece4836b393625f30525f19.E_input', dt_k2_pre_topc)).
fof(t7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,X2)=>r2_hidden(X4,X3)))=>r1_tarski(X2,X3)))), file('data/MPTP/proofs/b6027864fece4836b393625f30525f19.E_input', t7_subset_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/b6027864fece4836b393625f30525f19.E_input', l3_subset_1)).
fof(t45_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(r2_hidden(X3,u1_struct_0(X1))=>(r2_hidden(X3,k2_pre_topc(X1,X2))<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X4,X1)&r1_tarski(X2,X4))=>r2_hidden(X3,X4))))))), file('data/MPTP/proofs/b6027864fece4836b393625f30525f19.E_input', t45_pre_topc)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/b6027864fece4836b393625f30525f19.E_input', t3_subset)).
fof(c_0_6, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k2_pre_topc(X1,X2),k2_pre_topc(X1,X3))))))), inference(assume_negation,[status(cth)],[t49_pre_topc])).
fof(c_0_7, plain, ![X8, X9]:(~l1_pre_topc(X8)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|m1_subset_1(k2_pre_topc(X8,X9),k1_zfmisc_1(u1_struct_0(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_pre_topc])])).
fof(c_0_8, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(r1_tarski(esk2_0,esk3_0)&~r1_tarski(k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
fof(c_0_9, plain, ![X10, X11, X12]:((m1_subset_1(esk4_3(X10,X11,X12),X10)|r1_tarski(X11,X12)|~m1_subset_1(X12,k1_zfmisc_1(X10))|~m1_subset_1(X11,k1_zfmisc_1(X10)))&((r2_hidden(esk4_3(X10,X11,X12),X11)|r1_tarski(X11,X12)|~m1_subset_1(X12,k1_zfmisc_1(X10))|~m1_subset_1(X11,k1_zfmisc_1(X10)))&(~r2_hidden(esk4_3(X10,X11,X12),X12)|r1_tarski(X11,X12)|~m1_subset_1(X12,k1_zfmisc_1(X10))|~m1_subset_1(X11,k1_zfmisc_1(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_subset_1])])])])])).
cnf(c_0_10, plain, (m1_subset_1(k2_pre_topc(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_13, plain, ![X14, X15, X16]:(~m1_subset_1(X15,k1_zfmisc_1(X14))|(~r2_hidden(X16,X15)|r2_hidden(X16,X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|r1_tarski(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(k2_pre_topc(esk1_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])])).
cnf(c_0_17, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (m1_subset_1(k2_pre_topc(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_14]), c_0_12])])).
cnf(c_0_19, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),X1,k2_pre_topc(esk1_0,esk3_0)),X1)|r1_tarski(X1,k2_pre_topc(esk1_0,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, negated_conjecture, (~r1_tarski(k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_21, plain, ![X17, X18, X19, X20]:((~r2_hidden(X19,k2_pre_topc(X17,X18))|(~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X17)))|(~v4_pre_topc(X20,X17)|~r1_tarski(X18,X20)|r2_hidden(X19,X20)))|~r2_hidden(X19,u1_struct_0(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_pre_topc(X17))&((m1_subset_1(esk5_3(X17,X18,X19),k1_zfmisc_1(u1_struct_0(X17)))|r2_hidden(X19,k2_pre_topc(X17,X18))|~r2_hidden(X19,u1_struct_0(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_pre_topc(X17))&(((v4_pre_topc(esk5_3(X17,X18,X19),X17)|r2_hidden(X19,k2_pre_topc(X17,X18))|~r2_hidden(X19,u1_struct_0(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_pre_topc(X17))&(r1_tarski(X18,esk5_3(X17,X18,X19))|r2_hidden(X19,k2_pre_topc(X17,X18))|~r2_hidden(X19,u1_struct_0(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_pre_topc(X17)))&(~r2_hidden(X19,esk5_3(X17,X18,X19))|r2_hidden(X19,k2_pre_topc(X17,X18))|~r2_hidden(X19,u1_struct_0(X17))|~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X17)))|~l1_pre_topc(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_pre_topc])])])])])).
cnf(c_0_22, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k2_pre_topc(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_18]), c_0_20])).
cnf(c_0_24, plain, (r1_tarski(X2,X3)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, plain, (v4_pre_topc(esk5_3(X1,X2,X3),X1)|r2_hidden(X3,k2_pre_topc(X1,X2))|~r2_hidden(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_26, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27, negated_conjecture, (r1_tarski(X1,k2_pre_topc(esk1_0,esk3_0))|~r2_hidden(esk4_3(u1_struct_0(esk1_0),X1,k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_24, c_0_16])).
cnf(c_0_28, plain, (r2_hidden(X1,k2_pre_topc(X2,X3))|~r2_hidden(X1,esk5_3(X2,X3,X1))|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,k2_pre_topc(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~v4_pre_topc(X4,X2)|~r1_tarski(X3,X4)|~r2_hidden(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, negated_conjecture, (v4_pre_topc(esk5_3(esk1_0,X1,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))),esk1_0)|r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_12])])).
cnf(c_0_31, negated_conjecture, (~r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_18]), c_0_20])).
cnf(c_0_32, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,X1))|~r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),esk5_3(esk1_0,X1,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_26]), c_0_12])])).
cnf(c_0_33, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|r2_hidden(X3,k2_pre_topc(X1,X2))|~r2_hidden(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),X1)|~v4_pre_topc(X1,esk1_0)|~r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,X2))|~r1_tarski(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_26]), c_0_12])])).
cnf(c_0_35, negated_conjecture, (v4_pre_topc(esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))),esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_11]), c_0_31])).
cnf(c_0_36, negated_conjecture, (~r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_11]), c_0_31])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,X1))|m1_subset_1(esk5_3(esk1_0,X1,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_26]), c_0_12])])).
cnf(c_0_38, plain, (r1_tarski(X1,esk5_3(X2,X1,X3))|r2_hidden(X3,k2_pre_topc(X2,X1))|~r2_hidden(X3,u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_39, plain, ![X22, X23]:((~m1_subset_1(X22,k1_zfmisc_1(X23))|r1_tarski(X22,X23))&(~r1_tarski(X22,X23)|m1_subset_1(X22,k1_zfmisc_1(X23)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_40, negated_conjecture, (~r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,X1))|~r1_tarski(X1,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))|~m1_subset_1(esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_11]), c_0_31])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,X1))|r1_tarski(X1,esk5_3(esk1_0,X1,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_26]), c_0_12])])).
cnf(c_0_43, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44, negated_conjecture, (r1_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_45, negated_conjecture, (~r2_hidden(esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)),k2_pre_topc(esk1_0,X1))|~r1_tarski(X1,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])])).
cnf(c_0_46, negated_conjecture, (r1_tarski(esk3_0,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_11]), c_0_31])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk3_0))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),X1,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)))),X1)|r1_tarski(X1,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_15, c_0_41])).
cnf(c_0_49, negated_conjecture, (~r1_tarski(esk2_0,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_14]), c_0_23])])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)))))), inference(spm,[status(thm)],[c_0_43, c_0_46])).
cnf(c_0_51, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_17, c_0_47])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),esk2_0,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)))),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_14]), c_0_49])).
cnf(c_0_53, negated_conjecture, (r1_tarski(X1,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))|~r2_hidden(esk4_3(u1_struct_0(esk1_0),X1,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)))),esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_24, c_0_41])).
cnf(c_0_54, negated_conjecture, (r2_hidden(X1,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_17, c_0_50])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk4_3(u1_struct_0(esk1_0),esk2_0,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)))),esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56, negated_conjecture, (~r2_hidden(esk4_3(u1_struct_0(esk1_0),esk2_0,esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0)))),esk5_3(esk1_0,esk3_0,esk4_3(u1_struct_0(esk1_0),k2_pre_topc(esk1_0,esk2_0),k2_pre_topc(esk1_0,esk3_0))))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_14]), c_0_49])).
cnf(c_0_57, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 45
# Proof object formula steps           : 13
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 880
# ...of these trivial                  : 1
# ...subsumed                          : 115
# ...remaining for further processing  : 764
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 37
# Backward-rewritten                   : 72
# Generated clauses                    : 6763
# ...of the previous two non-trivial   : 6689
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6763
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 655
#    Positive orientable unit clauses  : 85
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 565
# Current number of unprocessed clauses: 5807
# ...number of literals in the above   : 22438
# Current number of archived formulas  : 0
# Current number of archived clauses   : 109
# Clause-clause subsumption calls (NU) : 33741
# Rec. Clause-clause subsumption calls : 20321
# Non-unit clause-clause subsumptions  : 151
# Unit Clause-clause subsumption calls : 1955
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 623
# BW rewrite match successes           : 8
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 376388

# -------------------------------------------------
# User time                : 0.108 s
# System time              : 0.000 s
# Total time               : 0.108 s
# Maximum resident set size: 28516 pages
