# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t68_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(~((~(X2=k1_xboole_0)&m1_orders_2(X2,X1,X2)))&~((~(m1_orders_2(X2,X1,X2))&X2=k1_xboole_0))))), file('data/MPTP/proofs/b644968ef2e349e78720de4a4f34595d.E_input', t68_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/b644968ef2e349e78720de4a4f34595d.E_input', t4_subset)).
fof(d15_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((~(X2=k1_xboole_0)=>(m1_orders_2(X3,X1,X2)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&(r2_hidden(X4,X2)&X3=k3_orders_2(X1,X2,X4)))))&(X2=k1_xboole_0=>(m1_orders_2(X3,X1,X2)<=>X3=k1_xboole_0)))))), file('data/MPTP/proofs/b644968ef2e349e78720de4a4f34595d.E_input', d15_orders_2)).
fof(dt_m1_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>![X3]:(m1_orders_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/b644968ef2e349e78720de4a4f34595d.E_input', dt_m1_orders_2)).
fof(irreflexivity_r2_orders_2, axiom, ![X1, X2, X3]:((l1_orders_2(X1)&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>~(r2_orders_2(X1,X2,X2))), file('data/MPTP/proofs/b644968ef2e349e78720de4a4f34595d.E_input', irreflexivity_r2_orders_2)).
fof(t57_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X2,k3_orders_2(X1,X4,X3))<=>(r2_orders_2(X1,X2,X3)&r2_hidden(X2,X4))))))), file('data/MPTP/proofs/b644968ef2e349e78720de4a4f34595d.E_input', t57_orders_2)).
fof(c_0_6, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(~((~(X2=k1_xboole_0)&m1_orders_2(X2,X1,X2)))&~((~(m1_orders_2(X2,X1,X2))&X2=k1_xboole_0)))))), inference(assume_negation,[status(cth)],[t68_orders_2])).
fof(c_0_7, plain, ![X7, X8, X9]:(~r2_hidden(X7,X8)|~m1_subset_1(X8,k1_zfmisc_1(X9))|m1_subset_1(X7,X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_8, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(((~m1_orders_2(esk2_0,esk1_0,esk2_0)|esk2_0!=k1_xboole_0)&(esk2_0=k1_xboole_0|esk2_0!=k1_xboole_0))&((~m1_orders_2(esk2_0,esk1_0,esk2_0)|m1_orders_2(esk2_0,esk1_0,esk2_0))&(esk2_0=k1_xboole_0|m1_orders_2(esk2_0,esk1_0,esk2_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_6])])])])])).
fof(c_0_9, plain, ![X10, X11, X12, X14]:((((m1_subset_1(esk3_3(X10,X11,X12),u1_struct_0(X10))|~m1_orders_2(X12,X10,X11)|X11=k1_xboole_0|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v2_struct_0(X10)|(~v3_orders_2(X10)|(~v4_orders_2(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10))))))&((r2_hidden(esk3_3(X10,X11,X12),X11)|~m1_orders_2(X12,X10,X11)|X11=k1_xboole_0|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v2_struct_0(X10)|(~v3_orders_2(X10)|(~v4_orders_2(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10))))))&(X12=k3_orders_2(X10,X11,esk3_3(X10,X11,X12))|~m1_orders_2(X12,X10,X11)|X11=k1_xboole_0|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v2_struct_0(X10)|(~v3_orders_2(X10)|(~v4_orders_2(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10))))))))&(~m1_subset_1(X14,u1_struct_0(X10))|(~r2_hidden(X14,X11)|X12!=k3_orders_2(X10,X11,X14))|m1_orders_2(X12,X10,X11)|X11=k1_xboole_0|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v2_struct_0(X10)|(~v3_orders_2(X10)|(~v4_orders_2(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10)))))))&((~m1_orders_2(X12,X10,X11)|X12=k1_xboole_0|X11!=k1_xboole_0|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v2_struct_0(X10)|(~v3_orders_2(X10)|(~v4_orders_2(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10))))))&(X12!=k1_xboole_0|m1_orders_2(X12,X10,X11)|X11!=k1_xboole_0|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X10)))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v2_struct_0(X10)|(~v3_orders_2(X10)|(~v4_orders_2(X10)|(~v5_orders_2(X10)|~l1_orders_2(X10)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d15_orders_2])])])])])])).
fof(c_0_10, plain, ![X15, X16, X17]:(v2_struct_0(X15)|(~v3_orders_2(X15)|(~v4_orders_2(X15)|(~v5_orders_2(X15)|~l1_orders_2(X15))))|~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))|(~m1_orders_2(X17,X15,X16)|m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_orders_2])])])])).
cnf(c_0_11, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (r2_hidden(esk3_3(X1,X2,X3),X2)|X2=k1_xboole_0|v2_struct_0(X1)|~m1_orders_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_orders_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (~m1_orders_2(esk2_0,esk1_0,esk2_0)|esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, plain, (m1_orders_2(X1,X2,X3)|v2_struct_0(X2)|X1!=k1_xboole_0|X3!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_18, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_19, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_20, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_21, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_22, plain, ![X18, X19, X20]:(~l1_orders_2(X18)|(~m1_subset_1(X19,u1_struct_0(X18))|~m1_subset_1(X20,u1_struct_0(X18)))|~r2_orders_2(X18,X19,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[irreflexivity_r2_orders_2])])])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_24, plain, (X1=k1_xboole_0|r2_hidden(esk3_3(X2,X1,X3),X1)|v2_struct_0(X2)|~m1_orders_2(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_25, negated_conjecture, (esk2_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_12]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
fof(c_0_26, plain, ![X21, X22, X23, X24]:(((r2_orders_2(X21,X22,X23)|~r2_hidden(X22,k3_orders_2(X21,X24,X23))|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X21)))|~m1_subset_1(X23,u1_struct_0(X21))|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))))&(r2_hidden(X22,X24)|~r2_hidden(X22,k3_orders_2(X21,X24,X23))|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X21)))|~m1_subset_1(X23,u1_struct_0(X21))|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21)))))))&(~r2_orders_2(X21,X22,X23)|~r2_hidden(X22,X24)|r2_hidden(X22,k3_orders_2(X21,X24,X23))|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(X21)))|~m1_subset_1(X23,u1_struct_0(X21))|~m1_subset_1(X22,u1_struct_0(X21))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t57_orders_2])])])])])).
cnf(c_0_27, plain, (X1=k3_orders_2(X2,X3,esk3_3(X2,X3,X1))|X3=k1_xboole_0|v2_struct_0(X2)|~m1_orders_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_28, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X1))|X2=k1_xboole_0|v2_struct_0(X1)|~m1_orders_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29, plain, (~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r2_orders_2(X1,X2,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk3_3(X1,esk2_0,X2),u1_struct_0(esk1_0))|v2_struct_0(X1)|~m1_orders_2(X2,X1,esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_31, plain, (r2_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r2_hidden(X2,k3_orders_2(X1,X4,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (k3_orders_2(X1,X2,esk3_3(X1,X2,X3))=X3|X2=k1_xboole_0|v2_struct_0(X1)|~m1_orders_2(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_27, c_0_14])).
cnf(c_0_33, plain, (X1=k1_xboole_0|m1_subset_1(esk3_3(X2,X1,X3),u1_struct_0(X2))|v2_struct_0(X2)|~m1_orders_2(X3,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_28, c_0_14])).
cnf(c_0_34, negated_conjecture, (v2_struct_0(X1)|~r2_orders_2(esk1_0,X2,X2)|~m1_orders_2(X3,X1,esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_17])])).
cnf(c_0_35, plain, (X1=k1_xboole_0|r2_orders_2(X2,X3,esk3_3(X2,X1,X4))|v2_struct_0(X2)|~r2_hidden(X3,X4)|~m1_orders_2(X4,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_36, negated_conjecture, (X1=k1_xboole_0|v2_struct_0(X2)|~r2_hidden(esk3_3(esk1_0,X1,X3),X3)|~m1_orders_2(X4,X2,esk2_0)|~m1_orders_2(X3,esk1_0,X1)|~m1_subset_1(esk3_3(esk1_0,X1,X3),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_37, negated_conjecture, (esk2_0=k1_xboole_0|m1_orders_2(esk2_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_38, negated_conjecture, (X1=k1_xboole_0|v2_struct_0(X2)|~m1_orders_2(X3,X2,esk2_0)|~m1_orders_2(X1,esk1_0,X1)|~m1_subset_1(esk3_3(esk1_0,X1,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_24]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_21])).
cnf(c_0_39, negated_conjecture, (m1_orders_2(esk2_0,esk1_0,esk2_0)), inference(sr,[status(thm)],[c_0_37, c_0_25])).
cnf(c_0_40, negated_conjecture, (v2_struct_0(X1)|~m1_orders_2(X2,X1,esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_30]), c_0_39]), c_0_12]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_25]), c_0_21])).
cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_39]), c_0_12]), c_0_17]), c_0_18]), c_0_19]), c_0_20])]), c_0_21]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 29
# Proof object formula steps           : 13
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 43
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 20
# Processed clauses                    : 76
# ...of these trivial                  : 0
# ...subsumed                          : 9
# ...remaining for further processing  : 67
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 1
# Generated clauses                    : 74
# ...of the previous two non-trivial   : 55
# Contextual simplify-reflections      : 12
# Paramodulations                      : 71
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 31
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 23
# Current number of unprocessed clauses: 19
# ...number of literals in the above   : 321
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 3221
# Rec. Clause-clause subsumption calls : 212
# Non-unit clause-clause subsumptions  : 33
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6772

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28680 pages
