# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', t3_boole)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', redefinition_k6_subset_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', t5_subset)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', t4_subset_1)).
fof(t29_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1)))), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', t29_yellow_6)).
fof(t10_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>~(r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', t10_waybel_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', d3_tarski)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', d12_waybel_0)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', existence_m1_subset_1)).
fof(t8_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3, X4]:(r1_tarski(X3,X4)=>((r1_waybel_0(X1,X2,X3)=>r1_waybel_0(X1,X2,X4))&(r2_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X4)))))), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', t8_waybel_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/b7b836d04f4b4668ab38bbbbdc4b9024.E_input', fc1_xboole_0)).
fof(c_0_11, plain, ![X17]:k4_xboole_0(X17,k1_xboole_0)=X17, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_12, plain, ![X9, X10]:k6_subset_1(X9,X10)=k4_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_13, plain, ![X11, X12, X13]:(~r2_hidden(X11,X12)|(~m1_subset_1(X12,k1_zfmisc_1(X13))|~v1_xboole_0(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_14, plain, ![X14]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X14)), inference(variable_rename,[status(thm)],[t4_subset_1])).
fof(c_0_15, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t29_yellow_6])).
fof(c_0_16, plain, ![X6, X7, X8]:((~r2_waybel_0(X6,X7,X8)|~r1_waybel_0(X6,X7,k6_subset_1(u1_struct_0(X6),X8))|(v2_struct_0(X7)|~l1_waybel_0(X7,X6))|(v2_struct_0(X6)|~l1_struct_0(X6)))&(r1_waybel_0(X6,X7,k6_subset_1(u1_struct_0(X6),X8))|r2_waybel_0(X6,X7,X8)|(v2_struct_0(X7)|~l1_waybel_0(X7,X6))|(v2_struct_0(X6)|~l1_struct_0(X6)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_waybel_0])])])])])).
cnf(c_0_17, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X18, X19, X20, X21, X22]:((~r1_tarski(X18,X19)|(~r2_hidden(X20,X18)|r2_hidden(X20,X19)))&((r2_hidden(esk3_2(X21,X22),X21)|r1_tarski(X21,X22))&(~r2_hidden(esk3_2(X21,X22),X22)|r1_tarski(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_22, plain, ![X24, X25, X26, X27, X29, X31]:(((m1_subset_1(esk4_4(X24,X25,X26,X27),u1_struct_0(X25))|~m1_subset_1(X27,u1_struct_0(X25))|~r2_waybel_0(X24,X25,X26)|(v2_struct_0(X25)|~l1_waybel_0(X25,X24))|(v2_struct_0(X24)|~l1_struct_0(X24)))&((r1_orders_2(X25,X27,esk4_4(X24,X25,X26,X27))|~m1_subset_1(X27,u1_struct_0(X25))|~r2_waybel_0(X24,X25,X26)|(v2_struct_0(X25)|~l1_waybel_0(X25,X24))|(v2_struct_0(X24)|~l1_struct_0(X24)))&(r2_hidden(k2_waybel_0(X24,X25,esk4_4(X24,X25,X26,X27)),X26)|~m1_subset_1(X27,u1_struct_0(X25))|~r2_waybel_0(X24,X25,X26)|(v2_struct_0(X25)|~l1_waybel_0(X25,X24))|(v2_struct_0(X24)|~l1_struct_0(X24)))))&((m1_subset_1(esk5_3(X24,X25,X29),u1_struct_0(X25))|r2_waybel_0(X24,X25,X29)|(v2_struct_0(X25)|~l1_waybel_0(X25,X24))|(v2_struct_0(X24)|~l1_struct_0(X24)))&(~m1_subset_1(X31,u1_struct_0(X25))|(~r1_orders_2(X25,esk5_3(X24,X25,X29),X31)|~r2_hidden(k2_waybel_0(X24,X25,X31),X29))|r2_waybel_0(X24,X25,X29)|(v2_struct_0(X25)|~l1_waybel_0(X25,X24))|(v2_struct_0(X24)|~l1_struct_0(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
fof(c_0_23, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])).
cnf(c_0_24, plain, (r1_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (k6_subset_1(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_26, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_27, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (r2_hidden(k2_waybel_0(X1,X2,esk4_4(X1,X2,X3,X4)),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_29, plain, ![X36]:m1_subset_1(esk6_1(X36),X36), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
fof(c_0_30, plain, ![X32, X33, X34, X35]:((~r1_waybel_0(X32,X33,X34)|r1_waybel_0(X32,X33,X35)|~r1_tarski(X34,X35)|(v2_struct_0(X33)|~l1_waybel_0(X33,X32))|(v2_struct_0(X32)|~l1_struct_0(X32)))&(~r2_waybel_0(X32,X33,X34)|r2_waybel_0(X32,X33,X35)|~r1_tarski(X34,X35)|(v2_struct_0(X33)|~l1_waybel_0(X33,X32))|(v2_struct_0(X32)|~l1_struct_0(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_waybel_0])])])])])).
cnf(c_0_31, negated_conjecture, (~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|r2_waybel_0(X1,X2,k1_xboole_0)|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_33, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37, plain, (r1_tarski(k1_xboole_0,X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_38, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_39, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(X3)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,k1_xboole_0)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_28])).
cnf(c_0_40, plain, (m1_subset_1(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (r2_waybel_0(X1,X2,X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r2_waybel_0(X1,X2,X3)|~r1_tarski(X3,X4)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33]), c_0_34])]), c_0_35]), c_0_36])).
cnf(c_0_43, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v1_xboole_0(X3)|~r2_waybel_0(X2,X1,k1_xboole_0)|~l1_waybel_0(X1,X2)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_45, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_33]), c_0_34])]), c_0_36]), c_0_35])).
cnf(c_0_46, negated_conjecture, (~v1_xboole_0(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_33]), c_0_34])]), c_0_36]), c_0_35])).
cnf(c_0_47, plain, ($false), inference(sr,[status(thm)],[c_0_38, c_0_46]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 26
# Proof object formula steps           : 22
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 24
# Processed clauses                    : 68
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 64
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 31
# ...of the previous two non-trivial   : 25
# Contextual simplify-reflections      : 0
# Paramodulations                      : 30
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 36
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 5
# ...number of literals in the above   : 47
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 156
# Rec. Clause-clause subsumption calls : 13
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 28
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2672

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28680 pages
