# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', commutativity_k3_xboole_0)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t21_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', commutativity_k2_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', d5_xboole_0)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t39_xboole_1)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t47_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t3_xboole_0)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t22_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', idempotence_k3_xboole_0)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t16_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t28_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t41_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', l32_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', d7_xboole_0)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t1_boole)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t3_boole)).
fof(t48_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/b7dd6449178b4cb1a81588981ac67295.E_input', t48_xboole_1)).
fof(c_0_18, plain, ![X37, X38]:r1_tarski(k3_xboole_0(X37,X38),X37), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_19, plain, ![X49, X50]:k3_xboole_0(X49,X50)=k3_xboole_0(X50,X49), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_20, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X15, X16]:k3_xboole_0(X15,k2_xboole_0(X15,X16))=X15, inference(variable_rename,[status(thm)],[t21_xboole_1])).
cnf(c_0_23, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_25, plain, ![X39, X40]:k2_xboole_0(X39,X40)=k2_xboole_0(X40,X39), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_26, plain, ![X20, X21, X22, X23, X24, X25, X26, X27]:((((r2_hidden(X23,X20)|~r2_hidden(X23,X22)|X22!=k4_xboole_0(X20,X21))&(~r2_hidden(X23,X21)|~r2_hidden(X23,X22)|X22!=k4_xboole_0(X20,X21)))&(~r2_hidden(X24,X20)|r2_hidden(X24,X21)|r2_hidden(X24,X22)|X22!=k4_xboole_0(X20,X21)))&((~r2_hidden(esk3_3(X25,X26,X27),X27)|(~r2_hidden(esk3_3(X25,X26,X27),X25)|r2_hidden(esk3_3(X25,X26,X27),X26))|X27=k4_xboole_0(X25,X26))&((r2_hidden(esk3_3(X25,X26,X27),X25)|r2_hidden(esk3_3(X25,X26,X27),X27)|X27=k4_xboole_0(X25,X26))&(~r2_hidden(esk3_3(X25,X26,X27),X26)|r2_hidden(esk3_3(X25,X26,X27),X27)|X27=k4_xboole_0(X25,X26))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_27, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_29, plain, ![X13, X14]:k2_xboole_0(X13,k4_xboole_0(X14,X13))=k2_xboole_0(X13,X14), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_30, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X11, X12]:k4_xboole_0(X11,k3_xboole_0(X11,X12))=k4_xboole_0(X11,X12), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_32, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_33, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_34, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_36, plain, ![X29, X30, X32, X33, X34]:(((r2_hidden(esk4_2(X29,X30),X29)|r1_xboole_0(X29,X30))&(r2_hidden(esk4_2(X29,X30),X30)|r1_xboole_0(X29,X30)))&(~r2_hidden(X34,X32)|~r2_hidden(X34,X33)|~r1_xboole_0(X32,X33))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_37, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_38, plain, ![X46, X47]:k2_xboole_0(X46,k3_xboole_0(X46,X47))=X46, inference(variable_rename,[status(thm)],[t22_xboole_1])).
fof(c_0_39, plain, ![X10]:k3_xboole_0(X10,X10)=X10, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_40, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41, plain, (r2_hidden(esk4_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_42, plain, ![X17, X18, X19]:k3_xboole_0(k3_xboole_0(X17,X18),X19)=k3_xboole_0(X17,k3_xboole_0(X18,X19)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
fof(c_0_43, plain, ![X8, X9]:(~r1_tarski(X8,X9)|k3_xboole_0(X8,X9)=X8), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_44, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_37, c_0_28])).
cnf(c_0_45, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_46, plain, ![X41, X42, X43]:k4_xboole_0(k4_xboole_0(X41,X42),X43)=k4_xboole_0(X41,k2_xboole_0(X42,X43)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_47, plain, ![X44, X45]:((k4_xboole_0(X44,X45)!=k1_xboole_0|r1_tarski(X44,X45))&(~r1_tarski(X44,X45)|k4_xboole_0(X44,X45)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_48, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_49, plain, ![X35, X36]:((~r1_xboole_0(X35,X36)|k3_xboole_0(X35,X36)=k1_xboole_0)&(k3_xboole_0(X35,X36)!=k1_xboole_0|r1_xboole_0(X35,X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_50, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r2_hidden(esk4_2(X1,k3_xboole_0(X2,X3)),k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_52, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_35])).
cnf(c_0_55, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_20, c_0_48])).
fof(c_0_58, plain, ![X7]:k2_xboole_0(X7,k1_xboole_0)=X7, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_59, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (r1_xboole_0(k4_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X3,k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_21, c_0_52])).
cnf(c_0_62, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_21])).
fof(c_0_63, plain, ![X48]:k4_xboole_0(X48,k1_xboole_0)=X48, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_64, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_52, c_0_48])).
cnf(c_0_65, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_66, plain, (k2_xboole_0(X1,k4_xboole_0(X2,k2_xboole_0(X3,X1)))=k2_xboole_0(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_33, c_0_55])).
cnf(c_0_67, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_68, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_69, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_61]), c_0_21]), c_0_62])).
cnf(c_0_70, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_71, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_64, c_0_21])).
cnf(c_0_72, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_65])).
cnf(c_0_73, plain, (k4_xboole_0(X1,k3_xboole_0(X2,X1))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_35, c_0_21])).
cnf(c_0_74, plain, (k2_xboole_0(X1,k4_xboole_0(k2_xboole_0(X2,X1),X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_75, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_69]), c_0_70])).
cnf(c_0_76, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_24, c_0_28])).
cnf(c_0_77, plain, (k3_xboole_0(X1,X2)=X2|k4_xboole_0(X2,X1)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])).
cnf(c_0_78, plain, (k2_xboole_0(X1,k4_xboole_0(k2_xboole_0(X1,X2),X2))=X1), inference(spm,[status(thm)],[c_0_74, c_0_28])).
cnf(c_0_79, plain, (k4_xboole_0(X1,k2_xboole_0(k4_xboole_0(X2,X1),X3))=k4_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_55, c_0_75])).
fof(c_0_80, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t48_xboole_1])).
cnf(c_0_81, plain, (k2_xboole_0(X1,X2)=X2|k4_xboole_0(k2_xboole_0(X1,X2),X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_82, plain, (k2_xboole_0(X1,k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1)))=X1), inference(spm,[status(thm)],[c_0_78, c_0_33])).
cnf(c_0_83, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X3,X1)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_33]), c_0_79])).
cnf(c_0_84, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_24]), c_0_67])).
fof(c_0_85, negated_conjecture, k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))!=k3_xboole_0(esk1_0,esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_80])])])).
cnf(c_0_86, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_87, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))!=k3_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_88, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_35]), c_0_28]), c_0_45])).
cnf(c_0_89, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_88])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 53
# Proof object formula steps           : 37
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 31
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 8078
# ...of these trivial                  : 2758
# ...subsumed                          : 3786
# ...remaining for further processing  : 1534
# Other redundant clauses eliminated   : 3919
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 10
# Backward-rewritten                   : 236
# Generated clauses                    : 482373
# ...of the previous two non-trivial   : 214725
# Contextual simplify-reflections      : 3
# Paramodulations                      : 478444
# Factorizations                       : 10
# Equation resolutions                 : 3919
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1285
#    Positive orientable unit clauses  : 948
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 330
# Current number of unprocessed clauses: 206215
# ...number of literals in the above   : 402564
# Current number of archived formulas  : 0
# Current number of archived clauses   : 246
# Clause-clause subsumption calls (NU) : 22473
# Rec. Clause-clause subsumption calls : 18304
# Non-unit clause-clause subsumptions  : 2230
# Unit Clause-clause subsumption calls : 360
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8095
# BW rewrite match successes           : 268
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3806245

# -------------------------------------------------
# User time                : 2.377 s
# System time              : 0.056 s
# Total time               : 2.433 s
# Maximum resident set size: 159336 pages
