# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', t4_xboole_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', t12_setfam_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', t2_boole)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', t7_xboole_0)).
fof(t173_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', t173_relat_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', t65_xboole_1)).
fof(t168_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k10_relat_1(X2,X1)=k10_relat_1(X2,k3_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', t168_relat_1)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', t166_relat_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', d4_xboole_0)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', d14_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/b7e1ddcb187d4022b0e7d4cb40406c69.E_input', d5_relat_1)).
fof(c_0_11, plain, ![X6, X7, X9, X10, X11]:((r1_xboole_0(X6,X7)|r2_hidden(esk1_2(X6,X7),k3_xboole_0(X6,X7)))&(~r2_hidden(X11,k3_xboole_0(X9,X10))|~r1_xboole_0(X9,X10))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_12, plain, ![X12, X13]:k1_setfam_1(k2_tarski(X12,X13))=k3_xboole_0(X12,X13), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_13, plain, ![X48]:k3_xboole_0(X48,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_14, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X44]:(X44=k1_xboole_0|r2_hidden(esk9_1(X44),X44)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_17, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), inference(assume_negation,[status(cth)],[t173_relat_1])).
cnf(c_0_18, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, ![X49]:r1_xboole_0(X49,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_20, plain, ![X19, X20]:(~v1_relat_1(X20)|k10_relat_1(X20,X19)=k10_relat_1(X20,k3_xboole_0(k2_relat_1(X20),X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t168_relat_1])])).
cnf(c_0_21, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22, plain, (X1=k1_xboole_0|r2_hidden(esk9_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_23, negated_conjecture, (v1_relat_1(esk11_0)&((k10_relat_1(esk11_0,esk10_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk11_0),esk10_0))&(k10_relat_1(esk11_0,esk10_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk11_0),esk10_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
cnf(c_0_24, plain, (k1_setfam_1(k2_tarski(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_18, c_0_15])).
cnf(c_0_25, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_26, plain, ![X14, X15, X16, X18]:(((r2_hidden(esk2_3(X14,X15,X16),k2_relat_1(X16))|~r2_hidden(X14,k10_relat_1(X16,X15))|~v1_relat_1(X16))&((r2_hidden(k4_tarski(X14,esk2_3(X14,X15,X16)),X16)|~r2_hidden(X14,k10_relat_1(X16,X15))|~v1_relat_1(X16))&(r2_hidden(esk2_3(X14,X15,X16),X15)|~r2_hidden(X14,k10_relat_1(X16,X15))|~v1_relat_1(X16))))&(~r2_hidden(X18,k2_relat_1(X16))|(~r2_hidden(k4_tarski(X14,X18),X16)|~r2_hidden(X18,X15))|r2_hidden(X14,k10_relat_1(X16,X15))|~v1_relat_1(X16))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
cnf(c_0_27, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k3_xboole_0(k2_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (k1_setfam_1(k2_tarski(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_29, negated_conjecture, (k10_relat_1(esk11_0,esk10_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk11_0),esk10_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_24]), c_0_25])])).
cnf(c_0_31, plain, (r2_hidden(esk2_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_32, plain, ![X50, X51, X52, X53, X54, X55, X56, X57]:((((r2_hidden(X53,X50)|~r2_hidden(X53,X52)|X52!=k3_xboole_0(X50,X51))&(r2_hidden(X53,X51)|~r2_hidden(X53,X52)|X52!=k3_xboole_0(X50,X51)))&(~r2_hidden(X54,X50)|~r2_hidden(X54,X51)|r2_hidden(X54,X52)|X52!=k3_xboole_0(X50,X51)))&((~r2_hidden(esk12_3(X55,X56,X57),X57)|(~r2_hidden(esk12_3(X55,X56,X57),X55)|~r2_hidden(esk12_3(X55,X56,X57),X56))|X57=k3_xboole_0(X55,X56))&((r2_hidden(esk12_3(X55,X56,X57),X55)|r2_hidden(esk12_3(X55,X56,X57),X57)|X57=k3_xboole_0(X55,X56))&(r2_hidden(esk12_3(X55,X56,X57),X56)|r2_hidden(esk12_3(X55,X56,X57),X57)|X57=k3_xboole_0(X55,X56))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
fof(c_0_33, plain, ![X21, X22, X23, X24, X26, X27, X28, X29, X31]:((((r2_hidden(k4_tarski(X24,esk3_4(X21,X22,X23,X24)),X21)|~r2_hidden(X24,X23)|X23!=k10_relat_1(X21,X22)|~v1_relat_1(X21))&(r2_hidden(esk3_4(X21,X22,X23,X24),X22)|~r2_hidden(X24,X23)|X23!=k10_relat_1(X21,X22)|~v1_relat_1(X21)))&(~r2_hidden(k4_tarski(X26,X27),X21)|~r2_hidden(X27,X22)|r2_hidden(X26,X23)|X23!=k10_relat_1(X21,X22)|~v1_relat_1(X21)))&((~r2_hidden(esk4_3(X21,X28,X29),X29)|(~r2_hidden(k4_tarski(esk4_3(X21,X28,X29),X31),X21)|~r2_hidden(X31,X28))|X29=k10_relat_1(X21,X28)|~v1_relat_1(X21))&((r2_hidden(k4_tarski(esk4_3(X21,X28,X29),esk5_3(X21,X28,X29)),X21)|r2_hidden(esk4_3(X21,X28,X29),X29)|X29=k10_relat_1(X21,X28)|~v1_relat_1(X21))&(r2_hidden(esk5_3(X21,X28,X29),X28)|r2_hidden(esk4_3(X21,X28,X29),X29)|X29=k10_relat_1(X21,X28)|~v1_relat_1(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
fof(c_0_34, plain, ![X33, X34, X35, X37, X38, X39, X40, X42]:(((~r2_hidden(X35,X34)|r2_hidden(k4_tarski(esk6_3(X33,X34,X35),X35),X33)|X34!=k2_relat_1(X33))&(~r2_hidden(k4_tarski(X38,X37),X33)|r2_hidden(X37,X34)|X34!=k2_relat_1(X33)))&((~r2_hidden(esk7_2(X39,X40),X40)|~r2_hidden(k4_tarski(X42,esk7_2(X39,X40)),X39)|X40=k2_relat_1(X39))&(r2_hidden(esk7_2(X39,X40),X40)|r2_hidden(k4_tarski(esk8_2(X39,X40),esk7_2(X39,X40)),X39)|X40=k2_relat_1(X39)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_35, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k1_setfam_1(k2_tarski(k2_relat_1(X1),X2)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_27, c_0_15])).
cnf(c_0_36, negated_conjecture, (k1_setfam_1(k2_tarski(k2_relat_1(esk11_0),esk10_0))=k1_xboole_0|k10_relat_1(esk11_0,esk10_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_37, negated_conjecture, (v1_relat_1(esk11_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, plain, (~v1_relat_1(X1)|~r2_hidden(X2,k10_relat_1(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (r2_hidden(k4_tarski(esk6_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, negated_conjecture, (k10_relat_1(esk11_0,esk10_0)=k10_relat_1(esk11_0,k1_xboole_0)|k10_relat_1(esk11_0,esk10_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])])).
cnf(c_0_43, plain, (k10_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_22])).
cnf(c_0_44, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X2,X4))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_39, c_0_15])).
cnf(c_0_45, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_46, plain, (r2_hidden(X1,k10_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X4,X3)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (r2_hidden(k4_tarski(esk6_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, (k10_relat_1(esk11_0,esk10_0)=k1_xboole_0|k10_relat_1(esk11_0,k1_xboole_0)!=k1_xboole_0), inference(ef,[status(thm)],[c_0_42])).
cnf(c_0_49, negated_conjecture, (k10_relat_1(esk11_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_43, c_0_37])).
cnf(c_0_50, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_51, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,X3)))), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_52, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk1_2(X1,X2),k1_setfam_1(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_45, c_0_15])).
cnf(c_0_53, plain, (r2_hidden(esk6_3(X1,k2_relat_1(X1),X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))|~r2_hidden(X2,X3)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_54, negated_conjecture, (k10_relat_1(esk11_0,esk10_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_49])])).
cnf(c_0_55, plain, (r2_hidden(X1,X2)|X3!=k1_setfam_1(k2_tarski(X4,X2))|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_50, c_0_15])).
cnf(c_0_56, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_57, negated_conjecture, (~r2_hidden(X1,k2_relat_1(esk11_0))|~r2_hidden(X1,esk10_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_37])]), c_0_30])).
cnf(c_0_58, plain, (r2_hidden(esk2_3(X1,X2,X3),k2_relat_1(X3))|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_59, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_setfam_1(k2_tarski(X3,X2)))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_60, plain, (r2_hidden(esk12_3(X1,X2,X3),X1)|r2_hidden(esk12_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_61, plain, (r1_xboole_0(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_30, c_0_56])).
cnf(c_0_62, negated_conjecture, (~r2_hidden(esk2_3(X1,X2,esk11_0),esk10_0)|~r2_hidden(X1,k10_relat_1(esk11_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_37])])).
cnf(c_0_63, plain, (r2_hidden(esk2_3(X1,k1_setfam_1(k2_tarski(X2,X3)),X4),X3)|~v1_relat_1(X4)|~r2_hidden(X1,k10_relat_1(X4,k1_setfam_1(k2_tarski(X2,X3))))), inference(spm,[status(thm)],[c_0_59, c_0_31])).
cnf(c_0_64, plain, (X3=k1_setfam_1(k2_tarski(X1,X2))|r2_hidden(esk12_3(X1,X2,X3),X3)|r2_hidden(esk12_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_60, c_0_15])).
cnf(c_0_65, plain, (k1_setfam_1(k2_tarski(k1_xboole_0,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_61])).
cnf(c_0_66, negated_conjecture, (~r2_hidden(X1,k10_relat_1(esk11_0,k1_setfam_1(k2_tarski(X2,esk10_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_37])])).
cnf(c_0_67, plain, (X1=k1_xboole_0|r2_hidden(esk12_3(k1_xboole_0,X2,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_64]), c_0_65])).
cnf(c_0_68, negated_conjecture, (k10_relat_1(esk11_0,k1_setfam_1(k2_tarski(X1,esk10_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_69, negated_conjecture, (~r2_hidden(X1,k1_setfam_1(k2_tarski(X2,esk10_0)))|~r2_hidden(X1,k2_relat_1(esk11_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_68]), c_0_37])]), c_0_30])).
cnf(c_0_70, negated_conjecture, (k10_relat_1(esk11_0,esk10_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk11_0),esk10_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_71, negated_conjecture, (r1_xboole_0(X1,esk10_0)|~r2_hidden(esk1_2(X1,esk10_0),k2_relat_1(esk11_0))), inference(spm,[status(thm)],[c_0_69, c_0_52])).
cnf(c_0_72, negated_conjecture, (~r1_xboole_0(k2_relat_1(esk11_0),esk10_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_54])])).
cnf(c_0_73, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_56]), c_0_72]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 74
# Proof object clause steps            : 51
# Proof object formula steps           : 23
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 25340
# ...of these trivial                  : 2
# ...subsumed                          : 23742
# ...remaining for further processing  : 1596
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 7
# Generated clauses                    : 417765
# ...of the previous two non-trivial   : 388060
# Contextual simplify-reflections      : 2
# Paramodulations                      : 417702
# Factorizations                       : 55
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1545
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 1528
# Current number of unprocessed clauses: 362640
# ...number of literals in the above   : 1631632
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 1572235
# Rec. Clause-clause subsumption calls : 995681
# Non-unit clause-clause subsumptions  : 21307
# Unit Clause-clause subsumption calls : 329
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 73
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7079622

# -------------------------------------------------
# User time                : 4.445 s
# System time              : 0.113 s
# Total time               : 4.558 s
# Maximum resident set size: 254476 pages
