# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', commutativity_k3_xboole_0)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t21_xboole_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t39_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t41_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', l32_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', commutativity_k2_xboole_0)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t1_boole)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', d5_xboole_0)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t16_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', idempotence_k3_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t28_xboole_1)).
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t46_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t36_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', d7_xboole_0)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t22_xboole_1)).
fof(t47_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/b838648f798e4745814292367ffe28c0.E_input', t47_xboole_1)).
fof(c_0_18, plain, ![X50, X51]:r1_tarski(k3_xboole_0(X50,X51),X50), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_19, plain, ![X42, X43]:k3_xboole_0(X42,X43)=k3_xboole_0(X43,X42), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_20, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X25, X26]:k3_xboole_0(X25,k2_xboole_0(X25,X26))=X25, inference(variable_rename,[status(thm)],[t21_xboole_1])).
fof(c_0_23, plain, ![X40, X41]:k2_xboole_0(X40,k4_xboole_0(X41,X40))=k2_xboole_0(X40,X41), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_24, plain, ![X36, X37, X38]:k4_xboole_0(k4_xboole_0(X36,X37),X38)=k4_xboole_0(X36,k2_xboole_0(X37,X38)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_25, plain, ![X7, X8]:((k4_xboole_0(X7,X8)!=k1_xboole_0|r1_tarski(X7,X8))&(~r1_tarski(X7,X8)|k4_xboole_0(X7,X8)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_26, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_27, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, ![X9, X10]:k2_xboole_0(X9,X10)=k2_xboole_0(X10,X9), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_29, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X24]:k2_xboole_0(X24,k1_xboole_0)=X24, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_33, plain, ![X27, X28, X29, X30, X31, X32, X33, X34]:((((r2_hidden(X30,X27)|~r2_hidden(X30,X29)|X29!=k4_xboole_0(X27,X28))&(~r2_hidden(X30,X28)|~r2_hidden(X30,X29)|X29!=k4_xboole_0(X27,X28)))&(~r2_hidden(X31,X27)|r2_hidden(X31,X28)|r2_hidden(X31,X29)|X29!=k4_xboole_0(X27,X28)))&((~r2_hidden(esk4_3(X32,X33,X34),X34)|(~r2_hidden(esk4_3(X32,X33,X34),X32)|r2_hidden(esk4_3(X32,X33,X34),X33))|X34=k4_xboole_0(X32,X33))&((r2_hidden(esk4_3(X32,X33,X34),X32)|r2_hidden(esk4_3(X32,X33,X34),X34)|X34=k4_xboole_0(X32,X33))&(~r2_hidden(esk4_3(X32,X33,X34),X33)|r2_hidden(esk4_3(X32,X33,X34),X34)|X34=k4_xboole_0(X32,X33))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_34, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (k2_xboole_0(X1,k4_xboole_0(X2,k2_xboole_0(X3,X1)))=k2_xboole_0(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_20])).
cnf(c_0_38, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_39, plain, ![X21, X22, X23]:k3_xboole_0(k3_xboole_0(X21,X22),X23)=k3_xboole_0(X21,k3_xboole_0(X22,X23)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
fof(c_0_40, plain, ![X39]:k3_xboole_0(X39,X39)=X39, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_41, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_42, plain, ![X15, X16, X18, X19, X20]:(((r2_hidden(esk3_2(X15,X16),X15)|r1_xboole_0(X15,X16))&(r2_hidden(esk3_2(X15,X16),X16)|r1_xboole_0(X15,X16)))&(~r2_hidden(X20,X18)|~r2_hidden(X20,X19)|~r1_xboole_0(X18,X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_43, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44, plain, (k2_xboole_0(X1,k4_xboole_0(k3_xboole_0(k2_xboole_0(X2,X1),X3),X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_45, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_47, plain, ![X46, X47]:(~r1_tarski(X46,X47)|k3_xboole_0(X46,X47)=X46), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_48, plain, ![X44, X45]:k4_xboole_0(X44,k2_xboole_0(X44,X45))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
cnf(c_0_49, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_50, plain, (r2_hidden(esk3_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X1),X2)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_52, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_55, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_56, plain, ![X13, X14]:r1_tarski(k4_xboole_0(X13,X14),X13), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_57, plain, ![X11, X12]:((~r1_xboole_0(X11,X12)|k3_xboole_0(X11,X12)=k1_xboole_0)&(k3_xboole_0(X11,X12)!=k1_xboole_0|r1_xboole_0(X11,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_58, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk3_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_60, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X2),X1)), inference(spm,[status(thm)],[c_0_51, c_0_35])).
cnf(c_0_61, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_52, c_0_21])).
cnf(c_0_62, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_63, plain, (k4_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(k4_xboole_0(X1,X2),X3)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_55, c_0_30])).
cnf(c_0_64, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_65, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_66, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_67, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_38, c_0_35])).
cnf(c_0_68, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X3),X2)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_69, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(k4_xboole_0(X1,X2),X3)))=X1), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_70, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_26, c_0_45])).
fof(c_0_71, plain, ![X48, X49]:k2_xboole_0(X48,k3_xboole_0(X48,X49))=X48, inference(variable_rename,[status(thm)],[t22_xboole_1])).
cnf(c_0_72, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k3_xboole_0(X1,X3)))=k3_xboole_0(X2,k3_xboole_0(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_61]), c_0_45]), c_0_45])).
cnf(c_0_73, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_64]), c_0_21])).
cnf(c_0_74, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_75, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_27, c_0_67])).
cnf(c_0_76, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(k4_xboole_0(X1,X2),X3)),X2)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_77, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_70, c_0_27])).
cnf(c_0_78, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_64, c_0_30])).
cnf(c_0_79, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_80, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X1,X3)))=k3_xboole_0(X2,k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_81, plain, (k3_xboole_0(X1,k3_xboole_0(k4_xboole_0(X2,X1),X3))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_74]), c_0_75])).
cnf(c_0_82, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_76, c_0_38])).
cnf(c_0_83, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_77, c_0_21])).
cnf(c_0_84, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_78, c_0_35])).
cnf(c_0_85, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_79, c_0_61])).
cnf(c_0_86, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X3))=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_87, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_82]), c_0_21])).
cnf(c_0_88, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),k2_xboole_0(k3_xboole_0(X1,X2),X4))), inference(spm,[status(thm)],[c_0_83, c_0_45])).
cnf(c_0_89, plain, (k3_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,X2)|k4_xboole_0(X1,k2_xboole_0(X2,X3))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_30])).
cnf(c_0_90, plain, (k2_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_79, c_0_27])).
cnf(c_0_91, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_92, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_93, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(k3_xboole_0(X2,X1),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_46]), c_0_45]), c_0_61])).
cnf(c_0_94, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k4_xboole_0(X2,X1)|k4_xboole_0(X2,X1)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_21])).
cnf(c_0_95, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_91])).
cnf(c_0_96, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_92]), c_0_73])).
cnf(c_0_97, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_84, c_0_79])).
cnf(c_0_98, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_21]), c_0_74]), c_0_67])).
cnf(c_0_99, plain, (k4_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_73])).
cnf(c_0_100, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_97])).
cnf(c_0_101, plain, (k3_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_53, c_0_98])).
cnf(c_0_102, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X1))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_103, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X3))=k3_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_45])).
cnf(c_0_104, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X3))=k3_xboole_0(X2,X1)|k3_xboole_0(X2,X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_103])).
cnf(c_0_105, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_87, c_0_104])).
cnf(c_0_106, plain, (k3_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,k3_xboole_0(X1,X2))))=k1_xboole_0), inference(spm,[status(thm)],[c_0_45, c_0_74])).
fof(c_0_107, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t47_xboole_1])).
cnf(c_0_108, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_105]), c_0_95])).
cnf(c_0_109, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_106])).
cnf(c_0_110, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_29]), c_0_67])).
fof(c_0_111, negated_conjecture, k4_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(esk1_0,esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_107])])])).
cnf(c_0_112, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_110])).
cnf(c_0_113, negated_conjecture, (k4_xboole_0(esk1_0,k3_xboole_0(esk1_0,esk2_0))!=k4_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_114, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_112]), c_0_30]), c_0_35]), c_0_85])).
cnf(c_0_115, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_114])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 116
# Proof object clause steps            : 79
# Proof object formula steps           : 37
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 20
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 57
# Proof object simplifying inferences  : 24
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 7659
# ...of these trivial                  : 2789
# ...subsumed                          : 3599
# ...remaining for further processing  : 1271
# Other redundant clauses eliminated   : 1054
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 40
# Generated clauses                    : 361424
# ...of the previous two non-trivial   : 136671
# Contextual simplify-reflections      : 2
# Paramodulations                      : 360364
# Factorizations                       : 6
# Equation resolutions                 : 1054
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1219
#    Positive orientable unit clauses  : 1050
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 162
# Current number of unprocessed clauses: 128772
# ...number of literals in the above   : 171534
# Current number of archived formulas  : 0
# Current number of archived clauses   : 49
# Clause-clause subsumption calls (NU) : 9574
# Rec. Clause-clause subsumption calls : 7775
# Non-unit clause-clause subsumptions  : 1705
# Unit Clause-clause subsumption calls : 196
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7416
# BW rewrite match successes           : 132
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2794791

# -------------------------------------------------
# User time                : 1.763 s
# System time              : 0.040 s
# Total time               : 1.803 s
# Maximum resident set size: 119312 pages
