# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t7_ordinal1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t2_subset)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', d19_relat_1)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', existence_m1_subset_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k2_relset_1(X1,X2,X3)=k2_relat_1(X3)), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', redefinition_k2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', cc1_relset_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', l13_xboole_0)).
fof(fc4_zfmisc_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>v1_xboole_0(k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', fc4_zfmisc_1)).
fof(t186_funct_2, conjecture, ![X1, X2, X3]:(~(v1_xboole_0(X3))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>![X5]:((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))=>![X6]:(m1_subset_1(X6,X2)=>(r1_tarski(k2_relset_1(X2,X3,X4),k1_relset_1(X3,X1,X5))=>(X2=k1_xboole_0|k1_funct_1(k8_funct_2(X2,X3,X1,X4,X5),X6)=k7_partfun1(X1,X5,k1_funct_1(X4,X6)))))))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t186_funct_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t3_subset)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', fc1_xboole_0)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', cc2_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', redefinition_k1_relset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', cc1_subset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', d10_xboole_0)).
fof(t8_funct_2, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r1_tarski(k2_relset_1(X1,X2,X4),X3)=>((X2=k1_xboole_0&~(X1=k1_xboole_0))|(v1_funct_1(X4)&(v1_funct_2(X4,X1,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))))))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t8_funct_2)).
fof(fc9_relat_1, axiom, ![X1]:((~(v1_xboole_0(X1))&v1_relat_1(X1))=>~(v1_xboole_0(k2_relat_1(X1)))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', fc9_relat_1)).
fof(cc6_funct_2, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&(~(v1_xboole_0(X3))&v1_funct_2(X3,X1,X2)))))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', cc6_funct_2)).
fof(cc1_funct_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_funct_1(X1)), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', cc1_funct_1)).
fof(t218_relat_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>![X3]:((v1_relat_1(X3)&v5_relat_1(X3,X1))=>v5_relat_1(X3,X2))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t218_relat_1)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', d18_relat_1)).
fof(cc5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X2))=>v4_relat_1(X3,X1))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', cc5_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', fc6_relat_1)).
fof(t185_funct_2, axiom, ![X1, X2, X3]:(~(v1_xboole_0(X3))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>![X5]:((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))=>![X6]:(m1_subset_1(X6,X2)=>(r1_tarski(k2_relset_1(X2,X3,X4),k1_relset_1(X3,X1,X5))=>(X2=k1_xboole_0|k1_funct_1(k8_funct_2(X2,X3,X1,X4,X5),X6)=k1_funct_1(X5,k1_funct_1(X4,X6)))))))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t185_funct_2)).
fof(t7_funct_2, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r2_hidden(X3,X1)=>(X2=k1_xboole_0|r2_hidden(k1_funct_1(X4,X3),X2)))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t7_funct_2)).
fof(d8_partfun1, axiom, ![X1, X2]:((v1_relat_1(X2)&(v5_relat_1(X2,X1)&v1_funct_1(X2)))=>![X3]:(r2_hidden(X3,k1_relat_1(X2))=>k7_partfun1(X1,X2,X3)=k1_funct_1(X2,X3))), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', d8_partfun1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/b88030f949e74459a9a14d7491af160b.E_input', t2_xboole_1)).
fof(c_0_27, plain, ![X7, X8]:(~r2_hidden(X7,X8)|~r1_tarski(X8,X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_28, plain, ![X9, X10]:(~m1_subset_1(X9,X10)|(v1_xboole_0(X10)|r2_hidden(X9,X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_29, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_31, plain, ![X11, X12]:((~v5_relat_1(X12,X11)|r1_tarski(k2_relat_1(X12),X11)|~v1_relat_1(X12))&(~r1_tarski(k2_relat_1(X12),X11)|v5_relat_1(X12,X11)|~v1_relat_1(X12))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
cnf(c_0_32, plain, (v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_34, plain, ![X13]:m1_subset_1(esk1_1(X13),X13), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_35, plain, (v1_xboole_0(k2_relat_1(X1))|~v5_relat_1(X1,X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k2_relat_1(X1))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36, plain, (m1_subset_1(esk1_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_37, plain, ![X15, X16, X17]:(~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(X15,X16)))|k2_relset_1(X15,X16,X17)=k2_relat_1(X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
fof(c_0_38, plain, ![X18, X19, X20]:(~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19)))|v1_relat_1(X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_39, plain, ![X21]:(~v1_xboole_0(X21)|X21=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
fof(c_0_40, plain, ![X22, X23]:(~v1_xboole_0(X22)|v1_xboole_0(k2_zfmisc_1(X22,X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_zfmisc_1])])).
fof(c_0_41, negated_conjecture, ~(![X1, X2, X3]:(~(v1_xboole_0(X3))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X2,X3)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))=>![X5]:((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))=>![X6]:(m1_subset_1(X6,X2)=>(r1_tarski(k2_relset_1(X2,X3,X4),k1_relset_1(X3,X1,X5))=>(X2=k1_xboole_0|k1_funct_1(k8_funct_2(X2,X3,X1,X4,X5),X6)=k7_partfun1(X1,X5,k1_funct_1(X4,X6))))))))), inference(assume_negation,[status(cth)],[t186_funct_2])).
cnf(c_0_42, plain, (v1_xboole_0(k2_relat_1(X1))|~v5_relat_1(X1,esk1_1(k2_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_43, plain, (k2_relset_1(X2,X3,X1)=k2_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_47, plain, ![X30, X31]:((~m1_subset_1(X30,k1_zfmisc_1(X31))|r1_tarski(X30,X31))&(~r1_tarski(X30,X31)|m1_subset_1(X30,k1_zfmisc_1(X31)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_48, negated_conjecture, (~v1_xboole_0(esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,esk3_0,esk4_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0)))))&((v1_funct_1(esk6_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0))))&(m1_subset_1(esk7_0,esk3_0)&(r1_tarski(k2_relset_1(esk3_0,esk4_0,esk5_0),k1_relset_1(esk4_0,esk2_0,esk6_0))&(esk3_0!=k1_xboole_0&k1_funct_1(k8_funct_2(esk3_0,esk4_0,esk2_0,esk5_0,esk6_0),esk7_0)!=k7_partfun1(esk2_0,esk6_0,k1_funct_1(esk5_0,esk7_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_41])])])])).
cnf(c_0_49, plain, (v1_xboole_0(k2_relset_1(X1,X2,X3))|~v5_relat_1(X3,esk1_1(k2_relset_1(X1,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_50, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_52, plain, ![X32, X33, X34]:((v4_relat_1(X34,X32)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))&(v5_relat_1(X34,X33)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_53, plain, ![X37, X38, X39]:(~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X37,X38)))|k1_relset_1(X37,X38,X39)=k1_relat_1(X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_54, plain, ![X35, X36]:(~v1_xboole_0(X35)|(~m1_subset_1(X36,k1_zfmisc_1(X35))|v1_xboole_0(X36))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_55, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, negated_conjecture, (r1_tarski(k2_relset_1(esk3_0,esk4_0,esk5_0),k1_relset_1(esk4_0,esk2_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (v1_xboole_0(k2_relat_1(X1))|~v5_relat_1(X1,esk1_1(k2_relat_1(X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(spm,[status(thm)],[c_0_49, c_0_43])).
cnf(c_0_58, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_59, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_60, plain, ![X54, X55]:(((r1_tarski(X54,X55)|X54!=X55)&(r1_tarski(X55,X54)|X54!=X55))&(~r1_tarski(X54,X55)|~r1_tarski(X55,X54)|X54=X55)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_61, plain, ![X41, X42, X43, X44]:(((v1_funct_1(X44)|X42=k1_xboole_0|~r1_tarski(k2_relset_1(X41,X42,X44),X43)|(~v1_funct_1(X44)|(~v1_funct_2(X44,X41,X42)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))))&((v1_funct_2(X44,X41,X43)|X42=k1_xboole_0|~r1_tarski(k2_relset_1(X41,X42,X44),X43)|(~v1_funct_1(X44)|(~v1_funct_2(X44,X41,X42)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))))&(m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X43)))|X42=k1_xboole_0|~r1_tarski(k2_relset_1(X41,X42,X44),X43)|(~v1_funct_1(X44)|(~v1_funct_2(X44,X41,X42)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))))))&((v1_funct_1(X44)|X41!=k1_xboole_0|~r1_tarski(k2_relset_1(X41,X42,X44),X43)|(~v1_funct_1(X44)|(~v1_funct_2(X44,X41,X42)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))))&((v1_funct_2(X44,X41,X43)|X41!=k1_xboole_0|~r1_tarski(k2_relset_1(X41,X42,X44),X43)|(~v1_funct_1(X44)|(~v1_funct_2(X44,X41,X42)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))))&(m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X43)))|X41!=k1_xboole_0|~r1_tarski(k2_relset_1(X41,X42,X44),X43)|(~v1_funct_1(X44)|(~v1_funct_2(X44,X41,X42)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_funct_2])])])])).
cnf(c_0_62, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_64, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(k2_relset_1(esk3_0,esk4_0,esk5_0),k1_zfmisc_1(k1_relset_1(esk4_0,esk2_0,esk6_0)))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_66, plain, ![X40]:(v1_xboole_0(X40)|~v1_relat_1(X40)|~v1_xboole_0(k2_relat_1(X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc9_relat_1])])])).
cnf(c_0_67, plain, (v1_xboole_0(k2_relat_1(X1))|~v5_relat_1(X1,esk1_1(k2_relat_1(X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_68, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_59, c_0_58])).
fof(c_0_69, plain, ![X45, X46, X47]:((v1_funct_1(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,X45,X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X45,X46)))|(v1_xboole_0(X45)|v1_xboole_0(X46)))&((~v1_xboole_0(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,X45,X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X45,X46)))|(v1_xboole_0(X45)|v1_xboole_0(X46)))&(v1_funct_2(X47,X45,X46)|(~v1_funct_1(X47)|~v1_funct_2(X47,X45,X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X45,X46)))|(v1_xboole_0(X45)|v1_xboole_0(X46))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc6_funct_2])])])])])).
fof(c_0_70, plain, ![X48]:(~v1_xboole_0(X48)|v1_funct_1(X48)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_1])])).
fof(c_0_71, plain, ![X49, X50, X51]:(~r1_tarski(X49,X50)|(~v1_relat_1(X51)|~v5_relat_1(X51,X49)|v5_relat_1(X51,X50))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t218_relat_1])])])).
fof(c_0_72, plain, ![X52, X53]:((~v4_relat_1(X53,X52)|r1_tarski(k1_relat_1(X53),X52)|~v1_relat_1(X53))&(~r1_tarski(k1_relat_1(X53),X52)|v4_relat_1(X53,X52)|~v1_relat_1(X53))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
cnf(c_0_73, plain, (v5_relat_1(X1,X2)|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_74, plain, ![X66, X67, X68]:(~v1_relat_1(X67)|~v4_relat_1(X67,X66)|(~m1_subset_1(X68,k1_zfmisc_1(X67))|v4_relat_1(X68,X66))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_relat_1])])])).
fof(c_0_75, plain, ![X69, X70]:v1_relat_1(k2_zfmisc_1(X69,X70)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_76, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_77, plain, ![X56, X57, X58, X59, X60, X61]:(v1_xboole_0(X58)|(~v1_funct_1(X59)|(~v1_funct_2(X59,X57,X58)|~m1_subset_1(X59,k1_zfmisc_1(k2_zfmisc_1(X57,X58))))|(~v1_funct_1(X60)|~m1_subset_1(X60,k1_zfmisc_1(k2_zfmisc_1(X58,X56)))|(~m1_subset_1(X61,X57)|(~r1_tarski(k2_relset_1(X57,X58,X59),k1_relset_1(X58,X56,X60))|(X57=k1_xboole_0|k1_funct_1(k8_funct_2(X57,X58,X56,X59,X60),X61)=k1_funct_1(X60,k1_funct_1(X59,X61)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t185_funct_2])])])])).
fof(c_0_78, plain, ![X62, X63, X64, X65]:(~v1_funct_1(X65)|(~v1_funct_2(X65,X62,X63)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X62,X63))))|(~r2_hidden(X64,X62)|(X63=k1_xboole_0|r2_hidden(k1_funct_1(X65,X64),X63)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_funct_2])])).
cnf(c_0_79, plain, (v1_funct_2(X1,X2,X3)|X4=k1_xboole_0|~r1_tarski(k2_relset_1(X2,X4,X1),X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_80, negated_conjecture, (r1_tarski(k2_relset_1(esk3_0,esk4_0,esk5_0),k1_relat_1(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_62]), c_0_63])])).
cnf(c_0_81, negated_conjecture, (v1_funct_2(esk5_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_82, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_83, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_84, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|X4=k1_xboole_0|~r1_tarski(k2_relset_1(X2,X4,X1),X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_85, negated_conjecture, (v1_xboole_0(k2_relset_1(esk3_0,esk4_0,esk5_0))|~v1_xboole_0(k1_relset_1(esk4_0,esk2_0,esk6_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_86, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(k2_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_87, plain, (v1_xboole_0(k2_relat_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_88, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_44, c_0_58])).
cnf(c_0_89, plain, (v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_90, plain, (v1_funct_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_91, plain, (v5_relat_1(X3,X2)|~r1_tarski(X1,X2)|~v1_relat_1(X3)|~v5_relat_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_92, plain, (r1_tarski(k1_relat_1(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_93, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~r1_tarski(k2_relset_1(X3,X4,X1),X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_43]), c_0_44])).
cnf(c_0_94, plain, (v4_relat_1(X3,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_95, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_96, plain, (v4_relat_1(X1,X2)|~r1_tarski(k1_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_97, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_98, plain, (v1_xboole_0(X1)|X3=k1_xboole_0|k1_funct_1(k8_funct_2(X3,X1,X5,X2,X4),X6)=k1_funct_1(X4,k1_funct_1(X2,X6))|~v1_funct_1(X2)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X5)))|~m1_subset_1(X6,X3)|~r1_tarski(k2_relset_1(X3,X1,X2),k1_relset_1(X1,X5,X4))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_99, negated_conjecture, (v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_100, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_101, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_102, plain, ![X71, X72, X73]:(~v1_relat_1(X72)|(~v5_relat_1(X72,X71)|~v1_funct_1(X72))|(~r2_hidden(X73,k1_relat_1(X72))|k7_partfun1(X71,X72,X73)=k1_funct_1(X72,X73))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_partfun1])])])).
cnf(c_0_103, plain, (X3=k1_xboole_0|r2_hidden(k1_funct_1(X1,X4),X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r2_hidden(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_104, negated_conjecture, (esk4_0=k1_xboole_0|v1_funct_2(esk5_0,esk3_0,k1_relat_1(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81]), c_0_82]), c_0_83])])).
cnf(c_0_105, negated_conjecture, (esk4_0=k1_xboole_0|m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,k1_relat_1(esk6_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_80]), c_0_81]), c_0_82]), c_0_83])])).
cnf(c_0_106, negated_conjecture, (v1_xboole_0(k2_relset_1(esk3_0,esk4_0,esk5_0))|~v1_xboole_0(k1_relat_1(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_62]), c_0_63])])).
cnf(c_0_107, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])).
cnf(c_0_108, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_funct_2(X3,X1,X2)|~v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[c_0_89, c_0_90])).
fof(c_0_109, plain, ![X74]:r1_tarski(k1_xboole_0,X74), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_110, plain, (v5_relat_1(X1,X2)|~v4_relat_1(X3,X2)|~v5_relat_1(X1,k1_relat_1(X3))|~v1_relat_1(X1)|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_111, negated_conjecture, (v5_relat_1(esk5_0,k1_relat_1(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_80]), c_0_83])])).
cnf(c_0_112, negated_conjecture, (v1_relat_1(esk5_0)), inference(spm,[status(thm)],[c_0_44, c_0_83])).
cnf(c_0_113, negated_conjecture, (v1_relat_1(esk6_0)), inference(spm,[status(thm)],[c_0_44, c_0_63])).
cnf(c_0_114, negated_conjecture, (v4_relat_1(esk6_0,X1)|~v4_relat_1(k2_zfmisc_1(esk4_0,esk2_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_63]), c_0_95])])).
cnf(c_0_115, plain, (v4_relat_1(X1,k1_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_116, negated_conjecture, (k1_funct_1(k8_funct_2(esk3_0,esk4_0,esk2_0,esk5_0,esk6_0),esk7_0)!=k7_partfun1(esk2_0,esk6_0,k1_funct_1(esk5_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_117, negated_conjecture, (k1_funct_1(k8_funct_2(esk3_0,esk4_0,esk2_0,esk5_0,esk6_0),X1)=k1_funct_1(esk6_0,k1_funct_1(esk5_0,X1))|~m1_subset_1(X1,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_56]), c_0_81]), c_0_99]), c_0_82]), c_0_63]), c_0_83])]), c_0_100]), c_0_101])).
cnf(c_0_118, negated_conjecture, (m1_subset_1(esk7_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_119, plain, (k7_partfun1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)|~v1_funct_1(X1)|~r2_hidden(X3,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_120, negated_conjecture, (k1_relat_1(esk6_0)=k1_xboole_0|esk4_0=k1_xboole_0|r2_hidden(k1_funct_1(esk5_0,X1),k1_relat_1(esk6_0))|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_82])]), c_0_105])).
cnf(c_0_121, negated_conjecture, (v1_xboole_0(k2_relset_1(esk3_0,esk4_0,esk5_0))|~m1_subset_1(k1_relat_1(esk6_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_122, plain, (m1_subset_1(k1_relat_1(X1),k1_zfmisc_1(X2))|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_55, c_0_92])).
cnf(c_0_123, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(esk5_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_81]), c_0_83])]), c_0_101])).
cnf(c_0_124, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_125, plain, (m1_subset_1(k2_relat_1(X1),k1_zfmisc_1(X2))|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_55, c_0_33])).
cnf(c_0_126, negated_conjecture, (v5_relat_1(esk5_0,X1)|~v4_relat_1(esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112]), c_0_113])])).
cnf(c_0_127, negated_conjecture, (v4_relat_1(esk6_0,k1_relat_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_95])])).
cnf(c_0_128, negated_conjecture, (k7_partfun1(esk2_0,esk6_0,k1_funct_1(esk5_0,esk7_0))!=k1_funct_1(esk6_0,k1_funct_1(esk5_0,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118])])).
cnf(c_0_129, negated_conjecture, (k7_partfun1(X1,esk6_0,k1_funct_1(esk5_0,X2))=k1_funct_1(esk6_0,k1_funct_1(esk5_0,X2))|k1_relat_1(esk6_0)=k1_xboole_0|esk4_0=k1_xboole_0|~v5_relat_1(esk6_0,X1)|~r2_hidden(X2,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_99]), c_0_113])])).
cnf(c_0_130, negated_conjecture, (v5_relat_1(esk6_0,esk2_0)), inference(spm,[status(thm)],[c_0_59, c_0_63])).
cnf(c_0_131, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_relset_1(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_43]), c_0_44])).
cnf(c_0_132, negated_conjecture, (v1_xboole_0(k2_relset_1(esk3_0,esk4_0,esk5_0))|~v4_relat_1(esk6_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_113])])).
cnf(c_0_133, negated_conjecture, (~v1_xboole_0(esk5_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_123]), c_0_100])).
cnf(c_0_134, negated_conjecture, (esk4_0=k1_xboole_0|m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk3_0,k1_relset_1(esk4_0,esk2_0,esk6_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_56]), c_0_81]), c_0_82]), c_0_83])])).
cnf(c_0_135, plain, (v1_xboole_0(k1_relat_1(X1))|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_32, c_0_92])).
cnf(c_0_136, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_137, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_55, c_0_124])).
cnf(c_0_138, plain, (v1_xboole_0(k2_relat_1(X1))|~v5_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_64, c_0_125])).
cnf(c_0_139, negated_conjecture, (v5_relat_1(esk5_0,k1_relat_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_140, negated_conjecture, (k1_relat_1(esk6_0)=k1_xboole_0|esk4_0=k1_xboole_0|~r2_hidden(esk7_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_130])])).
cnf(c_0_141, negated_conjecture, (~v4_relat_1(esk6_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_83])]), c_0_133])).
cnf(c_0_142, negated_conjecture, (esk4_0=k1_xboole_0|~v1_xboole_0(k2_zfmisc_1(esk3_0,k1_relset_1(esk4_0,esk2_0,esk6_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_134]), c_0_133])).
cnf(c_0_143, plain, (v1_xboole_0(k1_relat_1(X1))|~v4_relat_1(X1,esk1_1(k1_relat_1(X1)))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_135, c_0_36])).
cnf(c_0_144, plain, (v4_relat_1(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_145, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_95, c_0_58])).
cnf(c_0_146, negated_conjecture, (v1_xboole_0(k2_relat_1(esk5_0))|~v1_xboole_0(k1_relat_1(k2_zfmisc_1(esk4_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_112])])).
cnf(c_0_147, negated_conjecture, (esk4_0=k1_xboole_0|~r2_hidden(esk7_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_140]), c_0_113])]), c_0_141])).
cnf(c_0_148, negated_conjecture, (esk4_0=k1_xboole_0|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_142, c_0_46])).
cnf(c_0_149, plain, (v1_xboole_0(k1_relat_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_145])])).
cnf(c_0_150, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_relat_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(spm,[status(thm)],[c_0_131, c_0_43])).
cnf(c_0_151, negated_conjecture, (v1_xboole_0(k2_relat_1(esk5_0))|~m1_subset_1(k1_relat_1(k2_zfmisc_1(esk4_0,esk2_0)),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_146, c_0_107])).
cnf(c_0_152, negated_conjecture, (esk4_0=k1_xboole_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_30]), c_0_118])]), c_0_148])).
cnf(c_0_153, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_45, c_0_149])).
cnf(c_0_154, negated_conjecture, (~v1_xboole_0(k2_relat_1(esk5_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_83]), c_0_133])).
cnf(c_0_155, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_151, c_0_152]), c_0_58]), c_0_153]), c_0_137])]), c_0_154]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 156
# Proof object clause steps            : 102
# Proof object formula steps           : 54
# Proof object conjectures             : 46
# Proof object clause conjectures      : 43
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 40
# Proof object initial formulas used   : 27
# Proof object generating inferences   : 59
# Proof object simplifying inferences  : 76
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 27
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 49
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 45
# Processed clauses                    : 2694
# ...of these trivial                  : 5
# ...subsumed                          : 2001
# ...remaining for further processing  : 688
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 40
# Backward-rewritten                   : 211
# Generated clauses                    : 8700
# ...of the previous two non-trivial   : 5421
# Contextual simplify-reflections      : 52
# Paramodulations                      : 8696
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 389
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 342
# Current number of unprocessed clauses: 2646
# ...number of literals in the above   : 13016
# Current number of archived formulas  : 0
# Current number of archived clauses   : 295
# Clause-clause subsumption calls (NU) : 51803
# Rec. Clause-clause subsumption calls : 26067
# Non-unit clause-clause subsumptions  : 1259
# Unit Clause-clause subsumption calls : 948
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 28
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 122707

# -------------------------------------------------
# User time                : 0.092 s
# System time              : 0.010 s
# Total time               : 0.102 s
# Maximum resident set size: 28804 pages
