# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g1_orders_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3, X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/b8863ebcdd6e455ba387b18c23b60f3c.E_input', free_g1_orders_2)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/b8863ebcdd6e455ba387b18c23b60f3c.E_input', dt_u1_orders_2)).
fof(t26_yellow_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(r1_yellow_0(X1,X3)=>k1_yellow_0(X1,X3)=k1_yellow_0(X2,X3))))), file('data/MPTP/proofs/b8863ebcdd6e455ba387b18c23b60f3c.E_input', t26_yellow_0)).
fof(t1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X3=X5&X4=X6)=>((r1_orders_2(X1,X3,X4)=>r1_orders_2(X2,X5,X6))&(r2_orders_2(X1,X3,X4)=>r2_orders_2(X2,X5,X6))))))))))), file('data/MPTP/proofs/b8863ebcdd6e455ba387b18c23b60f3c.E_input', t1_yellow_0)).
fof(d9_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_yellow_0(X1,X2)=>(X3=k1_yellow_0(X1,X2)<=>(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), file('data/MPTP/proofs/b8863ebcdd6e455ba387b18c23b60f3c.E_input', d9_yellow_0)).
fof(dt_k1_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/b8863ebcdd6e455ba387b18c23b60f3c.E_input', dt_k1_yellow_0)).
fof(t2_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X4=X5=>((r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))&(r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5))))))))), file('data/MPTP/proofs/b8863ebcdd6e455ba387b18c23b60f3c.E_input', t2_yellow_0)).
fof(t14_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:((r1_yellow_0(X1,X3)=>r1_yellow_0(X2,X3))&(r2_yellow_0(X1,X3)=>r2_yellow_0(X2,X3)))))), file('data/MPTP/proofs/b8863ebcdd6e455ba387b18c23b60f3c.E_input', t14_yellow_0)).
fof(c_0_8, plain, ![X7, X8, X9, X10]:((X7=X9|g1_orders_2(X7,X8)!=g1_orders_2(X9,X10)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X7,X7))))&(X8=X10|g1_orders_2(X7,X8)!=g1_orders_2(X9,X10)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X7,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])).
fof(c_0_9, plain, ![X11]:(~l1_orders_2(X11)|m1_subset_1(u1_orders_2(X11),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X11))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
fof(c_0_10, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(r1_yellow_0(X1,X3)=>k1_yellow_0(X1,X3)=k1_yellow_0(X2,X3)))))), inference(assume_negation,[status(cth)],[t26_yellow_0])).
cnf(c_0_11, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, negated_conjecture, (l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&(r1_yellow_0(esk1_0,esk3_0)&k1_yellow_0(esk1_0,esk3_0)!=k1_yellow_0(esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_14, plain, (u1_orders_2(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, negated_conjecture, (u1_orders_2(esk2_0)=X1|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).
fof(c_0_18, plain, ![X15, X16, X17, X18, X19, X20]:((~r1_orders_2(X15,X17,X18)|r1_orders_2(X16,X19,X20)|(X17!=X19|X18!=X20)|~m1_subset_1(X20,u1_struct_0(X16))|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X15))|g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))!=g1_orders_2(u1_struct_0(X16),u1_orders_2(X16))|~l1_orders_2(X16)|~l1_orders_2(X15))&(~r2_orders_2(X15,X17,X18)|r2_orders_2(X16,X19,X20)|(X17!=X19|X18!=X20)|~m1_subset_1(X20,u1_struct_0(X16))|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X15))|~m1_subset_1(X17,u1_struct_0(X15))|g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))!=g1_orders_2(u1_struct_0(X16),u1_orders_2(X16))|~l1_orders_2(X16)|~l1_orders_2(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_yellow_0])])])])).
fof(c_0_19, plain, ![X21, X22, X23, X24]:(((r2_lattice3(X21,X22,X23)|X23!=k1_yellow_0(X21,X22)|~r1_yellow_0(X21,X22)|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21))&(~m1_subset_1(X24,u1_struct_0(X21))|(~r2_lattice3(X21,X22,X24)|r1_orders_2(X21,X23,X24))|X23!=k1_yellow_0(X21,X22)|~r1_yellow_0(X21,X22)|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21)))&((m1_subset_1(esk4_3(X21,X22,X23),u1_struct_0(X21))|~r2_lattice3(X21,X22,X23)|X23=k1_yellow_0(X21,X22)|~r1_yellow_0(X21,X22)|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21))&((r2_lattice3(X21,X22,esk4_3(X21,X22,X23))|~r2_lattice3(X21,X22,X23)|X23=k1_yellow_0(X21,X22)|~r1_yellow_0(X21,X22)|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21))&(~r1_orders_2(X21,X23,esk4_3(X21,X22,X23))|~r2_lattice3(X21,X22,X23)|X23=k1_yellow_0(X21,X22)|~r1_yellow_0(X21,X22)|~m1_subset_1(X23,u1_struct_0(X21))|~l1_orders_2(X21))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_yellow_0])])])])])).
fof(c_0_20, plain, ![X26, X27]:(~l1_orders_2(X26)|m1_subset_1(k1_yellow_0(X26,X27),u1_struct_0(X26))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_yellow_0])])).
cnf(c_0_21, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_22, negated_conjecture, (u1_orders_2(esk2_0)=u1_orders_2(esk1_0)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (r1_orders_2(X4,X5,X6)|~r1_orders_2(X1,X2,X3)|X2!=X5|X3!=X6|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (r1_orders_2(X2,X4,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|X4!=k1_yellow_0(X2,X3)|~r1_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (u1_struct_0(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X2,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_21, c_0_12])).
cnf(c_0_27, negated_conjecture, (g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))), inference(rw,[status(thm)],[c_0_15, c_0_22])).
fof(c_0_28, plain, ![X28, X29, X30, X31, X32]:((~r2_lattice3(X28,X30,X31)|r2_lattice3(X29,X30,X32)|X31!=X32|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X28))|g1_orders_2(u1_struct_0(X28),u1_orders_2(X28))!=g1_orders_2(u1_struct_0(X29),u1_orders_2(X29))|~l1_orders_2(X29)|~l1_orders_2(X28))&(~r1_lattice3(X28,X30,X31)|r1_lattice3(X29,X30,X32)|X31!=X32|~m1_subset_1(X32,u1_struct_0(X29))|~m1_subset_1(X31,u1_struct_0(X28))|g1_orders_2(u1_struct_0(X28),u1_orders_2(X28))!=g1_orders_2(u1_struct_0(X29),u1_orders_2(X29))|~l1_orders_2(X29)|~l1_orders_2(X28))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_yellow_0])])])])).
cnf(c_0_29, plain, (r1_orders_2(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_orders_2(X4,X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_23])])).
cnf(c_0_30, plain, (r1_orders_2(X1,k1_yellow_0(X1,X2),X3)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_24]), c_0_25])).
cnf(c_0_31, negated_conjecture, (u1_struct_0(esk2_0)=X1|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_22]), c_0_27]), c_0_16])])).
cnf(c_0_32, plain, (r2_lattice3(X4,X2,X5)|~r2_lattice3(X1,X2,X3)|X3!=X5|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (r2_lattice3(X1,X2,X3)|X3!=k1_yellow_0(X1,X2)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_34, plain, (r1_orders_2(X1,k1_yellow_0(X2,X3),X4)|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X2,X3,X4)|~r1_yellow_0(X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X2)|~m1_subset_1(k1_yellow_0(X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_25])).
cnf(c_0_35, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (r2_lattice3(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X4,X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (r2_lattice3(X1,X2,esk4_3(X1,X2,X3))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_38, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X1))|X3=k1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_39, plain, (r2_lattice3(X1,X2,k1_yellow_0(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_33]), c_0_25])).
cnf(c_0_40, negated_conjecture, (r1_orders_2(esk2_0,k1_yellow_0(X1,X2),X3)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~r2_lattice3(X1,X2,X3)|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)|~m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_22]), c_0_35]), c_0_16]), c_0_35]), c_0_35])])).
cnf(c_0_41, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_42, plain, (X1=k1_yellow_0(X2,X3)|r2_lattice3(X4,X3,esk4_3(X2,X3,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X4)|~l1_orders_2(X2)|~m1_subset_1(esk4_3(X2,X3,X1),u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_43, plain, (r2_lattice3(X1,X2,k1_yellow_0(X3,X2))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_yellow_0(X3,X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(k1_yellow_0(X3,X2),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_39]), c_0_25])).
cnf(c_0_44, plain, (X2=k1_yellow_0(X1,X3)|~r1_orders_2(X1,X2,esk4_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~r1_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_45, negated_conjecture, (r1_orders_2(esk2_0,k1_yellow_0(esk1_0,X1),X2)|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_40]), c_0_41])])).
cnf(c_0_46, negated_conjecture, (X1=k1_yellow_0(esk2_0,X2)|m1_subset_1(esk4_3(esk2_0,X2,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk2_0,X2,X1)|~r1_yellow_0(esk2_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_35]), c_0_16])])).
cnf(c_0_47, negated_conjecture, (X1=k1_yellow_0(esk2_0,X2)|r2_lattice3(X3,X2,esk4_3(esk2_0,X2,X1))|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~r2_lattice3(esk2_0,X2,X1)|~r1_yellow_0(esk2_0,X2)|~l1_orders_2(X3)|~m1_subset_1(esk4_3(esk2_0,X2,X1),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_22]), c_0_35]), c_0_16]), c_0_35])])).
cnf(c_0_48, negated_conjecture, (r2_lattice3(esk2_0,X1,k1_yellow_0(X2,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))|~r1_yellow_0(X2,X1)|~l1_orders_2(X2)|~m1_subset_1(k1_yellow_0(X2,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_22]), c_0_35]), c_0_16]), c_0_35])])).
fof(c_0_49, plain, ![X33, X34, X35]:((~r1_yellow_0(X33,X35)|r1_yellow_0(X34,X35)|g1_orders_2(u1_struct_0(X33),u1_orders_2(X33))!=g1_orders_2(u1_struct_0(X34),u1_orders_2(X34))|~l1_orders_2(X34)|~l1_orders_2(X33))&(~r2_yellow_0(X33,X35)|r2_yellow_0(X34,X35)|g1_orders_2(u1_struct_0(X33),u1_orders_2(X33))!=g1_orders_2(u1_struct_0(X34),u1_orders_2(X34))|~l1_orders_2(X34)|~l1_orders_2(X33))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t14_yellow_0])])])])).
cnf(c_0_50, negated_conjecture, (k1_yellow_0(esk1_0,X1)=k1_yellow_0(esk2_0,X2)|~r2_lattice3(esk1_0,X1,esk4_3(esk2_0,X2,k1_yellow_0(esk1_0,X1)))|~r2_lattice3(esk2_0,X2,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk2_0,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_16]), c_0_35])]), c_0_46])).
cnf(c_0_51, negated_conjecture, (X1=k1_yellow_0(esk2_0,X2)|r2_lattice3(esk1_0,X2,esk4_3(esk2_0,X2,X1))|~r2_lattice3(esk2_0,X2,X1)|~r1_yellow_0(esk2_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_47]), c_0_41])]), c_0_46])).
cnf(c_0_52, negated_conjecture, (r2_lattice3(esk2_0,X1,k1_yellow_0(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_48]), c_0_41])])).
cnf(c_0_53, plain, (r1_yellow_0(X3,X2)|~r1_yellow_0(X1,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~l1_orders_2(X3)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_55, negated_conjecture, (k1_yellow_0(esk2_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(k1_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_56, negated_conjecture, (r1_yellow_0(X1,esk3_0)|g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_41])])).
cnf(c_0_57, negated_conjecture, (k1_yellow_0(esk1_0,esk3_0)!=k1_yellow_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_58, negated_conjecture, (k1_yellow_0(esk2_0,X1)=k1_yellow_0(esk1_0,X1)|~r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_25]), c_0_41])])).
cnf(c_0_59, negated_conjecture, (r1_yellow_0(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_22]), c_0_27]), c_0_16])])).
cnf(c_0_60, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_54])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 61
# Proof object clause steps            : 44
# Proof object formula steps           : 17
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 53
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 104
# ...of these trivial                  : 1
# ...subsumed                          : 21
# ...remaining for further processing  : 82
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 5
# Generated clauses                    : 91
# ...of the previous two non-trivial   : 83
# Contextual simplify-reflections      : 10
# Paramodulations                      : 70
# Factorizations                       : 0
# Equation resolutions                 : 23
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 50
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 41
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 83
# Current number of archived formulas  : 0
# Current number of archived clauses   : 26
# Clause-clause subsumption calls (NU) : 1129
# Rec. Clause-clause subsumption calls : 55
# Non-unit clause-clause subsumptions  : 32
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5160

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28736 pages
