# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_k3_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&l1_waybel_0(k3_yellow19(X1,X2,X3),X1)))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', dt_k3_yellow19)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', d2_yellow_1)).
fof(t17_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3))))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', t17_yellow19)).
fof(fc4_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v3_orders_2(k3_yellow19(X1,X2,X3))&(v4_orders_2(k3_yellow19(X1,X2,X3))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', fc4_yellow19)).
fof(dt_k2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', dt_k2_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', dt_l1_pre_topc)).
fof(t15_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', t15_yellow19)).
fof(fc5_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&v7_waybel_0(k3_yellow19(X1,X2,X3))))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', fc5_yellow19)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', d5_xboole_0)).
fof(t12_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,X2,X3)<=>r1_waybel_7(X1,k2_yellow19(X1,X2),X3))))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', t12_yellow19)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', d1_xboole_0)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', redefinition_k7_subset_1)).
fof(t22_pre_topc, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2)), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', t22_pre_topc)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', d4_subset_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/ba5793e5c599487cbd066b29c8d96570.E_input', fc2_struct_0)).
fof(c_0_16, plain, ![X20, X21, X22]:((~v2_struct_0(k3_yellow19(X20,X21,X22))|(v2_struct_0(X20)|~l1_struct_0(X20)|(v1_xboole_0(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v1_xboole_0(X22)|(~v2_waybel_0(X22,k3_yellow_1(X21))|(~v13_waybel_0(X22,k3_yellow_1(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X21))))))))))&((v6_waybel_0(k3_yellow19(X20,X21,X22),X20)|(v2_struct_0(X20)|~l1_struct_0(X20)|(v1_xboole_0(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v1_xboole_0(X22)|(~v2_waybel_0(X22,k3_yellow_1(X21))|(~v13_waybel_0(X22,k3_yellow_1(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X21))))))))))&(l1_waybel_0(k3_yellow19(X20,X21,X22),X20)|(v2_struct_0(X20)|~l1_struct_0(X20)|(v1_xboole_0(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(v1_xboole_0(X22)|(~v2_waybel_0(X22,k3_yellow_1(X21))|(~v13_waybel_0(X22,k3_yellow_1(X21))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X21)))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_yellow19])])])])).
fof(c_0_17, plain, ![X9]:k3_yellow_1(X9)=k3_lattice3(k1_lattice3(X9)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_18, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t17_yellow19])).
fof(c_0_19, plain, ![X32, X33, X34]:((~v2_struct_0(k3_yellow19(X32,X33,X34))|(v2_struct_0(X32)|~l1_struct_0(X32)|(v1_xboole_0(X33)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|(v1_xboole_0(X34)|(~v2_waybel_0(X34,k3_yellow_1(X33))|(~v13_waybel_0(X34,k3_yellow_1(X33))|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X33))))))))))&((v3_orders_2(k3_yellow19(X32,X33,X34))|(v2_struct_0(X32)|~l1_struct_0(X32)|(v1_xboole_0(X33)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|(v1_xboole_0(X34)|(~v2_waybel_0(X34,k3_yellow_1(X33))|(~v13_waybel_0(X34,k3_yellow_1(X33))|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X33))))))))))&((v4_orders_2(k3_yellow19(X32,X33,X34))|(v2_struct_0(X32)|~l1_struct_0(X32)|(v1_xboole_0(X33)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|(v1_xboole_0(X34)|(~v2_waybel_0(X34,k3_yellow_1(X33))|(~v13_waybel_0(X34,k3_yellow_1(X33))|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X33))))))))))&(v6_waybel_0(k3_yellow19(X32,X33,X34),X32)|(v2_struct_0(X32)|~l1_struct_0(X32)|(v1_xboole_0(X33)|~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|(v1_xboole_0(X34)|(~v2_waybel_0(X34,k3_yellow_1(X33))|(~v13_waybel_0(X34,k3_yellow_1(X33))|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X33))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_yellow19])])])])).
cnf(c_0_20, plain, (l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, plain, ![X43]:(~l1_struct_0(X43)|m1_subset_1(k2_struct_0(X43),k1_zfmisc_1(u1_struct_0(X43)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_struct_0])])).
fof(c_0_23, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v1_xboole_0(esk2_0)&(v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))&(v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&(v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))))))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0))&(r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|r1_waybel_7(esk1_0,esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_18])])])])).
fof(c_0_24, plain, ![X38]:(~l1_pre_topc(X38)|l1_struct_0(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_25, plain, ![X23, X24]:(v2_struct_0(X23)|~l1_struct_0(X23)|(v1_xboole_0(X24)|(~v1_subset_1(X24,u1_struct_0(k3_yellow_1(k2_struct_0(X23))))|(~v2_waybel_0(X24,k3_yellow_1(k2_struct_0(X23)))|(~v13_waybel_0(X24,k3_yellow_1(k2_struct_0(X23)))|~m1_subset_1(X24,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X23))))))))|X24=k2_yellow19(X23,k3_yellow19(X23,k2_struct_0(X23),X24)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_yellow19])])])])).
cnf(c_0_26, plain, (v4_orders_2(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_28, plain, ![X25, X26, X27]:((~v2_struct_0(k3_yellow19(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v1_xboole_0(X27)|(~v1_subset_1(X27,u1_struct_0(k3_yellow_1(X26)))|(~v2_waybel_0(X27,k3_yellow_1(X26))|(~v13_waybel_0(X27,k3_yellow_1(X26))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X26)))))))))))&((v6_waybel_0(k3_yellow19(X25,X26,X27),X25)|(v2_struct_0(X25)|~l1_struct_0(X25)|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v1_xboole_0(X27)|(~v1_subset_1(X27,u1_struct_0(k3_yellow_1(X26)))|(~v2_waybel_0(X27,k3_yellow_1(X26))|(~v13_waybel_0(X27,k3_yellow_1(X26))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X26)))))))))))&(v7_waybel_0(k3_yellow19(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v1_xboole_0(X27)|(~v1_subset_1(X27,u1_struct_0(k3_yellow_1(X26)))|(~v2_waybel_0(X27,k3_yellow_1(X26))|(~v13_waybel_0(X27,k3_yellow_1(X26))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X26))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_yellow19])])])])).
fof(c_0_29, plain, ![X10, X11, X12, X13, X14, X15, X16, X17]:((((r2_hidden(X13,X10)|~r2_hidden(X13,X12)|X12!=k4_xboole_0(X10,X11))&(~r2_hidden(X13,X11)|~r2_hidden(X13,X12)|X12!=k4_xboole_0(X10,X11)))&(~r2_hidden(X14,X10)|r2_hidden(X14,X11)|r2_hidden(X14,X12)|X12!=k4_xboole_0(X10,X11)))&((~r2_hidden(esk4_3(X15,X16,X17),X17)|(~r2_hidden(esk4_3(X15,X16,X17),X15)|r2_hidden(esk4_3(X15,X16,X17),X16))|X17=k4_xboole_0(X15,X16))&((r2_hidden(esk4_3(X15,X16,X17),X15)|r2_hidden(esk4_3(X15,X16,X17),X17)|X17=k4_xboole_0(X15,X16))&(~r2_hidden(esk4_3(X15,X16,X17),X16)|r2_hidden(esk4_3(X15,X16,X17),X17)|X17=k4_xboole_0(X15,X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_30, plain, ![X35, X36, X37]:((~r3_waybel_9(X35,X36,X37)|r1_waybel_7(X35,k2_yellow19(X35,X36),X37)|~m1_subset_1(X37,u1_struct_0(X35))|(v2_struct_0(X36)|(~v4_orders_2(X36)|(~v7_waybel_0(X36)|~l1_waybel_0(X36,X35))))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))))&(~r1_waybel_7(X35,k2_yellow19(X35,X36),X37)|r3_waybel_9(X35,X36,X37)|~m1_subset_1(X37,u1_struct_0(X35))|(v2_struct_0(X36)|(~v4_orders_2(X36)|(~v7_waybel_0(X36)|~l1_waybel_0(X36,X35))))|(v2_struct_0(X35)|(~v2_pre_topc(X35)|~l1_pre_topc(X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_yellow19])])])])])).
cnf(c_0_31, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_32, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, negated_conjecture, (v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_35, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_37, negated_conjecture, (v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|~l1_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_40, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_41, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_42, plain, (v7_waybel_0(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_43, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_44, plain, ![X28, X29, X30]:((~v1_xboole_0(X28)|~r2_hidden(X29,X28))&(r2_hidden(esk5_1(X30),X30)|v1_xboole_0(X30))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_45, plain, (r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_47, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_48, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_49, plain, (l1_waybel_0(k3_yellow19(X1,k2_struct_0(X1),X2),X1)|v1_xboole_0(k2_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_50, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_33, c_0_21])).
cnf(c_0_51, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_52, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))))), inference(rw,[status(thm)],[c_0_36, c_0_21])).
cnf(c_0_53, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_37, c_0_21])).
cnf(c_0_54, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_55, plain, (X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|v2_struct_0(X1)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_21]), c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_56, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0)))))), inference(rw,[status(thm)],[c_0_39, c_0_21])).
cnf(c_0_57, plain, (v4_orders_2(k3_yellow19(X1,k2_struct_0(X1),X2))|v1_xboole_0(k2_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))), inference(spm,[status(thm)],[c_0_40, c_0_32])).
cnf(c_0_58, plain, (v1_xboole_0(k2_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v2_struct_0(k3_yellow19(X1,k2_struct_0(X1),X2))), inference(spm,[status(thm)],[c_0_41, c_0_32])).
cnf(c_0_59, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v7_waybel_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_21]), c_0_21]), c_0_21]), c_0_21])).
cnf(c_0_60, plain, (r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_61, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_62, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_63, plain, ![X40, X41, X42]:(~m1_subset_1(X41,k1_zfmisc_1(X40))|k7_subset_1(X40,X41,X42)=k4_xboole_0(X41,X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
fof(c_0_64, plain, ![X44, X45]:(~l1_struct_0(X44)|(~m1_subset_1(X45,k1_zfmisc_1(u1_struct_0(X44)))|k7_subset_1(u1_struct_0(X44),k2_struct_0(X44),k7_subset_1(u1_struct_0(X44),k2_struct_0(X44),X45))=X45)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_pre_topc])])])).
cnf(c_0_65, negated_conjecture, (r3_waybel_9(esk1_0,X1,esk3_0)|v2_struct_0(X1)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~r1_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_35]), c_0_47])]), c_0_48])).
cnf(c_0_66, negated_conjecture, (l1_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk1_0)|v1_xboole_0(k2_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52]), c_0_53])]), c_0_54]), c_0_48])).
cnf(c_0_67, negated_conjecture, (k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))=esk2_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_56])]), c_0_54]), c_0_48])).
cnf(c_0_68, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_69, negated_conjecture, (v4_orders_2(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))|v1_xboole_0(k2_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_50]), c_0_51]), c_0_52]), c_0_53])]), c_0_54]), c_0_48])).
cnf(c_0_70, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))|~v2_struct_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_50]), c_0_51]), c_0_52]), c_0_53])]), c_0_54]), c_0_48])).
cnf(c_0_71, plain, (v7_waybel_0(k3_yellow19(X1,k2_struct_0(X1),X2))|v1_xboole_0(k2_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1)))))), inference(spm,[status(thm)],[c_0_59, c_0_32])).
cnf(c_0_72, negated_conjecture, (r1_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)|v2_struct_0(X1)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~r3_waybel_9(esk1_0,X1,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_46]), c_0_35]), c_0_47])]), c_0_48])).
cnf(c_0_73, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_74, plain, (r2_hidden(esk5_1(k4_xboole_0(X1,X2)),X1)|v1_xboole_0(k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_75, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_76, plain, (k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_77, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|v1_xboole_0(k2_struct_0(esk1_0))|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_68]), c_0_69]), c_0_70])).
cnf(c_0_78, negated_conjecture, (v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))|v1_xboole_0(k2_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_50]), c_0_51]), c_0_52]), c_0_53]), c_0_56])]), c_0_54]), c_0_48])).
cnf(c_0_79, negated_conjecture, (r1_waybel_7(esk1_0,esk2_0,esk3_0)|v1_xboole_0(k2_struct_0(esk1_0))|~v7_waybel_0(k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_66]), c_0_67]), c_0_68]), c_0_69]), c_0_70])).
cnf(c_0_80, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_81, plain, (k4_xboole_0(k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_32])).
cnf(c_0_82, negated_conjecture, (~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_83, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|v1_xboole_0(k2_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_84, negated_conjecture, (r1_waybel_7(esk1_0,esk2_0,esk3_0)|v1_xboole_0(k2_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_79, c_0_78])).
fof(c_0_85, plain, ![X39]:m1_subset_1(k2_subset_1(X39),k1_zfmisc_1(X39)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_86, plain, ![X19]:k2_subset_1(X19)=X19, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_87, plain, ![X8]:(v2_struct_0(X8)|~l1_struct_0(X8)|~v1_xboole_0(u1_struct_0(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_88, plain, (v1_xboole_0(X1)|~l1_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_xboole_0(k2_struct_0(X2))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_89, negated_conjecture, (v1_xboole_0(k2_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])).
cnf(c_0_90, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_92, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_93, negated_conjecture, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_51])])).
cnf(c_0_94, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_95, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_51]), c_0_48])).
cnf(c_0_96, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_95]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 97
# Proof object clause steps            : 64
# Proof object formula steps           : 33
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 76
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 38
# Processed clauses                    : 2292
# ...of these trivial                  : 7
# ...subsumed                          : 1837
# ...remaining for further processing  : 448
# Other redundant clauses eliminated   : 114
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 20
# Backward-rewritten                   : 101
# Generated clauses                    : 6082
# ...of the previous two non-trivial   : 5526
# Contextual simplify-reflections      : 18
# Paramodulations                      : 5941
# Factorizations                       : 10
# Equation resolutions                 : 131
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 293
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 269
# Current number of unprocessed clauses: 3112
# ...number of literals in the above   : 17644
# Current number of archived formulas  : 0
# Current number of archived clauses   : 157
# Clause-clause subsumption calls (NU) : 117859
# Rec. Clause-clause subsumption calls : 16375
# Non-unit clause-clause subsumptions  : 1649
# Unit Clause-clause subsumption calls : 447
# Rewrite failures with RHS unbound    : 70
# BW rewrite match attempts            : 62
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 95223

# -------------------------------------------------
# User time                : 0.102 s
# System time              : 0.003 s
# Total time               : 0.105 s
# Maximum resident set size: 28728 pages
