# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.006 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t113_zfmisc_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', d1_funct_2)).
fof(cc1_partfun1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_partfun1(X3,X1))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', cc1_partfun1)).
fof(dt_k2_partfun1, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(v1_funct_1(k2_partfun1(X1,X2,X3,X4))&m1_subset_1(k2_partfun1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', dt_k2_partfun1)).
fof(redefinition_k2_partfun1, axiom, ![X1, X2, X3, X4]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>k2_partfun1(X1,X2,X3,X4)=k7_relat_1(X3,X4)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', redefinition_k2_partfun1)).
fof(t178_funct_2, conjecture, ![X1, X2, X3, X4]:(~(v1_xboole_0(X4))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X1,X4)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X4)))))=>((r1_tarski(X2,X1)&r1_tarski(k7_relset_1(X1,X4,X5,X2),X3))=>(v1_funct_1(k2_partfun1(X1,X4,X5,X2))&(v1_funct_2(k2_partfun1(X1,X4,X5,X2),X2,X3)&m1_subset_1(k2_partfun1(X1,X4,X5,X2),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))))))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t178_funct_2)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', redefinition_k1_relset_1)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', cc1_funct_2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', fc1_xboole_0)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', fc6_relat_1)).
fof(t110_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k7_relat_1(X1,k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t110_relat_1)).
fof(t91_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(X1,k1_relat_1(X2))=>k1_relat_1(k7_relat_1(X2,X1))=X1)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t91_relat_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t4_subset_1)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t87_relat_1)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', d19_relat_1)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t3_xboole_1)).
fof(t148_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k2_relat_1(k7_relat_1(X2,X1))=k9_relat_1(X2,X1)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t148_relat_1)).
fof(dt_k7_relat_1, axiom, ![X1, X2]:(v1_relat_1(X1)=>v1_relat_1(k7_relat_1(X1,X2))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', dt_k7_relat_1)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', redefinition_k7_relset_1)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t11_relset_1)).
fof(t152_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>~((~(X1=k1_xboole_0)&(r1_tarski(X1,k1_relat_1(X2))&k9_relat_1(X2,X1)=k1_xboole_0)))), file('data/MPTP/proofs/bab3d3f60a1945d69fa168199b98fa9f.E_input', t152_relat_1)).
fof(c_0_22, plain, ![X50, X51]:((k2_zfmisc_1(X50,X51)!=k1_xboole_0|(X50=k1_xboole_0|X51=k1_xboole_0))&((X50!=k1_xboole_0|k2_zfmisc_1(X50,X51)=k1_xboole_0)&(X51!=k1_xboole_0|k2_zfmisc_1(X50,X51)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_23, plain, ![X32, X33, X34]:((((~v1_funct_2(X34,X32,X33)|X32=k1_relset_1(X32,X33,X34)|X33=k1_xboole_0|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))&(X32!=k1_relset_1(X32,X33,X34)|v1_funct_2(X34,X32,X33)|X33=k1_xboole_0|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33)))))&((~v1_funct_2(X34,X32,X33)|X32=k1_relset_1(X32,X33,X34)|X32!=k1_xboole_0|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))&(X32!=k1_relset_1(X32,X33,X34)|v1_funct_2(X34,X32,X33)|X32!=k1_xboole_0|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))))&((~v1_funct_2(X34,X32,X33)|X34=k1_xboole_0|X32=k1_xboole_0|X33!=k1_xboole_0|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))&(X34!=k1_xboole_0|v1_funct_2(X34,X32,X33)|X32=k1_xboole_0|X33!=k1_xboole_0|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_24, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_25, plain, ![X44, X45, X46]:(~v1_xboole_0(X44)|(~m1_subset_1(X46,k1_zfmisc_1(k2_zfmisc_1(X44,X45)))|v1_partfun1(X46,X44))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_partfun1])])])).
fof(c_0_26, plain, ![X15, X16, X17, X18]:((v1_funct_1(k2_partfun1(X15,X16,X17,X18))|(~v1_funct_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(X15,X16)))))&(m1_subset_1(k2_partfun1(X15,X16,X17,X18),k1_zfmisc_1(k2_zfmisc_1(X15,X16)))|(~v1_funct_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(X15,X16)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_partfun1])])])).
fof(c_0_27, plain, ![X35, X36, X37, X38]:(~v1_funct_1(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(X35,X36)))|k2_partfun1(X35,X36,X37,X38)=k7_relat_1(X37,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_partfun1])])).
fof(c_0_28, negated_conjecture, ~(![X1, X2, X3, X4]:(~(v1_xboole_0(X4))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,X1,X4)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X4)))))=>((r1_tarski(X2,X1)&r1_tarski(k7_relset_1(X1,X4,X5,X2),X3))=>(v1_funct_1(k2_partfun1(X1,X4,X5,X2))&(v1_funct_2(k2_partfun1(X1,X4,X5,X2),X2,X3)&m1_subset_1(k2_partfun1(X1,X4,X5,X2),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))))))), inference(assume_negation,[status(cth)],[t178_funct_2])).
fof(c_0_29, plain, ![X47, X48, X49]:(~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|k1_relset_1(X47,X48,X49)=k1_relat_1(X49)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_30, plain, (X2=k1_relset_1(X2,X3,X1)|~v1_funct_2(X1,X2,X3)|X2!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_24])).
fof(c_0_32, plain, ![X29, X30, X31]:((v1_funct_1(X31)|(~v1_funct_1(X31)|~v1_partfun1(X31,X29))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30))))&(v1_funct_2(X31,X29,X30)|(~v1_funct_1(X31)|~v1_partfun1(X31,X29))|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X29,X30))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_33, plain, (v1_partfun1(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_35, plain, (v1_funct_1(k2_partfun1(X1,X2,X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (k2_partfun1(X2,X3,X1,X4)=k7_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_37, negated_conjecture, (~v1_xboole_0(esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,esk1_0,esk4_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk4_0)))))&((r1_tarski(esk2_0,esk1_0)&r1_tarski(k7_relset_1(esk1_0,esk4_0,esk5_0,esk2_0),esk3_0))&(~v1_funct_1(k2_partfun1(esk1_0,esk4_0,esk5_0,esk2_0))|(~v1_funct_2(k2_partfun1(esk1_0,esk4_0,esk5_0,esk2_0),esk2_0,esk3_0)|~m1_subset_1(k2_partfun1(esk1_0,esk4_0,esk5_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_28])])])])).
fof(c_0_38, plain, ![X19, X20]:(~v1_relat_1(X19)|(~m1_subset_1(X20,k1_zfmisc_1(X19))|v1_relat_1(X20))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_39, plain, ![X21, X22]:v1_relat_1(k2_zfmisc_1(X21,X22)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_40, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (k1_relset_1(k1_xboole_0,X1,X2)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))|~v1_funct_2(X2,k1_xboole_0,X1)), inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_30]), c_0_31])).
cnf(c_0_42, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (v1_partfun1(X1,k1_xboole_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_31]), c_0_34])])).
cnf(c_0_44, plain, (v1_funct_1(k7_relat_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45, plain, (m1_subset_1(k2_partfun1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_48, plain, ![X25]:(~v1_relat_1(X25)|k7_relat_1(X25,k1_xboole_0)=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t110_relat_1])])).
cnf(c_0_49, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_51, plain, ![X56, X57]:(~v1_relat_1(X57)|(~r1_tarski(X56,k1_relat_1(X57))|k1_relat_1(k7_relat_1(X57,X56))=X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t91_relat_1])])).
cnf(c_0_52, plain, (k1_relat_1(X1)=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_funct_2(X1,k1_xboole_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_31])])).
cnf(c_0_53, plain, (v1_funct_2(X1,k1_xboole_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_31]), c_0_43])).
cnf(c_0_54, plain, (v1_funct_1(k7_relat_1(k2_partfun1(X1,X2,X3,X4),X5))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_35])).
fof(c_0_55, plain, ![X43]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X43)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_56, negated_conjecture, (v1_funct_1(k7_relat_1(esk5_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_46]), c_0_47])])).
cnf(c_0_57, plain, (k7_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, negated_conjecture, (v1_relat_1(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_46]), c_0_50])])).
cnf(c_0_59, plain, (v1_relat_1(k2_partfun1(X1,X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_45]), c_0_50])])).
cnf(c_0_60, plain, (k1_relat_1(k7_relat_1(X1,X2))=X2|~v1_relat_1(X1)|~r1_tarski(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (k1_relat_1(X1)=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_62, plain, (v1_funct_1(k7_relat_1(k7_relat_1(X1,X2),X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_54, c_0_36])).
cnf(c_0_63, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64, negated_conjecture, (v1_funct_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])])).
cnf(c_0_65, plain, (v1_relat_1(k7_relat_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_59, c_0_36])).
fof(c_0_66, plain, ![X41, X42]:(~v1_relat_1(X42)|r1_tarski(k1_relat_1(k7_relat_1(X42,X41)),X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
cnf(c_0_67, plain, (k1_xboole_0=X1|~v1_relat_1(X2)|~r1_tarski(X1,k1_relat_1(X2))|~m1_subset_1(k7_relat_1(X2,X1),k1_zfmisc_1(k1_xboole_0))|~v1_funct_1(k7_relat_1(X2,X1))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_68, plain, (v1_funct_1(k7_relat_1(k7_relat_1(k1_xboole_0,X1),X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])])).
cnf(c_0_69, plain, (v1_relat_1(k7_relat_1(k1_xboole_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_63]), c_0_64])])).
cnf(c_0_70, plain, (m1_subset_1(k2_partfun1(k1_xboole_0,X1,X2,X3),k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))|~v1_funct_1(X2)), inference(spm,[status(thm)],[c_0_45, c_0_31])).
cnf(c_0_71, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_72, plain, (k1_xboole_0=X1|~r1_tarski(X1,k1_relat_1(k7_relat_1(k1_xboole_0,X2)))|~m1_subset_1(k7_relat_1(k7_relat_1(k1_xboole_0,X2),X1),k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69])])).
cnf(c_0_73, plain, (m1_subset_1(k7_relat_1(X1,X2),k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_36]), c_0_31])])).
cnf(c_0_74, plain, (v1_funct_1(k7_relat_1(k1_xboole_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_63]), c_0_64])])).
fof(c_0_75, plain, ![X23, X24]:((~v5_relat_1(X24,X23)|r1_tarski(k2_relat_1(X24),X23)|~v1_relat_1(X24))&(~r1_tarski(k2_relat_1(X24),X23)|v5_relat_1(X24,X23)|~v1_relat_1(X24))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
cnf(c_0_76, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_77, plain, ![X58]:(~r1_tarski(X58,k1_xboole_0)|X58=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_78, plain, (r1_tarski(k1_relat_1(k1_xboole_0),k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_71, c_0_57])).
cnf(c_0_79, plain, (k1_xboole_0=X1|~r1_tarski(X1,k1_relat_1(k7_relat_1(k1_xboole_0,X2)))|~m1_subset_1(k7_relat_1(k1_xboole_0,X2),k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74])])).
cnf(c_0_80, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_81, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_82, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k1_relat_1(k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_78, c_0_58])).
cnf(c_0_84, plain, (k2_relat_1(X1)=k1_xboole_0|~v5_relat_1(X1,k1_relat_1(k7_relat_1(k1_xboole_0,X2)))|~v1_relat_1(X1)|~m1_subset_1(k7_relat_1(k1_xboole_0,X2),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_85, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_50, c_0_81])).
cnf(c_0_86, negated_conjecture, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_82, c_0_83])).
fof(c_0_87, plain, ![X54, X55]:(~v1_relat_1(X55)|k2_relat_1(k7_relat_1(X55,X54))=k9_relat_1(X55,X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t148_relat_1])])).
fof(c_0_88, plain, ![X52, X53]:(~v1_relat_1(X52)|v1_relat_1(k7_relat_1(X52,X53))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_relat_1])])).
cnf(c_0_89, plain, (k2_relat_1(X1)=k1_xboole_0|~v5_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_tarski(X2,k1_xboole_0)|~m1_subset_1(k7_relat_1(k1_xboole_0,X2),k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_60]), c_0_85]), c_0_86])])).
cnf(c_0_90, plain, (v5_relat_1(X1,X2)|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_91, plain, (k2_relat_1(k7_relat_1(X1,X2))=k9_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_92, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_93, plain, ![X11, X12, X13, X14]:(~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X11,X12)))|k7_relset_1(X11,X12,X13,X14)=k9_relat_1(X13,X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
fof(c_0_94, plain, ![X26, X27, X28]:(~v1_relat_1(X28)|(~r1_tarski(k1_relat_1(X28),X26)|~r1_tarski(k2_relat_1(X28),X27)|m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_95, plain, (k2_relat_1(X1)=k1_xboole_0|~v5_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_tarski(X2,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_73]), c_0_63]), c_0_64])])).
cnf(c_0_96, plain, (v5_relat_1(k7_relat_1(X1,X2),X3)|~v1_relat_1(X1)|~r1_tarski(k9_relat_1(X1,X2),X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92])).
cnf(c_0_97, negated_conjecture, (r1_tarski(k7_relset_1(esk1_0,esk4_0,esk5_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_98, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_99, plain, (r1_tarski(k1_xboole_0,X1)|~v1_relat_1(X2)|~m1_subset_1(k7_relat_1(X2,X1),k1_zfmisc_1(k1_xboole_0))|~v1_funct_1(k7_relat_1(X2,X1))), inference(spm,[status(thm)],[c_0_71, c_0_61])).
cnf(c_0_100, plain, (v1_funct_2(X3,X1,X2)|X2=k1_xboole_0|X1!=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_101, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_102, negated_conjecture, (~v1_funct_1(k2_partfun1(esk1_0,esk4_0,esk5_0,esk2_0))|~v1_funct_2(k2_partfun1(esk1_0,esk4_0,esk5_0,esk2_0),esk2_0,esk3_0)|~m1_subset_1(k2_partfun1(esk1_0,esk4_0,esk5_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_103, plain, (k2_relat_1(k7_relat_1(X1,X2))=k1_xboole_0|~v1_relat_1(X1)|~r1_tarski(k9_relat_1(X1,X2),X3)|~r1_tarski(X3,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_92])).
cnf(c_0_104, negated_conjecture, (r1_tarski(k9_relat_1(esk5_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_46])])).
cnf(c_0_105, plain, (r1_tarski(k1_xboole_0,X1)|~m1_subset_1(k7_relat_1(k1_xboole_0,X1),k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_74]), c_0_85])])).
cnf(c_0_106, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_107, plain, (X1=k1_xboole_0|v1_funct_2(X2,k1_relat_1(X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1)))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_40])])).
cnf(c_0_108, plain, (m1_subset_1(k7_relat_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X3)|~r1_tarski(k9_relat_1(X1,X2),X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_91]), c_0_92])).
cnf(c_0_109, negated_conjecture, (~m1_subset_1(k7_relat_1(esk5_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(k7_relat_1(esk5_0,esk2_0),esk2_0,esk3_0)|~v1_funct_1(k7_relat_1(esk5_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_36]), c_0_46]), c_0_47])])).
cnf(c_0_110, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v5_relat_1(X1,X3)|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)), inference(spm,[status(thm)],[c_0_101, c_0_80])).
cnf(c_0_111, negated_conjecture, (k2_relat_1(k7_relat_1(esk5_0,esk2_0))=k1_xboole_0|~r1_tarski(esk3_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_58])])).
cnf(c_0_112, negated_conjecture, (v1_relat_1(k7_relat_1(esk5_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_46]), c_0_47])])).
cnf(c_0_113, plain, (r1_tarski(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_73]), c_0_63]), c_0_64])])).
fof(c_0_114, plain, ![X39, X40]:(~v1_relat_1(X40)|(X39=k1_xboole_0|(~r1_tarski(X39,k1_relat_1(X40))|k9_relat_1(X40,X39)!=k1_xboole_0))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_relat_1])])])).
cnf(c_0_115, plain, (X1=k1_relat_1(X2)|X3=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~v1_funct_2(X2,X1,X3)), inference(spm,[status(thm)],[c_0_40, c_0_106])).
cnf(c_0_116, negated_conjecture, (v1_funct_2(esk5_0,esk1_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_117, plain, (X1=k1_xboole_0|v1_funct_2(k7_relat_1(X2,X3),X3,X1)|~v1_relat_1(X2)|~r1_tarski(X3,k1_relat_1(X2))|~m1_subset_1(k7_relat_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(spm,[status(thm)],[c_0_107, c_0_60])).
cnf(c_0_118, plain, (m1_subset_1(k7_relat_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_relat_1(X1)|~r1_tarski(k9_relat_1(X1,X2),X4)|~r1_tarski(X2,k1_relat_1(X1))|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_108, c_0_60])).
cnf(c_0_119, plain, (r1_tarski(X1,X1)|~v1_relat_1(X2)|~r1_tarski(X1,k1_relat_1(X2))), inference(spm,[status(thm)],[c_0_71, c_0_60])).
cnf(c_0_120, negated_conjecture, (~m1_subset_1(k7_relat_1(esk5_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,esk3_0)))|~v1_funct_2(k7_relat_1(esk5_0,esk2_0),esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_56])])).
cnf(c_0_121, plain, (m1_subset_1(k7_relat_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k9_relat_1(X1,X2),X3)), inference(spm,[status(thm)],[c_0_108, c_0_71])).
cnf(c_0_122, plain, (m1_subset_1(k7_relat_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v5_relat_1(k7_relat_1(X1,X2),X3)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_71]), c_0_92])).
cnf(c_0_123, negated_conjecture, (v5_relat_1(k7_relat_1(esk5_0,esk2_0),X1)|~r1_tarski(esk3_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_111]), c_0_112]), c_0_113])])).
cnf(c_0_124, plain, (X2=k1_xboole_0|~v1_relat_1(X1)|~r1_tarski(X2,k1_relat_1(X1))|k9_relat_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_125, negated_conjecture, (k1_relat_1(esk5_0)=esk1_0|esk4_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_46]), c_0_116])])).
cnf(c_0_126, plain, (X1=k1_xboole_0|v1_funct_2(k7_relat_1(X2,X3),X3,X1)|~v1_relat_1(X2)|~r1_tarski(k9_relat_1(X2,X3),X1)|~r1_tarski(X3,k1_relat_1(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])).
cnf(c_0_127, negated_conjecture, (~v1_funct_2(k7_relat_1(esk5_0,esk2_0),esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_58]), c_0_104])])).
cnf(c_0_128, negated_conjecture, (m1_subset_1(k7_relat_1(esk5_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk2_0,X1)))|~r1_tarski(esk3_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_58])])).
cnf(c_0_129, plain, (v1_partfun1(k7_relat_1(X1,X2),X2)|~v1_relat_1(X1)|~r1_tarski(k9_relat_1(X1,X2),X3)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_33, c_0_121])).
cnf(c_0_130, negated_conjecture, (esk4_0=k1_xboole_0|X1=k1_xboole_0|k9_relat_1(esk5_0,X1)!=k1_xboole_0|~r1_tarski(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_58])])).
cnf(c_0_131, negated_conjecture, (k9_relat_1(esk5_0,esk2_0)=k1_xboole_0|~r1_tarski(esk3_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_111]), c_0_58])])).
cnf(c_0_132, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_133, negated_conjecture, (esk3_0=k1_xboole_0|~r1_tarski(esk2_0,k1_relat_1(esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_104]), c_0_58])]), c_0_127])).
cnf(c_0_134, negated_conjecture, (v1_funct_2(k7_relat_1(esk5_0,esk2_0),esk2_0,X1)|~v1_partfun1(k7_relat_1(esk5_0,esk2_0),esk2_0)|~r1_tarski(esk3_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_128]), c_0_56])])).
cnf(c_0_135, negated_conjecture, (v1_partfun1(k7_relat_1(esk5_0,esk2_0),esk2_0)|~v1_xboole_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_104]), c_0_58])])).
cnf(c_0_136, negated_conjecture, (esk2_0=k1_xboole_0|esk4_0=k1_xboole_0|~r1_tarski(esk3_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_132])])).
cnf(c_0_137, negated_conjecture, (esk4_0=k1_xboole_0|esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_125]), c_0_132])])).
cnf(c_0_138, negated_conjecture, (v1_funct_2(k7_relat_1(esk5_0,esk2_0),esk2_0,X1)|~r1_tarski(esk3_0,k1_xboole_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_134, c_0_135])).
cnf(c_0_139, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_140, negated_conjecture, (esk4_0=k1_xboole_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_113])])).
cnf(c_0_141, negated_conjecture, (~r1_tarski(esk3_0,k1_xboole_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_127, c_0_138])).
cnf(c_0_142, negated_conjecture, (esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140]), c_0_34])])).
cnf(c_0_143, negated_conjecture, (~r1_tarski(esk3_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_141, c_0_142]), c_0_34])])).
cnf(c_0_144, negated_conjecture, (esk3_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_142]), c_0_113])])).
cnf(c_0_145, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_143, c_0_144]), c_0_113])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 146
# Proof object clause steps            : 102
# Proof object formula steps           : 44
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 62
# Proof object simplifying inferences  : 96
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 38
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 37
# Processed clauses                    : 10627
# ...of these trivial                  : 27
# ...subsumed                          : 8780
# ...remaining for further processing  : 1820
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 334
# Backward-rewritten                   : 436
# Generated clauses                    : 35954
# ...of the previous two non-trivial   : 33848
# Contextual simplify-reflections      : 200
# Paramodulations                      : 35941
# Factorizations                       : 3
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1007
#    Positive orientable unit clauses  : 30
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 976
# Current number of unprocessed clauses: 22366
# ...number of literals in the above   : 158519
# Current number of archived formulas  : 0
# Current number of archived clauses   : 807
# Clause-clause subsumption calls (NU) : 485858
# Rec. Clause-clause subsumption calls : 68368
# Non-unit clause-clause subsumptions  : 9298
# Unit Clause-clause subsumption calls : 1004
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 39
# BW rewrite match successes           : 21
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 724521

# -------------------------------------------------
# User time                : 0.441 s
# System time              : 0.003 s
# Total time               : 0.444 s
# Maximum resident set size: 28768 pages
