# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', d4_tarski)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', d1_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', fc1_xboole_0)).
fof(t115_funct_2, conjecture, ![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:~((r2_hidden(X6,X1)&(r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6))))))), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', t115_funct_2)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', t2_subset)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', fc1_subset_1)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', t99_zfmisc_1)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', t143_relat_1)).
fof(cc1_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_relat_1(X1)), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', cc1_relat_1)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', redefinition_k7_relset_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', t8_funct_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', cc1_relset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', redefinition_k1_relset_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(((X2=k1_xboole_0=>X1=k1_xboole_0)=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X2,X3)))&(X2=k1_xboole_0=>(X1=k1_xboole_0|(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0))))), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', d1_funct_2)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/bace4004d867452b9cdda1a69f57b300.E_input', t113_zfmisc_1)).
fof(c_0_15, plain, ![X45, X46, X47, X49, X50, X51, X52, X54]:((((r2_hidden(X47,esk8_3(X45,X46,X47))|~r2_hidden(X47,X46)|X46!=k3_tarski(X45))&(r2_hidden(esk8_3(X45,X46,X47),X45)|~r2_hidden(X47,X46)|X46!=k3_tarski(X45)))&(~r2_hidden(X49,X50)|~r2_hidden(X50,X45)|r2_hidden(X49,X46)|X46!=k3_tarski(X45)))&((~r2_hidden(esk9_2(X51,X52),X52)|(~r2_hidden(esk9_2(X51,X52),X54)|~r2_hidden(X54,X51))|X52=k3_tarski(X51))&((r2_hidden(esk9_2(X51,X52),esk10_2(X51,X52))|r2_hidden(esk9_2(X51,X52),X52)|X52=k3_tarski(X51))&(r2_hidden(esk10_2(X51,X52),X51)|r2_hidden(esk9_2(X51,X52),X52)|X52=k3_tarski(X51))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
fof(c_0_16, plain, ![X36, X37, X38]:((~v1_xboole_0(X36)|~r2_hidden(X37,X36))&(r2_hidden(esk7_1(X38),X38)|v1_xboole_0(X38))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_17, plain, (r2_hidden(esk8_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|X2!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_19, plain, (r2_hidden(esk8_3(X1,k3_tarski(X1),X2),X1)|~r2_hidden(X2,k3_tarski(X1))), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_20, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21, plain, (r2_hidden(esk10_2(X1,X2),X1)|r2_hidden(esk9_2(X1,X2),X2)|X2=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(X1)))), inference(spm,[status(thm)],[c_0_20, c_0_19])).
cnf(c_0_23, plain, (X1=k3_tarski(k3_tarski(X2))|r2_hidden(esk9_2(k3_tarski(X2),X1),X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(X1))))), inference(spm,[status(thm)],[c_0_22, c_0_19])).
cnf(c_0_25, plain, (r2_hidden(esk7_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (X1=k3_tarski(k3_tarski(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_18, c_0_23])).
cnf(c_0_27, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_28, negated_conjecture, ~(![X1, X2, X3, X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X5]:~((r2_hidden(X5,k7_relset_1(X1,X2,X4,X3))&![X6]:~((r2_hidden(X6,X1)&(r2_hidden(X6,X3)&X5=k1_funct_1(X4,X6)))))))), inference(assume_negation,[status(cth)],[t115_funct_2])).
cnf(c_0_29, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(k3_tarski(X1)))))), inference(spm,[status(thm)],[c_0_24, c_0_19])).
cnf(c_0_30, plain, (v1_xboole_0(k3_tarski(k3_tarski(k3_tarski(X1))))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, plain, (k3_tarski(k3_tarski(X1))=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_32, plain, ![X20, X21]:(~m1_subset_1(X20,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_33, negated_conjecture, ![X12]:((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,esk2_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))&(r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))&(~r2_hidden(X12,esk1_0)|(~r2_hidden(X12,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X12))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])])).
fof(c_0_34, plain, ![X22]:~v1_xboole_0(k1_zfmisc_1(X22)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_35, plain, (v1_xboole_0(k3_tarski(k3_tarski(k3_tarski(k3_tarski(X1)))))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_29, c_0_25])).
cnf(c_0_36, plain, (v1_xboole_0(k3_tarski(k1_xboole_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (r2_hidden(X1,X4)|~r2_hidden(X1,X2)|~r2_hidden(X2,X3)|X4!=k3_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_38, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_41, plain, ![X40]:k3_tarski(k1_zfmisc_1(X40))=X40, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
fof(c_0_42, plain, ![X23, X24, X25, X27]:(((r2_hidden(esk6_3(X23,X24,X25),k1_relat_1(X25))|~r2_hidden(X23,k9_relat_1(X25,X24))|~v1_relat_1(X25))&((r2_hidden(k4_tarski(esk6_3(X23,X24,X25),X23),X25)|~r2_hidden(X23,k9_relat_1(X25,X24))|~v1_relat_1(X25))&(r2_hidden(esk6_3(X23,X24,X25),X24)|~r2_hidden(X23,k9_relat_1(X25,X24))|~v1_relat_1(X25))))&(~r2_hidden(X27,k1_relat_1(X25))|(~r2_hidden(k4_tarski(X27,X23),X25)|~r2_hidden(X27,X24))|r2_hidden(X23,k9_relat_1(X25,X24))|~v1_relat_1(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
fof(c_0_43, plain, ![X44]:(~v1_xboole_0(X44)|v1_relat_1(X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relat_1])])).
fof(c_0_44, plain, ![X16, X17, X18, X19]:(~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X16,X17)))|k7_relset_1(X16,X17,X18,X19)=k9_relat_1(X18,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
cnf(c_0_45, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_22, c_0_31])).
cnf(c_0_46, plain, (X1=k3_tarski(X2)|r2_hidden(esk9_2(X2,X1),X1)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_18, c_0_21])).
cnf(c_0_47, plain, (v1_xboole_0(k3_tarski(k3_tarski(k1_xboole_0)))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_35, c_0_31])).
cnf(c_0_48, plain, (v1_xboole_0(k3_tarski(k1_xboole_0))), inference(spm,[status(thm)],[c_0_36, c_0_27])).
cnf(c_0_49, plain, (r2_hidden(X1,k3_tarski(X2))|~r2_hidden(X3,X2)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_51, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (r2_hidden(k4_tarski(esk6_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, negated_conjecture, (r2_hidden(esk5_0,k7_relset_1(esk1_0,esk2_0,esk4_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_55, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, plain, (k3_tarski(X1)=k1_xboole_0|~v1_xboole_0(X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_57, plain, (v1_xboole_0(k3_tarski(k3_tarski(k1_xboole_0)))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_58, negated_conjecture, (r2_hidden(X1,k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X1,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
cnf(c_0_59, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_52]), c_0_53])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk5_0,k9_relat_1(esk4_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_39])])).
cnf(c_0_61, plain, (k3_tarski(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_62, negated_conjecture, (r2_hidden(X1,k3_tarski(k2_zfmisc_1(esk1_0,esk2_0)))|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_49, c_0_58])).
cnf(c_0_63, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
fof(c_0_64, plain, ![X41, X42, X43]:(((r2_hidden(X41,k1_relat_1(X43))|~r2_hidden(k4_tarski(X41,X42),X43)|(~v1_relat_1(X43)|~v1_funct_1(X43)))&(X42=k1_funct_1(X43,X41)|~r2_hidden(k4_tarski(X41,X42),X43)|(~v1_relat_1(X43)|~v1_funct_1(X43))))&(~r2_hidden(X41,k1_relat_1(X43))|X42!=k1_funct_1(X43,X41)|r2_hidden(k4_tarski(X41,X42),X43)|(~v1_relat_1(X43)|~v1_funct_1(X43)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
fof(c_0_65, plain, ![X31, X32, X33]:(~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))|v1_relat_1(X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_66, plain, ![X28, X29, X30]:(~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|k1_relset_1(X28,X29,X30)=k1_relat_1(X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_67, plain, (k3_tarski(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_61, c_0_27])).
cnf(c_0_68, negated_conjecture, (r2_hidden(X1,k3_tarski(k2_zfmisc_1(esk1_0,esk2_0)))|~r2_hidden(X1,esk7_1(esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_25]), c_0_63])).
cnf(c_0_69, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_70, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_71, plain, ![X13, X14, X15]:((((~v1_funct_2(X15,X13,X14)|X13=k1_relset_1(X13,X14,X15)|X14=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X13!=k1_relset_1(X13,X14,X15)|v1_funct_2(X15,X13,X14)|X14=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))&((~v1_funct_2(X15,X13,X14)|X13=k1_relset_1(X13,X14,X15)|X13!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X13!=k1_relset_1(X13,X14,X15)|v1_funct_2(X15,X13,X14)|X13!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))))&((~v1_funct_2(X15,X13,X14)|X15=k1_xboole_0|X13=k1_xboole_0|X14!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14))))&(X15!=k1_xboole_0|v1_funct_2(X15,X13,X14)|X13=k1_xboole_0|X14!=k1_xboole_0|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(X13,X14)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
cnf(c_0_72, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_73, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_67]), c_0_27])])).
cnf(c_0_74, plain, (X1=k3_tarski(X2)|r2_hidden(esk10_2(X2,X1),X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_18, c_0_21])).
cnf(c_0_75, negated_conjecture, (v1_xboole_0(esk7_1(esk4_0))|r2_hidden(esk7_1(esk7_1(esk4_0)),k3_tarski(k2_zfmisc_1(esk1_0,esk2_0)))), inference(spm,[status(thm)],[c_0_68, c_0_25])).
cnf(c_0_76, negated_conjecture, (~r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)|esk5_0!=k1_funct_1(esk4_0,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_77, plain, (k1_funct_1(X1,esk6_3(X2,X3,X1))=X2|~v1_relat_1(X1)|~r2_hidden(X2,k9_relat_1(X1,X3))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_69, c_0_52])).
cnf(c_0_78, negated_conjecture, (v1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_70, c_0_39])).
cnf(c_0_79, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_80, plain, (X2=k1_relset_1(X2,X3,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_81, negated_conjecture, (k1_relset_1(esk1_0,esk2_0,esk4_0)=k1_relat_1(esk4_0)), inference(spm,[status(thm)],[c_0_72, c_0_39])).
cnf(c_0_82, negated_conjecture, (v1_funct_2(esk4_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_83, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_67])).
cnf(c_0_84, negated_conjecture, (v1_xboole_0(esk7_1(esk4_0))|~v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_20, c_0_75])).
cnf(c_0_85, negated_conjecture, (~r2_hidden(esk6_3(esk5_0,X1,esk4_0),esk1_0)|~r2_hidden(esk6_3(esk5_0,X1,esk4_0),esk3_0)|~r2_hidden(esk5_0,k9_relat_1(esk4_0,X1))), inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), c_0_79])])])).
cnf(c_0_86, plain, (r2_hidden(esk6_3(X1,X2,X3),X2)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_87, plain, (r2_hidden(esk6_3(X1,X2,X3),k1_relat_1(X3))|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_88, negated_conjecture, (k1_relat_1(esk4_0)=esk1_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_39]), c_0_81]), c_0_82])])).
fof(c_0_89, plain, ![X34, X35]:((k2_zfmisc_1(X34,X35)!=k1_xboole_0|(X34=k1_xboole_0|X35=k1_xboole_0))&((X34!=k1_xboole_0|k2_zfmisc_1(X34,X35)=k1_xboole_0)&(X35!=k1_xboole_0|k2_zfmisc_1(X34,X35)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_90, negated_conjecture, (esk7_1(esk4_0)=k1_xboole_0|~v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_91, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X1,esk4_0)), inference(spm,[status(thm)],[c_0_18, c_0_58])).
cnf(c_0_92, negated_conjecture, (~r2_hidden(esk6_3(esk5_0,esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_60]), c_0_78])])).
cnf(c_0_93, negated_conjecture, (esk2_0=k1_xboole_0|r2_hidden(esk6_3(X1,X2,esk4_0),esk1_0)|~r2_hidden(X1,k9_relat_1(esk4_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_78])])).
cnf(c_0_94, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_95, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_90]), c_0_63]), c_0_91])).
cnf(c_0_96, negated_conjecture, (esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_60])])).
cnf(c_0_97, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_94])).
cnf(c_0_98, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_96]), c_0_97]), c_0_27])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 99
# Proof object clause steps            : 69
# Proof object formula steps           : 30
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 41
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 37
# Processed clauses                    : 475
# ...of these trivial                  : 2
# ...subsumed                          : 197
# ...remaining for further processing  : 276
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 66
# Generated clauses                    : 1652
# ...of the previous two non-trivial   : 1576
# Contextual simplify-reflections      : 11
# Paramodulations                      : 1641
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 161
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 143
# Current number of unprocessed clauses: 1124
# ...number of literals in the above   : 4383
# Current number of archived formulas  : 0
# Current number of archived clauses   : 105
# Clause-clause subsumption calls (NU) : 15698
# Rec. Clause-clause subsumption calls : 7398
# Non-unit clause-clause subsumptions  : 55
# Unit Clause-clause subsumption calls : 194
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 25180

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.000 s
# Total time               : 0.030 s
# Maximum resident set size: 28720 pages
