# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/bb5048c9eecb4622897a9aaffc042e7c.E_input', d5_funct_1)).
fof(t14_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(k1_relat_1(X2)=k1_tarski(X1)=>k2_relat_1(X2)=k1_tarski(k1_funct_1(X2,X1)))), file('data/MPTP/proofs/bb5048c9eecb4622897a9aaffc042e7c.E_input', t14_funct_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/bb5048c9eecb4622897a9aaffc042e7c.E_input', d1_tarski)).
fof(t18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>~((r2_hidden(X1,k1_relat_1(X2))&![X3]:~(r2_hidden(X3,k2_relat_1(X2)))))), file('data/MPTP/proofs/bb5048c9eecb4622897a9aaffc042e7c.E_input', t18_relat_1)).
fof(l44_zfmisc_1, axiom, ![X1, X2]:~((~(X1=k1_tarski(X2))&(~(X1=k1_xboole_0)&![X3]:~((r2_hidden(X3,X1)&~(X3=X2)))))), file('data/MPTP/proofs/bb5048c9eecb4622897a9aaffc042e7c.E_input', l44_zfmisc_1)).
fof(t49_zfmisc_1, axiom, ![X1, X2]:~(k2_xboole_0(k1_tarski(X1),X2)=k1_xboole_0), file('data/MPTP/proofs/bb5048c9eecb4622897a9aaffc042e7c.E_input', t49_zfmisc_1)).
fof(l22_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>k2_xboole_0(k1_tarski(X1),X2)=X2), file('data/MPTP/proofs/bb5048c9eecb4622897a9aaffc042e7c.E_input', l22_zfmisc_1)).
fof(c_0_7, plain, ![X19, X20, X21, X23, X24, X25, X27]:((((r2_hidden(esk5_3(X19,X20,X21),k1_relat_1(X19))|~r2_hidden(X21,X20)|X20!=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))&(X21=k1_funct_1(X19,esk5_3(X19,X20,X21))|~r2_hidden(X21,X20)|X20!=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19))))&(~r2_hidden(X24,k1_relat_1(X19))|X23!=k1_funct_1(X19,X24)|r2_hidden(X23,X20)|X20!=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19))))&((~r2_hidden(esk6_2(X19,X25),X25)|(~r2_hidden(X27,k1_relat_1(X19))|esk6_2(X19,X25)!=k1_funct_1(X19,X27))|X25=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))&((r2_hidden(esk7_2(X19,X25),k1_relat_1(X19))|r2_hidden(esk6_2(X19,X25),X25)|X25=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))&(esk6_2(X19,X25)=k1_funct_1(X19,esk7_2(X19,X25))|r2_hidden(esk6_2(X19,X25),X25)|X25=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
fof(c_0_8, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(k1_relat_1(X2)=k1_tarski(X1)=>k2_relat_1(X2)=k1_tarski(k1_funct_1(X2,X1))))), inference(assume_negation,[status(cth)],[t14_funct_1])).
fof(c_0_9, plain, ![X7, X8, X9, X10, X11, X12]:(((~r2_hidden(X9,X8)|X9=X7|X8!=k1_tarski(X7))&(X10!=X7|r2_hidden(X10,X8)|X8!=k1_tarski(X7)))&((~r2_hidden(esk3_2(X11,X12),X12)|esk3_2(X11,X12)!=X11|X12=k1_tarski(X11))&(r2_hidden(esk3_2(X11,X12),X12)|esk3_2(X11,X12)=X11|X12=k1_tarski(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_10, plain, (r2_hidden(esk5_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&(k1_relat_1(esk2_0)=k1_tarski(esk1_0)&k2_relat_1(esk2_0)!=k1_tarski(k1_funct_1(esk2_0,esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])])).
cnf(c_0_12, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, plain, (r2_hidden(esk5_3(X1,k2_relat_1(X1),X2),k1_relat_1(X1))|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_14, negated_conjecture, (k1_relat_1(esk2_0)=k1_tarski(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (X1=k1_funct_1(X2,esk5_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_18, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (r2_hidden(esk5_3(esk2_0,k2_relat_1(esk2_0),X1),k1_tarski(esk1_0))|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_20, plain, (k1_funct_1(X1,esk5_3(X1,k2_relat_1(X1),X2))=X2|~r2_hidden(X2,k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21, negated_conjecture, (esk5_3(esk2_0,k2_relat_1(esk2_0),X1)=esk1_0|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_22, plain, ![X16, X17]:(~v1_relat_1(X17)|(~r2_hidden(X16,k1_relat_1(X17))|r2_hidden(esk4_2(X16,X17),k2_relat_1(X17)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_relat_1])])])])).
cnf(c_0_23, negated_conjecture, (k1_funct_1(esk2_0,esk1_0)=X1|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_15]), c_0_16])])).
cnf(c_0_24, plain, (r2_hidden(esk4_2(X2,X1),k2_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_25, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26, negated_conjecture, (k1_funct_1(esk2_0,esk1_0)=esk4_2(X1,esk2_0)|~r2_hidden(X1,k1_tarski(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_14]), c_0_16])])).
cnf(c_0_27, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_25])])).
cnf(c_0_28, negated_conjecture, (k1_funct_1(esk2_0,esk1_0)=esk4_2(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_29, plain, ![X29, X30]:((r2_hidden(esk8_2(X29,X30),X29)|X29=k1_xboole_0|X29=k1_tarski(X30))&(esk8_2(X29,X30)!=X30|X29=k1_xboole_0|X29=k1_tarski(X30))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l44_zfmisc_1])])])])])).
cnf(c_0_30, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31, negated_conjecture, (esk4_2(esk1_0,esk2_0)=X1|~r2_hidden(X1,k2_relat_1(esk2_0))), inference(rw,[status(thm)],[c_0_23, c_0_28])).
cnf(c_0_32, plain, (r2_hidden(esk8_2(X1,X2),X1)|X1=k1_xboole_0|X1=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33, negated_conjecture, (k2_relat_1(esk2_0)!=k1_tarski(k1_funct_1(esk2_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_34, plain, ![X32, X33]:k2_xboole_0(k1_tarski(X32),X33)!=k1_xboole_0, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_zfmisc_1])])).
fof(c_0_35, plain, ![X14, X15]:(~r2_hidden(X14,X15)|k2_xboole_0(k1_tarski(X14),X15)=X15), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l22_zfmisc_1])])).
cnf(c_0_36, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_30])])).
cnf(c_0_37, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|esk8_2(X1,X2)!=X2), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, negated_conjecture, (esk4_2(esk1_0,esk2_0)=esk8_2(k2_relat_1(esk2_0),X1)|k1_tarski(X1)=k2_relat_1(esk2_0)|k2_relat_1(esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39, negated_conjecture, (k1_tarski(esk4_2(esk1_0,esk2_0))!=k2_relat_1(esk2_0)), inference(rw,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_40, plain, (k2_xboole_0(k1_tarski(X1),X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (k2_xboole_0(k1_tarski(X1),X2)=X2|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk4_2(esk1_0,esk2_0),k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_28]), c_0_14]), c_0_27]), c_0_15]), c_0_16])])).
cnf(c_0_43, negated_conjecture, (k2_relat_1(esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38])]), c_0_39])).
cnf(c_0_44, plain, (~r2_hidden(X1,k1_xboole_0)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41])])).
cnf(c_0_45, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_43]), c_0_44]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 46
# Proof object clause steps            : 31
# Proof object formula steps           : 15
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 57
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 52
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 22
# Generated clauses                    : 96
# ...of the previous two non-trivial   : 84
# Contextual simplify-reflections      : 0
# Paramodulations                      : 88
# Factorizations                       : 0
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 25
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 17
# Current number of unprocessed clauses: 34
# ...number of literals in the above   : 127
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 168
# Rec. Clause-clause subsumption calls : 88
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 37
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2961

# -------------------------------------------------
# User time                : 0.002 s
# System time              : 0.008 s
# Total time               : 0.010 s
# Maximum resident set size: 28648 pages
