# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t34_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2))))))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', t34_lattice3)).
fof(t45_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(X2=k15_lattice3(X1,a_2_3_lattice3(X1,X2))&X2=k16_lattice3(X1,a_2_4_lattice3(X1,X2))))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', t45_lattice3)).
fof(fraenkel_a_2_5_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_5_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r3_lattices(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', fraenkel_a_2_5_lattice3)).
fof(t48_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r3_lattices(X1,X4,X5)&r2_hidden(X5,X3)))))))=>r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3)))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', t48_lattice3)).
fof(t38_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r2_hidden(X2,X3)=>(r3_lattices(X1,X2,k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),X2))))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', t38_lattice3)).
fof(t47_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))&k16_lattice3(X1,X2)=k16_lattice3(X1,a_2_6_lattice3(X1,X2)))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', t47_lattice3)).
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', d16_lattice3)).
fof(fraenkel_a_2_4_lattice3, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))&m1_subset_1(X3,u1_struct_0(X2)))=>(r2_hidden(X1,a_2_4_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattices(X2,X3,X4))))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', fraenkel_a_2_4_lattice3)).
fof(t37_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', t37_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', d17_lattice3)).
fof(fraenkel_a_2_2_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_2_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r4_lattice3(X2,X4,X3))))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', fraenkel_a_2_2_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/bb8be2d1ee5f4360bbd1ffd9a2e40a27.E_input', dt_k15_lattice3)).
fof(c_0_12, plain, ![X11, X12, X13, X14, X15]:(((r3_lattice3(X11,X12,X13)|X12!=k16_lattice3(X11,X13)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v10_lattices(X11)|(~v4_lattice3(X11)|~l3_lattices(X11)))))&(~m1_subset_1(X14,u1_struct_0(X11))|(~r3_lattice3(X11,X14,X13)|r3_lattices(X11,X14,X12))|X12!=k16_lattice3(X11,X13)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v10_lattices(X11)|(~v4_lattice3(X11)|~l3_lattices(X11))))))&((m1_subset_1(esk5_3(X11,X12,X15),u1_struct_0(X11))|~r3_lattice3(X11,X12,X15)|X12=k16_lattice3(X11,X15)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v10_lattices(X11)|(~v4_lattice3(X11)|~l3_lattices(X11)))))&((r3_lattice3(X11,esk5_3(X11,X12,X15),X15)|~r3_lattice3(X11,X12,X15)|X12=k16_lattice3(X11,X15)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v10_lattices(X11)|(~v4_lattice3(X11)|~l3_lattices(X11)))))&(~r3_lattices(X11,esk5_3(X11,X12,X15),X12)|~r3_lattice3(X11,X12,X15)|X12=k16_lattice3(X11,X15)|~m1_subset_1(X12,u1_struct_0(X11))|(v2_struct_0(X11)|(~v10_lattices(X11)|(~v4_lattice3(X11)|~l3_lattices(X11)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])).
cnf(c_0_13, plain, (r3_lattices(X2,X1,X4)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_lattice3(X2,X1,X3)|X4!=k16_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_14, plain, ![X30, X31]:((X31=k15_lattice3(X30,a_2_3_lattice3(X30,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|(~v10_lattices(X30)|(~v4_lattice3(X30)|~l3_lattices(X30)))))&(X31=k16_lattice3(X30,a_2_4_lattice3(X30,X31))|~m1_subset_1(X31,u1_struct_0(X30))|(v2_struct_0(X30)|(~v10_lattices(X30)|(~v4_lattice3(X30)|~l3_lattices(X30)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t45_lattice3])])])])])).
cnf(c_0_15, plain, (r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|X2!=k16_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_16, plain, ![X22, X23, X24, X27, X28, X29]:(((m1_subset_1(esk7_3(X22,X23,X24),u1_struct_0(X23))|~r2_hidden(X22,a_2_5_lattice3(X23,X24))|(v2_struct_0(X23)|(~v10_lattices(X23)|(~v4_lattice3(X23)|~l3_lattices(X23)))))&((X22=esk7_3(X22,X23,X24)|~r2_hidden(X22,a_2_5_lattice3(X23,X24))|(v2_struct_0(X23)|(~v10_lattices(X23)|(~v4_lattice3(X23)|~l3_lattices(X23)))))&((m1_subset_1(esk8_3(X22,X23,X24),u1_struct_0(X23))|~r2_hidden(X22,a_2_5_lattice3(X23,X24))|(v2_struct_0(X23)|(~v10_lattices(X23)|(~v4_lattice3(X23)|~l3_lattices(X23)))))&((r3_lattices(X23,esk7_3(X22,X23,X24),esk8_3(X22,X23,X24))|~r2_hidden(X22,a_2_5_lattice3(X23,X24))|(v2_struct_0(X23)|(~v10_lattices(X23)|(~v4_lattice3(X23)|~l3_lattices(X23)))))&(r2_hidden(esk8_3(X22,X23,X24),X24)|~r2_hidden(X22,a_2_5_lattice3(X23,X24))|(v2_struct_0(X23)|(~v10_lattices(X23)|(~v4_lattice3(X23)|~l3_lattices(X23)))))))))&(~m1_subset_1(X28,u1_struct_0(X23))|(X22!=X28|(~m1_subset_1(X29,u1_struct_0(X23))|(~r3_lattices(X23,X28,X29)|~r2_hidden(X29,X27))))|r2_hidden(X22,a_2_5_lattice3(X23,X27))|(v2_struct_0(X23)|(~v10_lattices(X23)|(~v4_lattice3(X23)|~l3_lattices(X23)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_5_lattice3])])])])])])])).
fof(c_0_17, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X2)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>~((r3_lattices(X1,X4,X5)&r2_hidden(X5,X3)))))))=>r3_lattices(X1,k15_lattice3(X1,X2),k15_lattice3(X1,X3))))), inference(assume_negation,[status(cth)],[t48_lattice3])).
cnf(c_0_18, plain, (r3_lattices(X1,X2,k16_lattice3(X1,X3))|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(k16_lattice3(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (X1=k16_lattice3(X2,a_2_4_lattice3(X2,X1))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (r3_lattice3(X1,k16_lattice3(X1,X2),X2)|v2_struct_0(X1)|~m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(er,[status(thm)],[c_0_15])).
fof(c_0_21, plain, ![X52, X53, X54]:((r3_lattices(X52,X53,k15_lattice3(X52,X54))|~r2_hidden(X53,X54)|~m1_subset_1(X53,u1_struct_0(X52))|(v2_struct_0(X52)|(~v10_lattices(X52)|(~v4_lattice3(X52)|~l3_lattices(X52)))))&(r3_lattices(X52,k16_lattice3(X52,X54),X53)|~r2_hidden(X53,X54)|~m1_subset_1(X53,u1_struct_0(X52))|(v2_struct_0(X52)|(~v10_lattices(X52)|(~v4_lattice3(X52)|~l3_lattices(X52)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_lattice3])])])])])).
fof(c_0_22, plain, ![X57, X58]:((k15_lattice3(X57,X58)=k15_lattice3(X57,a_2_5_lattice3(X57,X58))|(v2_struct_0(X57)|(~v10_lattices(X57)|(~v4_lattice3(X57)|~l3_lattices(X57)))))&(k16_lattice3(X57,X58)=k16_lattice3(X57,a_2_6_lattice3(X57,X58))|(v2_struct_0(X57)|(~v10_lattices(X57)|(~v4_lattice3(X57)|~l3_lattices(X57)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_lattice3])])])])])).
cnf(c_0_23, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (X1=esk7_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (r2_hidden(X3,a_2_5_lattice3(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r3_lattices(X2,X1,X4)|~r2_hidden(X4,X5)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_26, negated_conjecture, ![X9]:((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v4_lattice3(esk1_0)&l3_lattices(esk1_0))))&(((m1_subset_1(esk4_1(X9),u1_struct_0(esk1_0))|~r2_hidden(X9,esk2_0)|~m1_subset_1(X9,u1_struct_0(esk1_0)))&((r3_lattices(esk1_0,X9,esk4_1(X9))|~r2_hidden(X9,esk2_0)|~m1_subset_1(X9,u1_struct_0(esk1_0)))&(r2_hidden(esk4_1(X9),esk3_0)|~r2_hidden(X9,esk2_0)|~m1_subset_1(X9,u1_struct_0(esk1_0)))))&~r3_lattices(esk1_0,k15_lattice3(esk1_0,esk2_0),k15_lattice3(esk1_0,esk3_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
cnf(c_0_27, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattice3(X1,X2,a_2_4_lattice3(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_28, plain, (r3_lattice3(X1,X2,a_2_4_lattice3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_20, c_0_19])).
fof(c_0_29, plain, ![X38, X39, X40, X41, X42]:((~r3_lattice3(X38,X39,X40)|(~m1_subset_1(X41,u1_struct_0(X38))|(~r2_hidden(X41,X40)|r1_lattices(X38,X39,X41)))|~m1_subset_1(X39,u1_struct_0(X38))|(v2_struct_0(X38)|~l3_lattices(X38)))&((m1_subset_1(esk10_3(X38,X39,X42),u1_struct_0(X38))|r3_lattice3(X38,X39,X42)|~m1_subset_1(X39,u1_struct_0(X38))|(v2_struct_0(X38)|~l3_lattices(X38)))&((r2_hidden(esk10_3(X38,X39,X42),X42)|r3_lattice3(X38,X39,X42)|~m1_subset_1(X39,u1_struct_0(X38))|(v2_struct_0(X38)|~l3_lattices(X38)))&(~r1_lattices(X38,X39,esk10_3(X38,X39,X42))|r3_lattice3(X38,X39,X42)|~m1_subset_1(X39,u1_struct_0(X38))|(v2_struct_0(X38)|~l3_lattices(X38)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
fof(c_0_30, plain, ![X17, X18, X19, X21]:(((m1_subset_1(esk6_3(X17,X18,X19),u1_struct_0(X18))|~r2_hidden(X17,a_2_4_lattice3(X18,X19))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18)))|~m1_subset_1(X19,u1_struct_0(X18))))&((X17=esk6_3(X17,X18,X19)|~r2_hidden(X17,a_2_4_lattice3(X18,X19))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18)))|~m1_subset_1(X19,u1_struct_0(X18))))&(r3_lattices(X18,X19,esk6_3(X17,X18,X19))|~r2_hidden(X17,a_2_4_lattice3(X18,X19))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18)))|~m1_subset_1(X19,u1_struct_0(X18))))))&(~m1_subset_1(X21,u1_struct_0(X18))|(X17!=X21|~r3_lattices(X18,X19,X21))|r2_hidden(X17,a_2_4_lattice3(X18,X19))|(v2_struct_0(X18)|(~v10_lattices(X18)|(~v4_lattice3(X18)|~l3_lattices(X18)))|~m1_subset_1(X19,u1_struct_0(X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_4_lattice3])])])])])])).
cnf(c_0_31, plain, (r3_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (k15_lattice3(X1,X2)=k15_lattice3(X1,a_2_5_lattice3(X1,X2))|v2_struct_0(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_5_lattice3(X2,X3))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_34, plain, (r2_hidden(X1,a_2_5_lattice3(X2,X3))|v2_struct_0(X2)|~r3_lattices(X2,X1,X4)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_35, negated_conjecture, (r3_lattices(esk1_0,X1,esk4_1(X1))|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, negated_conjecture, (v4_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk4_1(X1),u1_struct_0(esk1_0))|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_41, plain, (r3_lattices(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_42, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, plain, (r2_hidden(X3,a_2_4_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattices(X2,X4,X1)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~m1_subset_1(X4,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44, plain, (r3_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_5_lattice3(X1,X3))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_45, negated_conjecture, (r2_hidden(X1,a_2_5_lattice3(esk1_0,X2))|~r2_hidden(esk4_1(X1),X2)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37]), c_0_38])]), c_0_39]), c_0_40])).
cnf(c_0_46, plain, (r2_hidden(X1,a_2_5_lattice3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_34, c_0_41])).
fof(c_0_47, plain, ![X50, X51]:(v2_struct_0(X50)|(~v10_lattices(X50)|(~v4_lattice3(X50)|~l3_lattices(X50)))|k15_lattice3(X50,X51)=k16_lattice3(X50,a_2_2_lattice3(X50,X51))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t37_lattice3])])])])).
fof(c_0_48, plain, ![X32, X33, X34, X35, X36]:((~r4_lattice3(X32,X33,X34)|(~m1_subset_1(X35,u1_struct_0(X32))|(~r2_hidden(X35,X34)|r1_lattices(X32,X35,X33)))|~m1_subset_1(X33,u1_struct_0(X32))|(v2_struct_0(X32)|~l3_lattices(X32)))&((m1_subset_1(esk9_3(X32,X33,X36),u1_struct_0(X32))|r4_lattice3(X32,X33,X36)|~m1_subset_1(X33,u1_struct_0(X32))|(v2_struct_0(X32)|~l3_lattices(X32)))&((r2_hidden(esk9_3(X32,X33,X36),X36)|r4_lattice3(X32,X33,X36)|~m1_subset_1(X33,u1_struct_0(X32))|(v2_struct_0(X32)|~l3_lattices(X32)))&(~r1_lattices(X32,esk9_3(X32,X33,X36),X33)|r4_lattice3(X32,X33,X36)|~m1_subset_1(X33,u1_struct_0(X32))|(v2_struct_0(X32)|~l3_lattices(X32)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_49, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r2_hidden(X3,a_2_4_lattice3(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_42, c_0_28])).
cnf(c_0_50, plain, (r2_hidden(X1,a_2_4_lattice3(X2,X3))|v2_struct_0(X2)|~r3_lattices(X2,X3,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_43])).
cnf(c_0_51, plain, (r3_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_5_lattice3(X1,a_2_5_lattice3(X1,X3)))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_44, c_0_32])).
cnf(c_0_52, negated_conjecture, (r2_hidden(X1,a_2_5_lattice3(esk1_0,a_2_5_lattice3(X2,X3)))|v2_struct_0(X2)|~r2_hidden(esk4_1(X1),X3)|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk4_1(X1),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, plain, (r3_lattices(X1,k16_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_54, plain, (v2_struct_0(X1)|k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_55, plain, ![X44, X45, X46, X48, X49]:(((m1_subset_1(esk11_3(X44,X45,X46),u1_struct_0(X45))|~r2_hidden(X44,a_2_2_lattice3(X45,X46))|(v2_struct_0(X45)|(~v10_lattices(X45)|(~v4_lattice3(X45)|~l3_lattices(X45)))))&((X44=esk11_3(X44,X45,X46)|~r2_hidden(X44,a_2_2_lattice3(X45,X46))|(v2_struct_0(X45)|(~v10_lattices(X45)|(~v4_lattice3(X45)|~l3_lattices(X45)))))&(r4_lattice3(X45,esk11_3(X44,X45,X46),X46)|~r2_hidden(X44,a_2_2_lattice3(X45,X46))|(v2_struct_0(X45)|(~v10_lattices(X45)|(~v4_lattice3(X45)|~l3_lattices(X45)))))))&(~m1_subset_1(X49,u1_struct_0(X45))|(X44!=X49|~r4_lattice3(X45,X49,X48))|r2_hidden(X44,a_2_2_lattice3(X45,X48))|(v2_struct_0(X45)|(~v10_lattices(X45)|(~v4_lattice3(X45)|~l3_lattices(X45)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_2_lattice3])])])])])])])).
cnf(c_0_56, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,esk9_3(X1,X2,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58, plain, (m1_subset_1(esk9_3(X1,X2,X3),u1_struct_0(X1))|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, negated_conjecture, (r3_lattices(esk1_0,X1,k15_lattice3(esk1_0,X2))|~r2_hidden(esk4_1(X1),X2)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_36]), c_0_37]), c_0_38])]), c_0_39]), c_0_40])).
cnf(c_0_60, negated_conjecture, (r2_hidden(esk4_1(X1),esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_61, negated_conjecture, (~r3_lattices(esk1_0,k15_lattice3(esk1_0,esk2_0),k15_lattice3(esk1_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_62, plain, (r3_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r2_hidden(X3,a_2_2_lattice3(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_63, plain, (r2_hidden(X3,a_2_2_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r4_lattice3(X2,X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_64, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,esk9_3(X1,X2,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])).
cnf(c_0_65, negated_conjecture, (r3_lattices(esk1_0,X1,k15_lattice3(esk1_0,esk3_0))|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, negated_conjecture, (~r2_hidden(k15_lattice3(esk1_0,esk3_0),a_2_2_lattice3(esk1_0,esk2_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_67, plain, (r2_hidden(X1,a_2_2_lattice3(X2,X3))|v2_struct_0(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)|~v4_lattice3(X2)|~v10_lattices(X2)), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_68, negated_conjecture, (r4_lattice3(esk1_0,k15_lattice3(esk1_0,esk3_0),X1)|~r2_hidden(esk9_3(esk1_0,k15_lattice3(esk1_0,esk3_0),X1),esk2_0)|~m1_subset_1(esk9_3(esk1_0,k15_lattice3(esk1_0,esk3_0),X1),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_69, plain, (r2_hidden(esk9_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_70, negated_conjecture, (~r4_lattice3(esk1_0,k15_lattice3(esk1_0,esk3_0),esk2_0)|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_36]), c_0_37]), c_0_38])]), c_0_39])).
cnf(c_0_71, negated_conjecture, (~m1_subset_1(esk9_3(esk1_0,k15_lattice3(esk1_0,esk3_0),esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_36])]), c_0_39]), c_0_70])).
fof(c_0_72, plain, ![X55, X56]:(v2_struct_0(X55)|~l3_lattices(X55)|m1_subset_1(k15_lattice3(X55,X56),u1_struct_0(X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_73, negated_conjecture, (~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_58]), c_0_36])]), c_0_39]), c_0_70])).
cnf(c_0_74, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_75, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_36])]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 76
# Proof object clause steps            : 51
# Proof object formula steps           : 25
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 45
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 43
# Processed clauses                    : 225
# ...of these trivial                  : 0
# ...subsumed                          : 41
# ...remaining for further processing  : 184
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 0
# Generated clauses                    : 335
# ...of the previous two non-trivial   : 320
# Contextual simplify-reflections      : 19
# Paramodulations                      : 330
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 125
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 118
# Current number of unprocessed clauses: 166
# ...number of literals in the above   : 1574
# Current number of archived formulas  : 0
# Current number of archived clauses   : 54
# Clause-clause subsumption calls (NU) : 4632
# Rec. Clause-clause subsumption calls : 246
# Non-unit clause-clause subsumptions  : 71
# Unit Clause-clause subsumption calls : 17
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 17202

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.003 s
# Total time               : 0.023 s
# Maximum resident set size: 28720 pages
