# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.011 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t64_pre_topc, conjecture, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))))=>(X3=X4=>(v5_pre_topc(X3,X1,X2)<=>v5_pre_topc(X4,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t64_pre_topc)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', cc5_funct_2)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', cc1_relset_1)).
fof(d4_partfun1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>(v1_partfun1(X2,X1)<=>k1_relat_1(X2)=X1)), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', d4_partfun1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', l13_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t113_zfmisc_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t5_subset)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', fc1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', d3_tarski)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', d1_xboole_0)).
fof(d12_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v4_pre_topc(X4,X2)=>v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1))))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', d12_pre_topc)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t3_subset)).
fof(t13_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))=>(r1_tarski(k1_relat_1(X4),X2)=>m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t13_relset_1)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', d18_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', d10_xboole_0)).
fof(dt_k6_subset_1, axiom, ![X1, X2]:m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1)), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', dt_k6_subset_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', redefinition_k6_subset_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', redefinition_k1_relset_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t7_ordinal1)).
fof(t87_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k1_relat_1(k7_relat_1(X2,X1)),X1)), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t87_relat_1)).
fof(t209_relat_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>X2=k7_relat_1(X2,X1)), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t209_relat_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t3_boole)).
fof(t62_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(X2))))))=>(X3=X4=>(v5_pre_topc(X3,X1,X2)<=>v5_pre_topc(X4,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2))))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t62_pre_topc)).
fof(t38_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(k7_relset_1(X1,X2,X3,X1)=k2_relset_1(X1,X2,X3)&k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t38_relset_1)).
fof(rc1_funct_2, axiom, ![X1, X2]:?[X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&(v1_relat_1(X3)&(v4_relat_1(X3,X1)&(v5_relat_1(X3,X2)&(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', rc1_funct_2)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', cc4_relset_1)).
fof(cc2_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_xboole_0(X2)=>v4_pre_topc(X2,X1)))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', cc2_pre_topc)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t7_xboole_0)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t8_boole)).
fof(t29_relset_1, axiom, ![X1]:m1_subset_1(k6_relat_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t29_relset_1)).
fof(redefinition_k6_partfun1, axiom, ![X1]:k6_partfun1(X1)=k6_relat_1(X1), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', redefinition_k6_partfun1)).
fof(t63_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))))=>(X3=X4=>(v5_pre_topc(X3,X1,X2)<=>v5_pre_topc(X4,X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t63_pre_topc)).
fof(t81_relat_1, axiom, k6_relat_1(k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t81_relat_1)).
fof(dt_g1_pre_topc, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(v1_pre_topc(g1_pre_topc(X1,X2))&l1_pre_topc(g1_pre_topc(X1,X2)))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', dt_g1_pre_topc)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', dt_u1_pre_topc)).
fof(fc6_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))&v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', fc6_pre_topc)).
fof(d19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v5_relat_1(X2,X1)<=>r1_tarski(k2_relat_1(X2),X1))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', d19_relat_1)).
fof(t4_funct_2, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(k2_relat_1(X2),X1)=>(v1_funct_1(X2)&(v1_funct_2(X2,k1_relat_1(X2),X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X2),X1))))))), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t4_funct_2)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/bc83a612358e4b8da0bc3f39b4e39bd7.E_input', t71_relat_1)).
fof(c_0_40, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))))=>(X3=X4=>(v5_pre_topc(X3,X1,X2)<=>v5_pre_topc(X4,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))))))))), inference(assume_negation,[status(cth)],[t64_pre_topc])).
fof(c_0_41, plain, ![X69, X70, X71]:((v1_funct_1(X71)|(~v1_funct_1(X71)|~v1_funct_2(X71,X69,X70))|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(X69,X70)))|v1_xboole_0(X70))&(v1_partfun1(X71,X69)|(~v1_funct_1(X71)|~v1_funct_2(X71,X69,X70))|~m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(X69,X70)))|v1_xboole_0(X70))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
fof(c_0_42, negated_conjecture, ((v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0))&((v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))))&(esk3_0=esk4_0&((~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk4_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))&(v5_pre_topc(esk3_0,esk1_0,esk2_0)|v5_pre_topc(esk4_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_40])])])).
fof(c_0_43, plain, ![X45, X46, X47]:((v4_relat_1(X47,X45)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X45,X46))))&(v5_relat_1(X47,X46)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X45,X46))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_44, plain, ![X127, X128, X129]:(~m1_subset_1(X129,k1_zfmisc_1(k2_zfmisc_1(X127,X128)))|v1_relat_1(X129)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_45, plain, ![X99, X100]:((~v1_partfun1(X100,X99)|k1_relat_1(X100)=X99|(~v1_relat_1(X100)|~v4_relat_1(X100,X99)))&(k1_relat_1(X100)!=X99|v1_partfun1(X100,X99)|(~v1_relat_1(X100)|~v4_relat_1(X100,X99)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_partfun1])])])).
cnf(c_0_46, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_50, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_52, plain, ![X83]:(~v1_xboole_0(X83)|X83=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_53, plain, (k1_relat_1(X1)=X2|~v1_partfun1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, negated_conjecture, (v1_partfun1(esk3_0,u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48]), c_0_49])])).
cnf(c_0_55, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_50, c_0_48])).
cnf(c_0_56, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_51, c_0_48])).
fof(c_0_57, plain, ![X67, X68]:((k2_zfmisc_1(X67,X68)!=k1_xboole_0|(X67=k1_xboole_0|X68=k1_xboole_0))&((X67!=k1_xboole_0|k2_zfmisc_1(X67,X68)=k1_xboole_0)&(X68!=k1_xboole_0|k2_zfmisc_1(X67,X68)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_58, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_59, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56])])).
cnf(c_0_60, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_61, plain, ![X74, X75, X76]:(~r2_hidden(X74,X75)|(~m1_subset_1(X75,k1_zfmisc_1(X76))|~v1_xboole_0(X76))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_62, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|u1_struct_0(esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_63, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_64, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_65, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_62]), c_0_63])).
cnf(c_0_66, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_67, plain, ![X85, X86, X87, X88, X89]:((~r1_tarski(X85,X86)|(~r2_hidden(X87,X85)|r2_hidden(X87,X86)))&((r2_hidden(esk8_2(X88,X89),X88)|r1_tarski(X88,X89))&(~r2_hidden(esk8_2(X88,X89),X89)|r1_tarski(X88,X89)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_69, negated_conjecture, (esk3_0=esk4_0), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_70, plain, ![X19, X20, X21]:((~v1_xboole_0(X19)|~r2_hidden(X20,X19))&(r2_hidden(esk5_1(X21),X21)|v1_xboole_0(X21))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_71, plain, ![X62, X63, X64, X65]:((~v5_pre_topc(X64,X62,X63)|(~m1_subset_1(X65,k1_zfmisc_1(u1_struct_0(X63)))|(~v4_pre_topc(X65,X63)|v4_pre_topc(k8_relset_1(u1_struct_0(X62),u1_struct_0(X63),X64,X65),X62)))|(~v1_funct_1(X64)|(~v1_funct_2(X64,u1_struct_0(X62),u1_struct_0(X63))|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X62),u1_struct_0(X63))))))|~l1_pre_topc(X63)|~l1_pre_topc(X62))&((m1_subset_1(esk7_3(X62,X63,X64),k1_zfmisc_1(u1_struct_0(X63)))|v5_pre_topc(X64,X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,u1_struct_0(X62),u1_struct_0(X63))|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X62),u1_struct_0(X63))))))|~l1_pre_topc(X63)|~l1_pre_topc(X62))&((v4_pre_topc(esk7_3(X62,X63,X64),X63)|v5_pre_topc(X64,X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,u1_struct_0(X62),u1_struct_0(X63))|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X62),u1_struct_0(X63))))))|~l1_pre_topc(X63)|~l1_pre_topc(X62))&(~v4_pre_topc(k8_relset_1(u1_struct_0(X62),u1_struct_0(X63),X64,esk7_3(X62,X63,X64)),X62)|v5_pre_topc(X64,X62,X63)|(~v1_funct_1(X64)|(~v1_funct_2(X64,u1_struct_0(X62),u1_struct_0(X63))|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X62),u1_struct_0(X63))))))|~l1_pre_topc(X63)|~l1_pre_topc(X62))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_pre_topc])])])])])).
fof(c_0_72, plain, ![X79, X80]:((~m1_subset_1(X79,k1_zfmisc_1(X80))|r1_tarski(X79,X80))&(~r1_tarski(X79,X80)|m1_subset_1(X79,k1_zfmisc_1(X80)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_73, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66])])).
cnf(c_0_74, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_75, plain, ![X149, X150, X151, X152]:(~m1_subset_1(X152,k1_zfmisc_1(k2_zfmisc_1(X149,X151)))|(~r1_tarski(k1_relat_1(X152),X150)|m1_subset_1(X152,k1_zfmisc_1(k2_zfmisc_1(X150,X151))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_relset_1])])).
fof(c_0_76, plain, ![X136, X137]:((~v4_relat_1(X137,X136)|r1_tarski(k1_relat_1(X137),X136)|~v1_relat_1(X137))&(~r1_tarski(k1_relat_1(X137),X136)|v4_relat_1(X137,X136)|~v1_relat_1(X137))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
cnf(c_0_77, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))), inference(rw,[status(thm)],[c_0_68, c_0_69])).
fof(c_0_78, plain, ![X130, X131]:(((r1_tarski(X130,X131)|X130!=X131)&(r1_tarski(X131,X130)|X130!=X131))&(~r1_tarski(X130,X131)|~r1_tarski(X131,X130)|X130=X131)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_79, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_80, plain, (m1_subset_1(esk7_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_81, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_82, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_83, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_84, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|r1_tarski(esk3_0,X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_85, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(k1_relat_1(X1),X4)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_86, plain, (r1_tarski(k1_relat_1(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_87, negated_conjecture, (v4_relat_1(esk3_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))), inference(spm,[status(thm)],[c_0_50, c_0_77])).
fof(c_0_88, plain, ![X40, X41]:m1_subset_1(k6_subset_1(X40,X41),k1_zfmisc_1(X40)), inference(variable_rename,[status(thm)],[dt_k6_subset_1])).
fof(c_0_89, plain, ![X142, X143]:k6_subset_1(X142,X143)=k4_xboole_0(X142,X143), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_90, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_91, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_79, c_0_74])).
cnf(c_0_92, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_93, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|m1_subset_1(esk7_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_47]), c_0_48]), c_0_49]), c_0_81]), c_0_82])])).
fof(c_0_94, plain, ![X12, X13, X14]:(~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X12,X13)))|k1_relset_1(X12,X13,X14)=k1_relat_1(X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
fof(c_0_95, plain, ![X77, X78]:(~r2_hidden(X77,X78)|~r1_tarski(X78,X77)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_96, plain, ![X28, X29]:(~v1_relat_1(X29)|r1_tarski(k1_relat_1(k7_relat_1(X29,X28)),X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_relat_1])])).
fof(c_0_97, plain, ![X50, X51]:(~v1_relat_1(X51)|~v4_relat_1(X51,X50)|X51=k7_relat_1(X51,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t209_relat_1])])).
cnf(c_0_98, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_99, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(esk2_0))))|~r1_tarski(k1_relat_1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_85, c_0_48])).
cnf(c_0_100, negated_conjecture, (r1_tarski(k1_relat_1(esk3_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_56])])).
cnf(c_0_101, plain, (m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_102, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_103, plain, ![X35]:k4_xboole_0(X35,k1_xboole_0)=X35, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_104, plain, ![X123, X124, X125, X126]:((~v5_pre_topc(X125,X123,X124)|v5_pre_topc(X126,g1_pre_topc(u1_struct_0(X123),u1_pre_topc(X123)),X124)|X125!=X126|(~v1_funct_1(X126)|(~v1_funct_2(X126,u1_struct_0(g1_pre_topc(u1_struct_0(X123),u1_pre_topc(X123))),u1_struct_0(X124))|~m1_subset_1(X126,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X123),u1_pre_topc(X123))),u1_struct_0(X124))))))|(~v1_funct_1(X125)|(~v1_funct_2(X125,u1_struct_0(X123),u1_struct_0(X124))|~m1_subset_1(X125,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X123),u1_struct_0(X124))))))|(~v2_pre_topc(X124)|~l1_pre_topc(X124))|(~v2_pre_topc(X123)|~l1_pre_topc(X123)))&(~v5_pre_topc(X126,g1_pre_topc(u1_struct_0(X123),u1_pre_topc(X123)),X124)|v5_pre_topc(X125,X123,X124)|X125!=X126|(~v1_funct_1(X126)|(~v1_funct_2(X126,u1_struct_0(g1_pre_topc(u1_struct_0(X123),u1_pre_topc(X123))),u1_struct_0(X124))|~m1_subset_1(X126,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X123),u1_pre_topc(X123))),u1_struct_0(X124))))))|(~v1_funct_1(X125)|(~v1_funct_2(X125,u1_struct_0(X123),u1_struct_0(X124))|~m1_subset_1(X125,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X123),u1_struct_0(X124))))))|(~v2_pre_topc(X124)|~l1_pre_topc(X124))|(~v2_pre_topc(X123)|~l1_pre_topc(X123)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t62_pre_topc])])])])).
cnf(c_0_105, plain, (X1=X2|~r1_tarski(X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_106, negated_conjecture, (r1_tarski(esk7_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
fof(c_0_107, plain, ![X113, X114, X115]:((k7_relset_1(X113,X114,X115,X113)=k2_relset_1(X113,X114,X115)|~m1_subset_1(X115,k1_zfmisc_1(k2_zfmisc_1(X113,X114))))&(k8_relset_1(X113,X114,X115,X114)=k1_relset_1(X113,X114,X115)|~m1_subset_1(X115,k1_zfmisc_1(k2_zfmisc_1(X113,X114))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_relset_1])])])).
cnf(c_0_108, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_109, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_110, plain, (r1_tarski(k1_relat_1(k7_relat_1(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_111, plain, (X1=k7_relat_1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_112, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v4_relat_1(esk3_0,X1)), inference(spm,[status(thm)],[c_0_50, c_0_98])).
fof(c_0_113, plain, ![X120, X121]:(m1_subset_1(esk10_2(X120,X121),k1_zfmisc_1(k2_zfmisc_1(X120,X121)))&(v1_relat_1(esk10_2(X120,X121))&(v4_relat_1(esk10_2(X120,X121),X120)&(v5_relat_1(esk10_2(X120,X121),X121)&(v1_funct_1(esk10_2(X120,X121))&v1_funct_2(esk10_2(X120,X121),X120,X121)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_funct_2])])).
cnf(c_0_114, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
fof(c_0_115, plain, ![X32, X33, X34]:(~v1_xboole_0(X32)|(~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X33,X32)))|v1_xboole_0(X34))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_116, plain, (m1_subset_1(k4_xboole_0(X1,X2),k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_117, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_118, plain, (v5_pre_topc(X4,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X3)|~v5_pre_topc(X1,X2,X3)|X1!=X4|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))))|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_119, plain, (v5_pre_topc(X3,X1,X2)|~v4_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk7_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_120, negated_conjecture, (esk7_3(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk2_0)|v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_121, plain, (k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_122, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_108, c_0_48])).
fof(c_0_123, plain, ![X15, X16]:(~v2_pre_topc(X15)|~l1_pre_topc(X15)|(~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X15)))|(~v1_xboole_0(X16)|v4_pre_topc(X16,X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_pre_topc])])])).
cnf(c_0_124, plain, (~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(k7_relat_1(X1,X2)))), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_125, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|k7_relat_1(esk3_0,X1)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_56])])).
cnf(c_0_126, plain, (m1_subset_1(esk10_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_127, plain, ![X165]:(X165=k1_xboole_0|r2_hidden(esk11_1(X165),X165)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_128, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_92, c_0_114])).
cnf(c_0_129, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_130, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_131, plain, (v5_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X3)|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_118])).
cnf(c_0_132, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_133, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_134, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v4_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,esk7_3(esk1_0,esk2_0,esk3_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_47]), c_0_48]), c_0_49]), c_0_81]), c_0_82])])).
cnf(c_0_135, negated_conjecture, (esk7_3(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk2_0)|u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_120, c_0_59])).
cnf(c_0_136, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,u1_struct_0(esk2_0))=k1_relat_1(esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_48]), c_0_122])).
cnf(c_0_137, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_138, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_83, c_0_91])).
cnf(c_0_139, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_56])])).
cnf(c_0_140, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_141, plain, ![X57, X58]:(~v1_xboole_0(X57)|(X57=X58|~v1_xboole_0(X58))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
cnf(c_0_142, plain, (~r2_hidden(X1,esk10_2(X2,X3))|~v1_xboole_0(k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_64, c_0_126])).
cnf(c_0_143, plain, (X1=k1_xboole_0|r2_hidden(esk11_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_127])).
fof(c_0_144, plain, ![X167]:m1_subset_1(k6_relat_1(X167),k1_zfmisc_1(k2_zfmisc_1(X167,X167))), inference(variable_rename,[status(thm)],[t29_relset_1])).
fof(c_0_145, plain, ![X72]:k6_partfun1(X72)=k6_relat_1(X72), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
fof(c_0_146, plain, ![X138, X139, X140, X141]:((~v5_pre_topc(X140,X138,X139)|v5_pre_topc(X141,X138,g1_pre_topc(u1_struct_0(X139),u1_pre_topc(X139)))|X140!=X141|(~v1_funct_1(X141)|(~v1_funct_2(X141,u1_struct_0(X138),u1_struct_0(g1_pre_topc(u1_struct_0(X139),u1_pre_topc(X139))))|~m1_subset_1(X141,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X138),u1_struct_0(g1_pre_topc(u1_struct_0(X139),u1_pre_topc(X139))))))))|(~v1_funct_1(X140)|(~v1_funct_2(X140,u1_struct_0(X138),u1_struct_0(X139))|~m1_subset_1(X140,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X138),u1_struct_0(X139))))))|(~v2_pre_topc(X139)|~l1_pre_topc(X139))|(~v2_pre_topc(X138)|~l1_pre_topc(X138)))&(~v5_pre_topc(X141,X138,g1_pre_topc(u1_struct_0(X139),u1_pre_topc(X139)))|v5_pre_topc(X140,X138,X139)|X140!=X141|(~v1_funct_1(X141)|(~v1_funct_2(X141,u1_struct_0(X138),u1_struct_0(g1_pre_topc(u1_struct_0(X139),u1_pre_topc(X139))))|~m1_subset_1(X141,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X138),u1_struct_0(g1_pre_topc(u1_struct_0(X139),u1_pre_topc(X139))))))))|(~v1_funct_1(X140)|(~v1_funct_2(X140,u1_struct_0(X138),u1_struct_0(X139))|~m1_subset_1(X140,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X138),u1_struct_0(X139))))))|(~v2_pre_topc(X139)|~l1_pre_topc(X139))|(~v2_pre_topc(X138)|~l1_pre_topc(X138)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t63_pre_topc])])])])).
cnf(c_0_147, negated_conjecture, (k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))=esk3_0|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_105, c_0_128])).
cnf(c_0_148, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_149, negated_conjecture, (v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_47]), c_0_48]), c_0_49]), c_0_81]), c_0_82]), c_0_132]), c_0_133])])).
cnf(c_0_150, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v4_pre_topc(k1_relat_1(esk3_0),esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_136])).
cnf(c_0_151, plain, (v4_pre_topc(X1,X2)|~v1_xboole_0(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_152, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v1_xboole_0(k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_153, plain, (v1_funct_2(esk10_2(X1,X2),X1,X2)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_154, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_155, plain, (esk10_2(X1,X2)=k1_xboole_0|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_156, negated_conjecture, (v1_funct_2(esk4_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_157, plain, (m1_subset_1(k6_relat_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_158, plain, (k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_145])).
cnf(c_0_159, plain, (v5_pre_topc(X4,X2,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|~v5_pre_topc(X1,X2,X3)|X1!=X4|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))))|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_146])).
cnf(c_0_160, plain, (r1_tarski(esk10_2(X1,X2),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_92, c_0_126])).
cnf(c_0_161, negated_conjecture, (k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))=esk3_0|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_147, c_0_148])).
cnf(c_0_162, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_129, c_0_48])).
cnf(c_0_163, negated_conjecture, (v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_funct_2(esk3_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_149, c_0_114])])).
cnf(c_0_164, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_82]), c_0_133])]), c_0_152])).
cnf(c_0_165, plain, (v1_funct_2(X1,X2,X3)|~v1_xboole_0(esk10_2(X2,X3))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_153, c_0_154])).
cnf(c_0_166, plain, (esk10_2(X1,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_63]), c_0_66])])).
cnf(c_0_167, negated_conjecture, (~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk4_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_168, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(rw,[status(thm)],[c_0_156, c_0_69])).
cnf(c_0_169, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(rw,[status(thm)],[c_0_157, c_0_158])).
cnf(c_0_170, plain, (k6_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t81_relat_1])).
cnf(c_0_171, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_172, plain, (v5_pre_topc(X1,X2,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|~v5_pre_topc(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_159])).
cnf(c_0_173, plain, (v1_funct_1(esk10_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_174, plain, (k2_zfmisc_1(X1,X2)=esk10_2(X1,X2)|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_105, c_0_160])).
cnf(c_0_175, negated_conjecture, (k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))=esk3_0|u1_struct_0(esk1_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_161, c_0_59])).
cnf(c_0_176, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_162, c_0_59])).
cnf(c_0_177, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk2_0)|~v1_funct_2(esk3_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_62]), c_0_164])).
cnf(c_0_178, plain, (v1_funct_2(X1,X2,k1_xboole_0)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_66])])).
cnf(c_0_179, negated_conjecture, (~v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_167, c_0_69])).
fof(c_0_180, plain, ![X59, X60]:((v1_pre_topc(g1_pre_topc(X59,X60))|~m1_subset_1(X60,k1_zfmisc_1(k1_zfmisc_1(X59))))&(l1_pre_topc(g1_pre_topc(X59,X60))|~m1_subset_1(X60,k1_zfmisc_1(k1_zfmisc_1(X59))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_g1_pre_topc])])])).
fof(c_0_181, plain, ![X84]:(~l1_pre_topc(X84)|m1_subset_1(u1_pre_topc(X84),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X84))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_182, negated_conjecture, (v1_partfun1(esk3_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))|v1_xboole_0(u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_168]), c_0_77]), c_0_49])])).
cnf(c_0_183, plain, (~r2_hidden(X1,k6_partfun1(X2))|~v1_xboole_0(k2_zfmisc_1(X2,X2))), inference(spm,[status(thm)],[c_0_64, c_0_169])).
cnf(c_0_184, plain, (k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_170, c_0_158])).
cnf(c_0_185, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_171])).
cnf(c_0_186, plain, (v5_pre_topc(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v5_pre_topc(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),X1,X2)|~m1_subset_1(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))|~v1_funct_2(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),u1_struct_0(X1),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_153]), c_0_126]), c_0_173])])).
cnf(c_0_187, negated_conjecture, (esk10_2(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))=esk3_0|u1_struct_0(esk1_0)=k1_relat_1(esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175]), c_0_176])).
cnf(c_0_188, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178]), c_0_176])).
cnf(c_0_189, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|~v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(spm,[status(thm)],[c_0_179, c_0_164])).
cnf(c_0_190, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_180])).
cnf(c_0_191, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_192, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))))), inference(spm,[status(thm)],[c_0_92, c_0_77])).
cnf(c_0_193, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))=k1_relat_1(esk3_0)|v1_xboole_0(u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_182]), c_0_87]), c_0_56])])).
cnf(c_0_194, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184]), c_0_185]), c_0_66])])).
cnf(c_0_195, negated_conjecture, (~r2_hidden(X1,esk3_0)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_64, c_0_48])).
cnf(c_0_196, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_77]), c_0_168]), c_0_81]), c_0_132])]), c_0_188]), c_0_189])).
cnf(c_0_197, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_190, c_0_191])).
fof(c_0_198, plain, ![X98]:((v1_pre_topc(g1_pre_topc(u1_struct_0(X98),u1_pre_topc(X98)))|(~v2_pre_topc(X98)|~l1_pre_topc(X98)))&(v2_pre_topc(g1_pre_topc(u1_struct_0(X98),u1_pre_topc(X98)))|(~v2_pre_topc(X98)|~l1_pre_topc(X98)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_pre_topc])])])).
fof(c_0_199, plain, ![X110, X111]:((~v5_relat_1(X111,X110)|r1_tarski(k2_relat_1(X111),X110)|~v1_relat_1(X111))&(~r1_tarski(k2_relat_1(X111),X110)|v5_relat_1(X111,X110)|~v1_relat_1(X111))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_relat_1])])])).
cnf(c_0_200, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_201, negated_conjecture, (k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))=esk3_0|~r1_tarski(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))),esk3_0)), inference(spm,[status(thm)],[c_0_90, c_0_192])).
cnf(c_0_202, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))=k1_relat_1(esk3_0)|u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))=k1_xboole_0), inference(spm,[status(thm)],[c_0_58, c_0_193])).
cnf(c_0_203, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_194, c_0_74])).
cnf(c_0_204, negated_conjecture, (k1_xboole_0=esk3_0|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_195, c_0_143])).
cnf(c_0_205, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_197]), c_0_82])])).
cnf(c_0_206, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_198])).
fof(c_0_207, plain, ![X81, X82]:((v1_funct_1(X82)|~r1_tarski(k2_relat_1(X82),X81)|(~v1_relat_1(X82)|~v1_funct_1(X82)))&((v1_funct_2(X82,k1_relat_1(X82),X81)|~r1_tarski(k2_relat_1(X82),X81)|(~v1_relat_1(X82)|~v1_funct_1(X82)))&(m1_subset_1(X82,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X82),X81)))|~r1_tarski(k2_relat_1(X82),X81)|(~v1_relat_1(X82)|~v1_funct_1(X82))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_funct_2])])])).
cnf(c_0_208, plain, (r1_tarski(k2_relat_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_199])).
cnf(c_0_209, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_200, c_0_48])).
cnf(c_0_210, plain, (v5_pre_topc(X4,X2,X3)|~v5_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X3)|X4!=X1|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))))|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_211, plain, (v5_pre_topc(X4,X2,X3)|~v5_pre_topc(X1,X2,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|X4!=X1|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))))|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_pre_topc(X3)|~l1_pre_topc(X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_146])).
cnf(c_0_212, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))=k1_relat_1(esk3_0)|k1_xboole_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_202]), c_0_63]), c_0_63]), c_0_203])])).
cnf(c_0_213, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)|k1_xboole_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_62]), c_0_63]), c_0_66])])).
cnf(c_0_214, negated_conjecture, (u1_struct_0(esk1_0)=k1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_206]), c_0_82]), c_0_133])])).
cnf(c_0_215, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(spm,[status(thm)],[c_0_200, c_0_77])).
cnf(c_0_216, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),X2)))|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_217, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209]), c_0_56])])).
cnf(c_0_218, plain, (v1_funct_2(X1,k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_219, plain, (v5_pre_topc(X1,X2,X3)|~v5_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_210])).
cnf(c_0_220, plain, (v5_pre_topc(X1,X2,X3)|~v5_pre_topc(X1,X2,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(er,[status(thm)],[c_0_211])).
cnf(c_0_221, negated_conjecture, (u1_struct_0(g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)))=k1_relat_1(esk3_0)|k1_xboole_0=esk3_0), inference(spm,[status(thm)],[c_0_212, c_0_213])).
cnf(c_0_222, negated_conjecture, (l1_pre_topc(g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_214]), c_0_82])])).
cnf(c_0_223, negated_conjecture, (v2_pre_topc(g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_214]), c_0_82]), c_0_133])])).
cnf(c_0_224, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_215]), c_0_56])])).
cnf(c_0_225, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|v5_pre_topc(esk4_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_226, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),u1_struct_0(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_217]), c_0_56]), c_0_49])])).
cnf(c_0_227, negated_conjecture, (v1_funct_2(esk3_0,k1_relat_1(esk3_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_217]), c_0_56]), c_0_49])])).
cnf(c_0_228, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))))|~v1_funct_2(esk3_0,u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_47]), c_0_48]), c_0_49]), c_0_81]), c_0_82]), c_0_132]), c_0_133])])).
cnf(c_0_229, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(X1,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),X2)|~v5_pre_topc(X1,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),u1_struct_0(X2))))|~v1_funct_2(X1,k1_relat_1(esk3_0),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))|~v1_funct_2(X1,k1_relat_1(esk3_0),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221]), c_0_222]), c_0_223])])).
cnf(c_0_230, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_224]), c_0_56]), c_0_49])])).
cnf(c_0_231, negated_conjecture, (v1_funct_2(esk3_0,k1_relat_1(esk3_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_224]), c_0_56]), c_0_49])])).
cnf(c_0_232, negated_conjecture, (v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_225, c_0_69])).
cnf(c_0_233, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_212]), c_0_226]), c_0_227])])).
cnf(c_0_234, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_212]), c_0_226]), c_0_227])])).
cnf(c_0_235, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),esk2_0)|~v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_227]), c_0_230]), c_0_226]), c_0_231]), c_0_49]), c_0_81]), c_0_132])])).
cnf(c_0_236, negated_conjecture, (v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_232, c_0_214])).
cnf(c_0_237, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),esk2_0)|~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_233, c_0_214])).
fof(c_0_238, plain, ![X73]:(k1_relat_1(k6_relat_1(X73))=X73&k2_relat_1(k6_relat_1(X73))=X73), inference(variable_rename,[status(thm)],[t71_relat_1])).
cnf(c_0_239, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),esk2_0)), inference(spm,[status(thm)],[c_0_234, c_0_213])).
cnf(c_0_240, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_236]), c_0_237])).
cnf(c_0_241, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_238])).
cnf(c_0_242, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(X1,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v5_pre_topc(X1,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(esk3_0),u1_struct_0(X2))))|~v1_funct_2(X1,k1_relat_1(esk3_0),u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))|~v1_funct_2(X1,k1_relat_1(esk3_0),u1_struct_0(X2))|~v1_funct_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_221]), c_0_222]), c_0_223])])).
cnf(c_0_243, negated_conjecture, (~v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_179, c_0_214])).
cnf(c_0_244, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_239, c_0_240])).
cnf(c_0_245, plain, (k1_relat_1(k6_partfun1(X1))=X1), inference(rw,[status(thm)],[c_0_241, c_0_158])).
cnf(c_0_246, negated_conjecture, (k1_xboole_0=esk3_0|v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_227]), c_0_230]), c_0_226]), c_0_231]), c_0_49]), c_0_81]), c_0_132])])).
cnf(c_0_247, negated_conjecture, (k1_xboole_0=esk3_0|~v5_pre_topc(esk3_0,g1_pre_topc(k1_relat_1(esk3_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(spm,[status(thm)],[c_0_243, c_0_244])).
cnf(c_0_248, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(g1_pre_topc(X1,u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))|~v1_xboole_0(u1_struct_0(esk1_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_168, c_0_154])).
cnf(c_0_249, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_245, c_0_184])).
cnf(c_0_250, negated_conjecture, (k1_xboole_0=esk3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_240]), c_0_247])).
cnf(c_0_251, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))|~r1_tarski(k1_relat_1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_85, c_0_77])).
cnf(c_0_252, negated_conjecture, (r1_tarski(k1_relat_1(esk3_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_55]), c_0_56])])).
cnf(c_0_253, plain, (v5_pre_topc(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),X1,X2)|~v5_pre_topc(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~m1_subset_1(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(X2))))|~v1_funct_2(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_153]), c_0_126]), c_0_173])])).
cnf(c_0_254, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(g1_pre_topc(X1,u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))|~v1_xboole_0(k1_relat_1(esk3_0))|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_248, c_0_214])).
cnf(c_0_255, plain, (k1_relat_1(esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_249, c_0_250]), c_0_250])).
cnf(c_0_256, plain, (v1_xboole_0(esk3_0)), inference(rw,[status(thm)],[c_0_66, c_0_250])).
cnf(c_0_257, plain, (esk10_2(k1_xboole_0,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_185]), c_0_66])])).
cnf(c_0_258, negated_conjecture, (v5_pre_topc(esk3_0,g1_pre_topc(X1,u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_232, c_0_154])).
cnf(c_0_259, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_47]), c_0_48]), c_0_49]), c_0_81]), c_0_82]), c_0_132]), c_0_133])])).
cnf(c_0_260, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))), inference(spm,[status(thm)],[c_0_251, c_0_252])).
cnf(c_0_261, plain, (v5_pre_topc(X1,X2,X3)|~v1_xboole_0(esk10_2(u1_struct_0(X2),u1_struct_0(X3)))|~v1_xboole_0(X1)|~v5_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X3)|~v1_funct_2(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_154]), c_0_138])).
cnf(c_0_262, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(g1_pre_topc(X1,u1_pre_topc(esk1_0))),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_254, c_0_255]), c_0_256])])).
cnf(c_0_263, plain, (esk10_2(esk3_0,X1)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_257, c_0_250]), c_0_250])).
cnf(c_0_264, negated_conjecture, (v5_pre_topc(esk3_0,g1_pre_topc(X1,u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_xboole_0(k1_relat_1(esk3_0))|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_258, c_0_214])).
cnf(c_0_265, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_259, c_0_260])])).
cnf(c_0_266, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v5_pre_topc(esk3_0,g1_pre_topc(esk3_0,u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262]), c_0_214]), c_0_255]), c_0_263]), c_0_256]), c_0_256]), c_0_214]), c_0_255]), c_0_82]), c_0_133]), c_0_214]), c_0_255]), c_0_256])])).
cnf(c_0_267, negated_conjecture, (v5_pre_topc(esk3_0,g1_pre_topc(X1,u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_264, c_0_255]), c_0_256])])).
cnf(c_0_268, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_265, c_0_214]), c_0_231])])).
cnf(c_0_269, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_47]), c_0_48]), c_0_49]), c_0_81]), c_0_82]), c_0_132]), c_0_133])])).
cnf(c_0_270, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267]), c_0_256])]), c_0_268])).
cnf(c_0_271, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_269, c_0_260])])).
cnf(c_0_272, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_197]), c_0_81])])).
cnf(c_0_273, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)|~v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_271, c_0_214]), c_0_231])])).
cnf(c_0_274, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_206]), c_0_81]), c_0_132])])).
cnf(c_0_275, plain, (v5_pre_topc(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~v5_pre_topc(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),X1,X2)|~m1_subset_1(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(X2))))|~v1_funct_2(esk10_2(u1_struct_0(X1),u1_struct_0(X2)),u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_153]), c_0_126]), c_0_173])])).
cnf(c_0_276, negated_conjecture, (~v5_pre_topc(esk3_0,g1_pre_topc(esk3_0,u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(rw,[status(thm)],[c_0_243, c_0_255])).
cnf(c_0_277, negated_conjecture, (v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_273, c_0_274])])).
cnf(c_0_278, plain, (v5_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X3)|~v1_xboole_0(esk10_2(u1_struct_0(X2),u1_struct_0(X3)))|~v1_xboole_0(X1)|~v5_pre_topc(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),u1_struct_0(X3))|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_154]), c_0_138])).
cnf(c_0_279, negated_conjecture, (~v5_pre_topc(esk3_0,g1_pre_topc(esk3_0,u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_276, c_0_277])])).
cnf(c_0_280, negated_conjecture, (~l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_262]), c_0_214]), c_0_255]), c_0_214]), c_0_255]), c_0_263]), c_0_256]), c_0_256]), c_0_274]), c_0_82]), c_0_133]), c_0_214]), c_0_255]), c_0_256])]), c_0_279])).
cnf(c_0_281, negated_conjecture, (~v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_197]), c_0_81])])).
cnf(c_0_282, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_206]), c_0_81]), c_0_132])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 283
# Proof object clause steps            : 204
# Proof object formula steps           : 79
# Proof object conjectures             : 120
# Proof object clause conjectures      : 117
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 61
# Proof object initial formulas used   : 40
# Proof object generating inferences   : 111
# Proof object simplifying inferences  : 246
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 69
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 120
# Removed in clause preprocessing      : 13
# Initial clauses in saturation        : 107
# Processed clauses                    : 37803
# ...of these trivial                  : 1399
# ...subsumed                          : 30081
# ...remaining for further processing  : 6323
# Other redundant clauses eliminated   : 18
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 794
# Backward-rewritten                   : 2541
# Generated clauses                    : 388105
# ...of the previous two non-trivial   : 351089
# Contextual simplify-reflections      : 191
# Paramodulations                      : 388076
# Factorizations                       : 11
# Equation resolutions                 : 18
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2977
#    Positive orientable unit clauses  : 305
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 66
#    Non-unit-clauses                  : 2605
# Current number of unprocessed clauses: 280554
# ...number of literals in the above   : 1324660
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3345
# Clause-clause subsumption calls (NU) : 6295145
# Rec. Clause-clause subsumption calls : 2668033
# Non-unit clause-clause subsumptions  : 19984
# Unit Clause-clause subsumption calls : 151967
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1076
# BW rewrite match successes           : 122
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6912225

# -------------------------------------------------
# User time                : 5.307 s
# System time              : 0.080 s
# Total time               : 5.387 s
# Maximum resident set size: 209672 pages
