# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', t3_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', t65_xboole_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', t7_tarski)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', fc6_relat_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', t113_zfmisc_1)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', d8_relat_1)).
fof(t56_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(![X2, X3]:~(r2_hidden(k4_tarski(X2,X3),X1))=>X1=k1_xboole_0)), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', t56_relat_1)).
fof(t62_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0)), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', t62_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/bc9fb37bd7d041c4a3ed3dbac6a02eb1.E_input', dt_k5_relat_1)).
fof(c_0_9, plain, ![X35, X36, X38, X39, X40]:(((r2_hidden(esk9_2(X35,X36),X35)|r1_xboole_0(X35,X36))&(r2_hidden(esk9_2(X35,X36),X36)|r1_xboole_0(X35,X36)))&(~r2_hidden(X40,X38)|~r2_hidden(X40,X39)|~r1_xboole_0(X38,X39))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_10, plain, ![X28]:r1_xboole_0(X28,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_11, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, ![X8, X9, X11]:((r2_hidden(esk2_2(X8,X9),X9)|~r2_hidden(X8,X9))&(~r2_hidden(X11,X9)|~r2_hidden(X11,esk2_2(X8,X9))|~r2_hidden(X8,X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
fof(c_0_14, plain, ![X29, X30]:v1_relat_1(k2_zfmisc_1(X29,X30)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_15, plain, ![X33, X34]:((k2_zfmisc_1(X33,X34)!=k1_xboole_0|(X33=k1_xboole_0|X34=k1_xboole_0))&((X33!=k1_xboole_0|k2_zfmisc_1(X33,X34)=k1_xboole_0)&(X34!=k1_xboole_0|k2_zfmisc_1(X33,X34)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_16, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17, plain, (r2_hidden(esk2_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (~r2_hidden(esk2_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_21, plain, ![X12, X13, X14, X15, X16, X18, X19, X20, X23]:((((r2_hidden(k4_tarski(X15,esk3_5(X12,X13,X14,X15,X16)),X12)|~r2_hidden(k4_tarski(X15,X16),X14)|X14!=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12))&(r2_hidden(k4_tarski(esk3_5(X12,X13,X14,X15,X16),X16),X13)|~r2_hidden(k4_tarski(X15,X16),X14)|X14!=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12)))&(~r2_hidden(k4_tarski(X18,X20),X12)|~r2_hidden(k4_tarski(X20,X19),X13)|r2_hidden(k4_tarski(X18,X19),X14)|X14!=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12)))&((~r2_hidden(k4_tarski(esk4_3(X12,X13,X14),esk5_3(X12,X13,X14)),X14)|(~r2_hidden(k4_tarski(esk4_3(X12,X13,X14),X23),X12)|~r2_hidden(k4_tarski(X23,esk5_3(X12,X13,X14)),X13))|X14=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12))&((r2_hidden(k4_tarski(esk4_3(X12,X13,X14),esk6_3(X12,X13,X14)),X12)|r2_hidden(k4_tarski(esk4_3(X12,X13,X14),esk5_3(X12,X13,X14)),X14)|X14=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12))&(r2_hidden(k4_tarski(esk6_3(X12,X13,X14),esk5_3(X12,X13,X14)),X13)|r2_hidden(k4_tarski(esk4_3(X12,X13,X14),esk5_3(X12,X13,X14)),X14)|X14=k5_relat_1(X12,X13)|~v1_relat_1(X14)|~v1_relat_1(X13)|~v1_relat_1(X12))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
cnf(c_0_22, plain, (v1_relat_1(k1_xboole_0)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_20, c_0_17])).
cnf(c_0_24, plain, (r2_hidden(k4_tarski(esk3_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (v1_relat_1(k1_xboole_0)), inference(er,[status(thm)],[c_0_22])).
fof(c_0_26, plain, ![X25]:(~v1_relat_1(X25)|(r2_hidden(k4_tarski(esk7_1(X25),esk8_1(X25)),X25)|X25=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t56_relat_1])])])])).
cnf(c_0_27, plain, (X1!=k5_relat_1(X2,k1_xboole_0)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(esk7_1(X1),esk8_1(X1)),X1)|X1=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,esk3_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_30, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>(k5_relat_1(k1_xboole_0,X1)=k1_xboole_0&k5_relat_1(X1,k1_xboole_0)=k1_xboole_0))), inference(assume_negation,[status(cth)],[t62_relat_1])).
cnf(c_0_31, plain, (X1=k1_xboole_0|X1!=k5_relat_1(X2,k1_xboole_0)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_32, plain, ![X31, X32]:(~v1_relat_1(X31)|~v1_relat_1(X32)|v1_relat_1(k5_relat_1(X31,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_33, plain, (X1!=k5_relat_1(k1_xboole_0,X2)|~r2_hidden(k4_tarski(X3,X4),X1)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_29]), c_0_25])])).
fof(c_0_34, negated_conjecture, (v1_relat_1(esk1_0)&(k5_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0|k5_relat_1(esk1_0,k1_xboole_0)!=k1_xboole_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
cnf(c_0_35, plain, (k5_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(k5_relat_1(X1,k1_xboole_0))|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (X1=k1_xboole_0|X1!=k5_relat_1(k1_xboole_0,X2)|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_38, negated_conjecture, (k5_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0|k5_relat_1(esk1_0,k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_39, plain, (k5_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_25])])).
cnf(c_0_40, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (k5_relat_1(k1_xboole_0,X1)=k1_xboole_0|~v1_relat_1(k5_relat_1(k1_xboole_0,X1))|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_37])).
cnf(c_0_42, negated_conjecture, (k5_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
cnf(c_0_43, plain, (k5_relat_1(k1_xboole_0,X1)=k1_xboole_0|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_36]), c_0_25])])).
cnf(c_0_44, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_40])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 26
# Proof object formula steps           : 19
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 67
# ...of these trivial                  : 2
# ...subsumed                          : 5
# ...remaining for further processing  : 60
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 59
# ...of the previous two non-trivial   : 45
# Contextual simplify-reflections      : 0
# Paramodulations                      : 53
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 37
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 18
# ...number of literals in the above   : 116
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 716
# Rec. Clause-clause subsumption calls : 217
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2583

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28720 pages
