# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t156_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>![X4]:((v1_relat_1(X4)&v1_funct_1(X4))=>((X1=k2_relat_1(X2)&(k1_relat_1(X3)=X1&(k1_relat_1(X4)=X1&k5_relat_1(X2,X3)=k5_relat_1(X2,X4))))=>X3=X4)))), file('data/MPTP/proofs/bca0d3080373462a967e05d9186d4e4a.E_input', t156_funct_1)).
fof(t9_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((k1_relat_1(X1)=k1_relat_1(X2)&![X3]:(r2_hidden(X3,k1_relat_1(X1))=>k1_funct_1(X1,X3)=k1_funct_1(X2,X3)))=>X1=X2))), file('data/MPTP/proofs/bca0d3080373462a967e05d9186d4e4a.E_input', t9_funct_1)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/bca0d3080373462a967e05d9186d4e4a.E_input', d5_funct_1)).
fof(t23_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k1_relat_1(X2))=>k1_funct_1(k5_relat_1(X2,X3),X1)=k1_funct_1(X3,k1_funct_1(X2,X1))))), file('data/MPTP/proofs/bca0d3080373462a967e05d9186d4e4a.E_input', t23_funct_1)).
fof(c_0_4, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>![X4]:((v1_relat_1(X4)&v1_funct_1(X4))=>((X1=k2_relat_1(X2)&(k1_relat_1(X3)=X1&(k1_relat_1(X4)=X1&k5_relat_1(X2,X3)=k5_relat_1(X2,X4))))=>X3=X4))))), inference(assume_negation,[status(cth)],[t156_funct_1])).
fof(c_0_5, plain, ![X9, X10]:((r2_hidden(esk5_2(X9,X10),k1_relat_1(X9))|k1_relat_1(X9)!=k1_relat_1(X10)|X9=X10|(~v1_relat_1(X10)|~v1_funct_1(X10))|(~v1_relat_1(X9)|~v1_funct_1(X9)))&(k1_funct_1(X9,esk5_2(X9,X10))!=k1_funct_1(X10,esk5_2(X9,X10))|k1_relat_1(X9)!=k1_relat_1(X10)|X9=X10|(~v1_relat_1(X10)|~v1_funct_1(X10))|(~v1_relat_1(X9)|~v1_funct_1(X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_funct_1])])])])])).
fof(c_0_6, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((v1_relat_1(esk4_0)&v1_funct_1(esk4_0))&((esk1_0=k2_relat_1(esk2_0)&(k1_relat_1(esk3_0)=esk1_0&(k1_relat_1(esk4_0)=esk1_0&k5_relat_1(esk2_0,esk3_0)=k5_relat_1(esk2_0,esk4_0))))&esk3_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
cnf(c_0_7, plain, (r2_hidden(esk5_2(X1,X2),k1_relat_1(X1))|X1=X2|k1_relat_1(X1)!=k1_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, negated_conjecture, (k1_relat_1(esk4_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_11, plain, ![X15, X16, X17, X19, X20, X21, X23]:((((r2_hidden(esk6_3(X15,X16,X17),k1_relat_1(X15))|~r2_hidden(X17,X16)|X16!=k2_relat_1(X15)|(~v1_relat_1(X15)|~v1_funct_1(X15)))&(X17=k1_funct_1(X15,esk6_3(X15,X16,X17))|~r2_hidden(X17,X16)|X16!=k2_relat_1(X15)|(~v1_relat_1(X15)|~v1_funct_1(X15))))&(~r2_hidden(X20,k1_relat_1(X15))|X19!=k1_funct_1(X15,X20)|r2_hidden(X19,X16)|X16!=k2_relat_1(X15)|(~v1_relat_1(X15)|~v1_funct_1(X15))))&((~r2_hidden(esk7_2(X15,X21),X21)|(~r2_hidden(X23,k1_relat_1(X15))|esk7_2(X15,X21)!=k1_funct_1(X15,X23))|X21=k2_relat_1(X15)|(~v1_relat_1(X15)|~v1_funct_1(X15)))&((r2_hidden(esk8_2(X15,X21),k1_relat_1(X15))|r2_hidden(esk7_2(X15,X21),X21)|X21=k2_relat_1(X15)|(~v1_relat_1(X15)|~v1_funct_1(X15)))&(esk7_2(X15,X21)=k1_funct_1(X15,esk8_2(X15,X21))|r2_hidden(esk7_2(X15,X21),X21)|X21=k2_relat_1(X15)|(~v1_relat_1(X15)|~v1_funct_1(X15)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
cnf(c_0_12, negated_conjecture, (X1=esk4_0|r2_hidden(esk5_2(X1,esk4_0),k1_relat_1(X1))|k1_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_13, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_14, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, negated_conjecture, (esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17, plain, (r2_hidden(esk6_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (r2_hidden(esk5_2(esk3_0,esk4_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15])]), c_0_16])).
cnf(c_0_19, plain, (X1=k1_funct_1(X2,esk6_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_20, plain, ![X12, X13, X14]:(~v1_relat_1(X13)|~v1_funct_1(X13)|(~v1_relat_1(X14)|~v1_funct_1(X14)|(~r2_hidden(X12,k1_relat_1(X13))|k1_funct_1(k5_relat_1(X13,X14),X12)=k1_funct_1(X14,k1_funct_1(X13,X12))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t23_funct_1])])])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk6_3(X1,esk1_0,esk5_2(esk3_0,esk4_0)),k1_relat_1(X1))|k2_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22, negated_conjecture, (esk1_0=k2_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_23, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25, negated_conjecture, (k1_funct_1(X1,esk6_3(X1,esk1_0,esk5_2(esk3_0,esk4_0)))=esk5_2(esk3_0,esk4_0)|k2_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_19, c_0_18])).
cnf(c_0_26, plain, (k1_funct_1(k5_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~r2_hidden(X3,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, negated_conjecture, (r2_hidden(esk6_3(esk2_0,esk1_0,esk5_2(esk3_0,esk4_0)),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_28, negated_conjecture, (k1_funct_1(esk2_0,esk6_3(esk2_0,esk1_0,esk5_2(esk3_0,esk4_0)))=esk5_2(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_29, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,X1),esk6_3(esk2_0,esk1_0,esk5_2(esk3_0,esk4_0)))=k1_funct_1(X1,esk5_2(esk3_0,esk4_0))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_23]), c_0_24])]), c_0_28])).
cnf(c_0_30, negated_conjecture, (k5_relat_1(esk2_0,esk3_0)=k5_relat_1(esk2_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_31, plain, (X1=X2|k1_funct_1(X1,esk5_2(X1,X2))!=k1_funct_1(X2,esk5_2(X1,X2))|k1_relat_1(X1)!=k1_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_32, negated_conjecture, (k1_funct_1(k5_relat_1(esk2_0,esk3_0),esk6_3(esk2_0,esk1_0,esk5_2(esk3_0,esk4_0)))=k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_8]), c_0_30]), c_0_10])])).
cnf(c_0_33, negated_conjecture, (X1=esk4_0|k1_funct_1(X1,esk5_2(X1,esk4_0))!=k1_funct_1(esk4_0,esk5_2(X1,esk4_0))|k1_relat_1(X1)!=esk1_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_34, negated_conjecture, (k1_funct_1(esk4_0,esk5_2(esk3_0,esk4_0))=k1_funct_1(esk3_0,esk5_2(esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_14]), c_0_15])]), c_0_32])).
cnf(c_0_35, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_13]), c_0_14]), c_0_15])]), c_0_16]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 27
# Proof object formula steps           : 9
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 145
# ...of these trivial                  : 0
# ...subsumed                          : 6
# ...remaining for further processing  : 139
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 24
# Generated clauses                    : 189
# ...of the previous two non-trivial   : 204
# Contextual simplify-reflections      : 0
# Paramodulations                      : 155
# Factorizations                       : 0
# Equation resolutions                 : 34
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 95
#    Positive orientable unit clauses  : 38
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 56
# Current number of unprocessed clauses: 80
# ...number of literals in the above   : 323
# Current number of archived formulas  : 0
# Current number of archived clauses   : 44
# Clause-clause subsumption calls (NU) : 569
# Rec. Clause-clause subsumption calls : 383
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 131
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6090

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 28748 pages
