# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d11_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(v1_relat_1(X3)=>(X3=k7_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X4,X2)&r2_hidden(k4_tarski(X4,X5),X1)))))), file('data/MPTP/proofs/bdfec7b036db4455bdbd27fbeb54eceb.E_input', d11_relat_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/bdfec7b036db4455bdbd27fbeb54eceb.E_input', d4_relat_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/bdfec7b036db4455bdbd27fbeb54eceb.E_input', t8_funct_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/bdfec7b036db4455bdbd27fbeb54eceb.E_input', d5_relat_1)).
fof(t72_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,X2)=>k1_funct_1(k7_relat_1(X3,X2),X1)=k1_funct_1(X3,X1))), file('data/MPTP/proofs/bdfec7b036db4455bdbd27fbeb54eceb.E_input', t72_funct_1)).
fof(fc8_funct_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k7_relat_1(X1,X2))&v1_funct_1(k7_relat_1(X1,X2)))), file('data/MPTP/proofs/bdfec7b036db4455bdbd27fbeb54eceb.E_input', fc8_funct_1)).
fof(t73_funct_1, conjecture, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X1,X2))=>r2_hidden(k1_funct_1(X3,X1),k2_relat_1(k7_relat_1(X3,X2))))), file('data/MPTP/proofs/bdfec7b036db4455bdbd27fbeb54eceb.E_input', t73_funct_1)).
fof(c_0_7, plain, ![X31, X32, X33, X34, X35, X36, X37]:((((r2_hidden(X34,X32)|~r2_hidden(k4_tarski(X34,X35),X33)|X33!=k7_relat_1(X31,X32)|~v1_relat_1(X33)|~v1_relat_1(X31))&(r2_hidden(k4_tarski(X34,X35),X31)|~r2_hidden(k4_tarski(X34,X35),X33)|X33!=k7_relat_1(X31,X32)|~v1_relat_1(X33)|~v1_relat_1(X31)))&(~r2_hidden(X36,X32)|~r2_hidden(k4_tarski(X36,X37),X31)|r2_hidden(k4_tarski(X36,X37),X33)|X33!=k7_relat_1(X31,X32)|~v1_relat_1(X33)|~v1_relat_1(X31)))&((~r2_hidden(k4_tarski(esk7_3(X31,X32,X33),esk8_3(X31,X32,X33)),X33)|(~r2_hidden(esk7_3(X31,X32,X33),X32)|~r2_hidden(k4_tarski(esk7_3(X31,X32,X33),esk8_3(X31,X32,X33)),X31))|X33=k7_relat_1(X31,X32)|~v1_relat_1(X33)|~v1_relat_1(X31))&((r2_hidden(esk7_3(X31,X32,X33),X32)|r2_hidden(k4_tarski(esk7_3(X31,X32,X33),esk8_3(X31,X32,X33)),X33)|X33=k7_relat_1(X31,X32)|~v1_relat_1(X33)|~v1_relat_1(X31))&(r2_hidden(k4_tarski(esk7_3(X31,X32,X33),esk8_3(X31,X32,X33)),X31)|r2_hidden(k4_tarski(esk7_3(X31,X32,X33),esk8_3(X31,X32,X33)),X33)|X33=k7_relat_1(X31,X32)|~v1_relat_1(X33)|~v1_relat_1(X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_relat_1])])])])])])).
fof(c_0_8, plain, ![X20, X21, X22, X24, X25, X26, X27, X29]:(((~r2_hidden(X22,X21)|r2_hidden(k4_tarski(X22,esk4_3(X20,X21,X22)),X20)|X21!=k1_relat_1(X20))&(~r2_hidden(k4_tarski(X24,X25),X20)|r2_hidden(X24,X21)|X21!=k1_relat_1(X20)))&((~r2_hidden(esk5_2(X26,X27),X27)|~r2_hidden(k4_tarski(esk5_2(X26,X27),X29),X26)|X27=k1_relat_1(X26))&(r2_hidden(esk5_2(X26,X27),X27)|r2_hidden(k4_tarski(esk5_2(X26,X27),esk6_2(X26,X27)),X26)|X27=k1_relat_1(X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_9, plain, ![X17, X18, X19]:(((r2_hidden(X17,k1_relat_1(X19))|~r2_hidden(k4_tarski(X17,X18),X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))&(X18=k1_funct_1(X19,X17)|~r2_hidden(k4_tarski(X17,X18),X19)|(~v1_relat_1(X19)|~v1_funct_1(X19))))&(~r2_hidden(X17,k1_relat_1(X19))|X18!=k1_funct_1(X19,X17)|r2_hidden(k4_tarski(X17,X18),X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_10, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X4!=k7_relat_1(X5,X2)|~v1_relat_1(X4)|~v1_relat_1(X5)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r2_hidden(k4_tarski(X1,esk4_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X6, X7, X8, X10, X11, X12, X13, X15]:(((~r2_hidden(X8,X7)|r2_hidden(k4_tarski(esk1_3(X6,X7,X8),X8),X6)|X7!=k2_relat_1(X6))&(~r2_hidden(k4_tarski(X11,X10),X6)|r2_hidden(X10,X7)|X7!=k2_relat_1(X6)))&((~r2_hidden(esk2_2(X12,X13),X13)|~r2_hidden(k4_tarski(X15,esk2_2(X12,X13)),X12)|X13=k2_relat_1(X12))&(r2_hidden(esk2_2(X12,X13),X13)|r2_hidden(k4_tarski(esk3_2(X12,X13),esk2_2(X12,X13)),X12)|X13=k2_relat_1(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_13, plain, (r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, plain, ![X43, X44, X45]:(~v1_relat_1(X45)|~v1_funct_1(X45)|(~r2_hidden(X43,X44)|k1_funct_1(k7_relat_1(X45,X44),X43)=k1_funct_1(X45,X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t72_funct_1])])).
cnf(c_0_15, plain, (r2_hidden(X1,X2)|~v1_relat_1(k7_relat_1(X3,X2))|~v1_relat_1(X3)|~r2_hidden(k4_tarski(X1,X4),k7_relat_1(X3,X2))), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(X1,esk4_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_11])).
fof(c_0_17, plain, ![X46, X47]:((v1_relat_1(k7_relat_1(X46,X47))|(~v1_relat_1(X46)|~v1_funct_1(X46)))&(v1_funct_1(k7_relat_1(X46,X47))|(~v1_relat_1(X46)|~v1_funct_1(X46)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_funct_1])])])).
fof(c_0_18, negated_conjecture, ~(![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X1,X2))=>r2_hidden(k1_funct_1(X3,X1),k2_relat_1(k7_relat_1(X3,X2)))))), inference(assume_negation,[status(cth)],[t73_funct_1])).
cnf(c_0_19, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (k1_funct_1(k7_relat_1(X1,X3),X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r2_hidden(X1,X2)|~v1_relat_1(k7_relat_1(X3,X2))|~v1_relat_1(X3)|~r2_hidden(X1,k1_relat_1(k7_relat_1(X3,X2)))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (v1_funct_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_26, plain, (r2_hidden(k4_tarski(X1,X3),X5)|~r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X5!=k7_relat_1(X4,X2)|~v1_relat_1(X5)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_27, negated_conjecture, ((v1_relat_1(esk11_0)&v1_funct_1(esk11_0))&((r2_hidden(esk9_0,k1_relat_1(esk11_0))&r2_hidden(esk9_0,esk10_0))&~r2_hidden(k1_funct_1(esk11_0,esk9_0),k2_relat_1(k7_relat_1(esk11_0,esk10_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
cnf(c_0_28, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),k7_relat_1(X2,X3))|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(X1,k1_relat_1(k7_relat_1(X2,X3)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23]), c_0_24])).
cnf(c_0_30, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (r2_hidden(k4_tarski(X1,X2),k7_relat_1(X3,X4))|~v1_relat_1(k7_relat_1(X3,X4))|~v1_relat_1(X3)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_32, negated_conjecture, (~r2_hidden(k1_funct_1(esk11_0,esk9_0),k2_relat_1(k7_relat_1(esk11_0,esk10_0)))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(k7_relat_1(X1,X3)))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(k7_relat_1(X1,X3)))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, negated_conjecture, (v1_funct_1(esk11_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (v1_relat_1(esk11_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (r2_hidden(X1,k1_relat_1(k7_relat_1(X2,X3)))|~v1_relat_1(k7_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, negated_conjecture, (~r2_hidden(esk9_0,k1_relat_1(k7_relat_1(esk11_0,esk10_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35])])).
cnf(c_0_38, plain, (r2_hidden(X1,k1_relat_1(k7_relat_1(X2,X3)))|~v1_relat_1(k7_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_36, c_0_16])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk9_0,k1_relat_1(esk11_0))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_40, negated_conjecture, (r2_hidden(esk9_0,esk10_0)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_41, negated_conjecture, (~v1_relat_1(k7_relat_1(esk11_0,esk10_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_35]), c_0_39]), c_0_40])])).
cnf(c_0_42, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_23]), c_0_34]), c_0_35])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 28
# Proof object formula steps           : 15
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 19
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 81
# ...of these trivial                  : 0
# ...subsumed                          : 7
# ...remaining for further processing  : 74
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 123
# ...of the previous two non-trivial   : 111
# Contextual simplify-reflections      : 4
# Paramodulations                      : 115
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 34
# Current number of unprocessed clauses: 79
# ...number of literals in the above   : 440
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 670
# Rec. Clause-clause subsumption calls : 165
# Non-unit clause-clause subsumptions  : 12
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4753

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 28760 pages
