# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t7_ordinal1)).
fof(l49_zfmisc_1, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', l49_zfmisc_1)).
fof(d4_tarski, axiom, ![X1, X2]:(X2=k3_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X3,X4)&r2_hidden(X4,X1)))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', d4_tarski)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t69_enumset1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', d3_xboole_0)).
fof(d2_ordinal1, axiom, ![X1]:(v1_ordinal1(X1)<=>![X2]:(r2_hidden(X2,X1)=>r1_tarski(X2,X1))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', d2_ordinal1)).
fof(cc1_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>(v1_ordinal1(X1)&v2_ordinal1(X1))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', cc1_ordinal1)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t10_ordinal1)).
fof(t41_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1))))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t41_ordinal1)).
fof(t31_zfmisc_1, axiom, ![X1]:k3_tarski(k1_tarski(X1))=X1, file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t31_zfmisc_1)).
fof(t94_zfmisc_1, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)=>r1_tarski(X3,X2))=>r1_tarski(k3_tarski(X1),X2)), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t94_zfmisc_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t12_xboole_1)).
fof(t24_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1))))))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t24_ordinal1)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t29_ordinal1)).
fof(d6_ordinal1, axiom, ![X1]:(v4_ordinal1(X1)<=>X1=k3_tarski(X1)), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', d6_ordinal1)).
fof(t96_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_xboole_0(X1,X2))=k2_xboole_0(k3_tarski(X1),k3_tarski(X2)), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t96_zfmisc_1)).
fof(t30_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k3_tarski(X1))), file('data/MPTP/proofs/be623f80e72043f6b716162925cb896f.E_input', t30_ordinal1)).
fof(c_0_18, plain, ![X5, X6]:(~r2_hidden(X5,X6)|~r1_tarski(X6,X5)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_19, plain, ![X7, X8]:(~r2_hidden(X7,X8)|r1_tarski(X7,k3_tarski(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l49_zfmisc_1])])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (r1_tarski(X1,k3_tarski(X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X9, X10, X11, X13, X14, X15, X16, X18]:((((r2_hidden(X11,esk1_3(X9,X10,X11))|~r2_hidden(X11,X10)|X10!=k3_tarski(X9))&(r2_hidden(esk1_3(X9,X10,X11),X9)|~r2_hidden(X11,X10)|X10!=k3_tarski(X9)))&(~r2_hidden(X13,X14)|~r2_hidden(X14,X9)|r2_hidden(X13,X10)|X10!=k3_tarski(X9)))&((~r2_hidden(esk2_2(X15,X16),X16)|(~r2_hidden(esk2_2(X15,X16),X18)|~r2_hidden(X18,X15))|X16=k3_tarski(X15))&((r2_hidden(esk2_2(X15,X16),esk3_2(X15,X16))|r2_hidden(esk2_2(X15,X16),X16)|X16=k3_tarski(X15))&(r2_hidden(esk3_2(X15,X16),X15)|r2_hidden(esk2_2(X15,X16),X16)|X16=k3_tarski(X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_tarski])])])])])])).
fof(c_0_23, plain, ![X20]:k1_ordinal1(X20)=k2_xboole_0(X20,k1_tarski(X20)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_24, plain, ![X21]:k2_tarski(X21,X21)=k1_tarski(X21), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_25, plain, (~r2_hidden(k3_tarski(X1),X2)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, plain, (r2_hidden(X1,esk1_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k3_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_27, plain, ![X22, X23, X24, X25, X26, X27, X28, X29]:(((~r2_hidden(X25,X24)|(r2_hidden(X25,X22)|r2_hidden(X25,X23))|X24!=k2_xboole_0(X22,X23))&((~r2_hidden(X26,X22)|r2_hidden(X26,X24)|X24!=k2_xboole_0(X22,X23))&(~r2_hidden(X26,X23)|r2_hidden(X26,X24)|X24!=k2_xboole_0(X22,X23))))&(((~r2_hidden(esk4_3(X27,X28,X29),X27)|~r2_hidden(esk4_3(X27,X28,X29),X29)|X29=k2_xboole_0(X27,X28))&(~r2_hidden(esk4_3(X27,X28,X29),X28)|~r2_hidden(esk4_3(X27,X28,X29),X29)|X29=k2_xboole_0(X27,X28)))&(r2_hidden(esk4_3(X27,X28,X29),X29)|(r2_hidden(esk4_3(X27,X28,X29),X27)|r2_hidden(esk4_3(X27,X28,X29),X28))|X29=k2_xboole_0(X27,X28)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
fof(c_0_28, plain, ![X31, X32, X33]:((~v1_ordinal1(X31)|(~r2_hidden(X32,X31)|r1_tarski(X32,X31)))&((r2_hidden(esk5_1(X33),X33)|v1_ordinal1(X33))&(~r1_tarski(esk5_1(X33),X33)|v1_ordinal1(X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_ordinal1])])])])])])).
fof(c_0_29, plain, ![X35]:((v1_ordinal1(X35)|~v3_ordinal1(X35))&(v2_ordinal1(X35)|~v3_ordinal1(X35))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_ordinal1])])])).
fof(c_0_30, plain, ![X36]:r2_hidden(X36,k1_ordinal1(X36)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
cnf(c_0_31, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_33, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>(v4_ordinal1(X1)<=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X2,X1)=>r2_hidden(k1_ordinal1(X2),X1)))))), inference(assume_negation,[status(cth)],[t41_ordinal1])).
cnf(c_0_34, plain, (X1!=k3_tarski(X2)|~r2_hidden(esk1_3(X2,X1,k3_tarski(X3)),X3)|~r2_hidden(k3_tarski(X3),X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_35, plain, (r2_hidden(esk1_3(X1,X2,X3),X1)|~r2_hidden(X3,X2)|X2!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_36, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_37, plain, ![X40]:k3_tarski(k1_tarski(X40))=X40, inference(variable_rename,[status(thm)],[t31_zfmisc_1])).
fof(c_0_38, plain, ![X41, X42]:((r2_hidden(esk8_2(X41,X42),X41)|r1_tarski(k3_tarski(X41),X42))&(~r1_tarski(esk8_2(X41,X42),X42)|r1_tarski(k3_tarski(X41),X42))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_zfmisc_1])])])])).
cnf(c_0_39, plain, (r1_tarski(X2,X1)|~v1_ordinal1(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_40, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_43, negated_conjecture, ![X39]:(v3_ordinal1(esk6_0)&(((v3_ordinal1(esk7_0)|~v4_ordinal1(esk6_0))&((r2_hidden(esk7_0,esk6_0)|~v4_ordinal1(esk6_0))&(~r2_hidden(k1_ordinal1(esk7_0),esk6_0)|~v4_ordinal1(esk6_0))))&(v4_ordinal1(esk6_0)|(~v3_ordinal1(X39)|(~r2_hidden(X39,esk6_0)|r2_hidden(k1_ordinal1(X39),esk6_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])])])).
cnf(c_0_44, plain, (X1!=k3_tarski(X2)|~r2_hidden(k3_tarski(X2),X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (k3_tarski(k1_tarski(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_47, plain, ![X44, X45]:(~r1_tarski(X44,X45)|k2_xboole_0(X44,X45)=X45), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_48, plain, ![X46, X47]:(~v3_ordinal1(X46)|(~v3_ordinal1(X47)|(r2_hidden(X46,X47)|(X46=X47|r2_hidden(X47,X46))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])).
fof(c_0_49, plain, ![X48]:(~v3_ordinal1(X48)|v3_ordinal1(k1_ordinal1(X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
cnf(c_0_50, plain, (r1_tarski(k3_tarski(X1),X2)|~r1_tarski(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51, plain, (r1_tarski(X1,X2)|~v3_ordinal1(X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_52, plain, (r2_hidden(X1,k2_xboole_0(X1,k2_tarski(X1,X1)))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_53, negated_conjecture, (v4_ordinal1(esk6_0)|r2_hidden(k1_ordinal1(X1),esk6_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, plain, (k2_xboole_0(X1,X2)!=k3_tarski(X3)|~r2_hidden(k3_tarski(X3),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_55, plain, (k3_tarski(k2_tarski(X1,X1))=X1), inference(rw,[status(thm)],[c_0_46, c_0_32])).
cnf(c_0_56, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, plain, (r2_hidden(X1,X2)|X1=X2|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, negated_conjecture, (v3_ordinal1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_59, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (r1_tarski(k3_tarski(X1),X2)|~v3_ordinal1(X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(k3_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_62, plain, ![X49]:((~v4_ordinal1(X49)|X49=k3_tarski(X49))&(X49!=k3_tarski(X49)|v4_ordinal1(X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_ordinal1])])).
cnf(c_0_63, plain, (~r2_hidden(k2_xboole_0(k3_tarski(X1),k2_tarski(k3_tarski(X1),k3_tarski(X1))),X1)), inference(spm,[status(thm)],[c_0_25, c_0_52])).
cnf(c_0_64, negated_conjecture, (v4_ordinal1(esk6_0)|r2_hidden(k2_xboole_0(X1,k2_tarski(X1,X1)),esk6_0)|~v3_ordinal1(X1)|~r2_hidden(X1,esk6_0)), inference(rw,[status(thm)],[c_0_53, c_0_42])).
cnf(c_0_65, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_66, plain, (k2_xboole_0(X1,X2)!=X3|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_67, plain, (k2_xboole_0(X1,k3_tarski(X2))=k3_tarski(X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_56, c_0_21])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(k1_ordinal1(esk7_0),esk6_0)|~v4_ordinal1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_69, negated_conjecture, (X1=esk6_0|r2_hidden(X1,esk6_0)|r2_hidden(esk6_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_70, plain, (v3_ordinal1(k2_xboole_0(X1,k2_tarski(X1,X1)))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[c_0_59, c_0_42])).
fof(c_0_71, plain, ![X50, X51]:k3_tarski(k2_xboole_0(X50,X51))=k2_xboole_0(k3_tarski(X50),k3_tarski(X51)), inference(variable_rename,[status(thm)],[t96_zfmisc_1])).
cnf(c_0_72, plain, (r1_tarski(k3_tarski(X1),X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
fof(c_0_73, plain, ![X52]:(~v3_ordinal1(X52)|v3_ordinal1(k3_tarski(X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_ordinal1])])).
cnf(c_0_74, plain, (X1=k3_tarski(X1)|~v4_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_75, negated_conjecture, (v4_ordinal1(esk6_0)|~v3_ordinal1(k3_tarski(esk6_0))|~r2_hidden(k3_tarski(esk6_0),esk6_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_76, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_65])).
cnf(c_0_77, plain, (k3_tarski(X1)!=X2|~r2_hidden(X2,X3)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_78, negated_conjecture, (r2_hidden(esk7_0,esk6_0)|~v4_ordinal1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_79, negated_conjecture, (~v4_ordinal1(esk6_0)|~r2_hidden(k2_xboole_0(esk7_0,k2_tarski(esk7_0,esk7_0)),esk6_0)), inference(rw,[status(thm)],[c_0_68, c_0_42])).
cnf(c_0_80, negated_conjecture, (k2_xboole_0(X1,k2_tarski(X1,X1))=esk6_0|r2_hidden(esk6_0,k2_xboole_0(X1,k2_tarski(X1,X1)))|r2_hidden(k2_xboole_0(X1,k2_tarski(X1,X1)),esk6_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_81, negated_conjecture, (v3_ordinal1(esk7_0)|~v4_ordinal1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_82, plain, (k3_tarski(k2_xboole_0(X1,X2))=k2_xboole_0(k3_tarski(X1),k3_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_83, plain, (k2_xboole_0(k3_tarski(X1),X1)=X1|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_56, c_0_72])).
cnf(c_0_84, plain, (v3_ordinal1(k3_tarski(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_85, plain, (~v3_ordinal1(X1)|~r2_hidden(X1,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_20, c_0_72])).
cnf(c_0_86, negated_conjecture, (k3_tarski(esk6_0)=esk6_0|~v3_ordinal1(k3_tarski(esk6_0))|~r2_hidden(k3_tarski(esk6_0),esk6_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_87, plain, (k2_xboole_0(X1,X2)!=k3_tarski(X3)|~r2_hidden(k3_tarski(X3),X2)), inference(spm,[status(thm)],[c_0_44, c_0_76])).
cnf(c_0_88, negated_conjecture, (k3_tarski(X1)!=esk7_0|~v4_ordinal1(esk6_0)|~r2_hidden(esk6_0,X1)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_89, negated_conjecture, (k2_xboole_0(esk7_0,k2_tarski(esk7_0,esk7_0))=esk6_0|r2_hidden(esk6_0,k2_xboole_0(esk7_0,k2_tarski(esk7_0,esk7_0)))|~v4_ordinal1(esk6_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_81])).
cnf(c_0_90, plain, (k2_xboole_0(k3_tarski(X1),X2)=k3_tarski(k2_xboole_0(X1,k2_tarski(X2,X2)))), inference(spm,[status(thm)],[c_0_82, c_0_55])).
cnf(c_0_91, negated_conjecture, (k2_xboole_0(k3_tarski(esk7_0),esk7_0)=esk7_0|~v4_ordinal1(esk6_0)), inference(spm,[status(thm)],[c_0_83, c_0_81])).
cnf(c_0_92, negated_conjecture, (k3_tarski(X1)=esk6_0|r2_hidden(esk6_0,k3_tarski(X1))|r2_hidden(k3_tarski(X1),esk6_0)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_69, c_0_84])).
cnf(c_0_93, negated_conjecture, (~r2_hidden(esk6_0,k3_tarski(esk6_0))), inference(spm,[status(thm)],[c_0_85, c_0_58])).
cnf(c_0_94, negated_conjecture, (k3_tarski(esk6_0)=esk6_0|~r2_hidden(k3_tarski(esk6_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_84]), c_0_58])])).
cnf(c_0_95, plain, (k2_xboole_0(X1,X2)!=X3|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_87, c_0_55])).
cnf(c_0_96, negated_conjecture, (k2_xboole_0(k3_tarski(esk6_0),esk6_0)=esk6_0), inference(spm,[status(thm)],[c_0_83, c_0_58])).
cnf(c_0_97, negated_conjecture, (k2_xboole_0(esk7_0,k2_tarski(esk7_0,esk7_0))=esk6_0|~v4_ordinal1(esk6_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90]), c_0_91])).
cnf(c_0_98, negated_conjecture, (k3_tarski(esk6_0)=esk6_0), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_58]), c_0_93]), c_0_94])).
cnf(c_0_99, negated_conjecture, (esk6_0!=X1|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_100, negated_conjecture, (k2_xboole_0(k3_tarski(esk7_0),esk7_0)=esk6_0|~v4_ordinal1(esk6_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_97]), c_0_98])).
cnf(c_0_101, negated_conjecture, (esk7_0!=esk6_0|~v4_ordinal1(esk6_0)), inference(spm,[status(thm)],[c_0_99, c_0_78])).
cnf(c_0_102, negated_conjecture, (~v4_ordinal1(esk6_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_100]), c_0_101])).
cnf(c_0_103, plain, (v4_ordinal1(X1)|X1!=k3_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_104, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_98])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 105
# Proof object clause steps            : 68
# Proof object formula steps           : 37
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 18
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 35
# Processed clauses                    : 2047
# ...of these trivial                  : 13
# ...subsumed                          : 1433
# ...remaining for further processing  : 601
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 31
# Backward-rewritten                   : 112
# Generated clauses                    : 7517
# ...of the previous two non-trivial   : 7440
# Contextual simplify-reflections      : 25
# Paramodulations                      : 7483
# Factorizations                       : 6
# Equation resolutions                 : 28
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 423
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 54
#    Non-unit-clauses                  : 356
# Current number of unprocessed clauses: 5398
# ...number of literals in the above   : 16530
# Current number of archived formulas  : 0
# Current number of archived clauses   : 180
# Clause-clause subsumption calls (NU) : 92989
# Rec. Clause-clause subsumption calls : 33246
# Non-unit clause-clause subsumptions  : 784
# Unit Clause-clause subsumption calls : 8973
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 60
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 87982

# -------------------------------------------------
# User time                : 0.084 s
# System time              : 0.000 s
# Total time               : 0.084 s
# Maximum resident set size: 28756 pages
