# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t27_yellow_1, conjecture, ![X1]:(l1_orders_2(X1)=>k6_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', t27_yellow_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', t70_enumset1)).
fof(t26_yellow_1, axiom, ![X1]:((v1_relat_1(X1)&(v4_relat_1(X1,k1_xboole_0)&(v1_funct_1(X1)&(v1_partfun1(X1,k1_xboole_0)&v1_yellow_1(X1)))))=>k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', t26_yellow_1)).
fof(fc2_funcop_1, axiom, ![X1]:v1_xboole_0(k2_funcop_1(k1_xboole_0,X1)), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', fc2_funcop_1)).
fof(redefinition_k7_funcop_1, axiom, ![X1, X2]:k7_funcop_1(X1,X2)=k2_funcop_1(X1,X2), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', redefinition_k7_funcop_1)).
fof(t134_pboole, axiom, ![X1]:((v1_relat_1(X1)&(v4_relat_1(X1,k1_xboole_0)&(v1_funct_1(X1)&v1_partfun1(X1,k1_xboole_0))))=>X1=k1_xboole_0), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', t134_pboole)).
fof(rc1_yellow_1, axiom, ![X1]:?[X2]:(v1_relat_1(X2)&(v4_relat_1(X2,X1)&(v1_funct_1(X2)&(v1_partfun1(X2,X1)&v1_yellow_1(X2))))), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', rc1_yellow_1)).
fof(d5_yellow_1, axiom, ![X1, X2]:(l1_orders_2(X2)=>k6_yellow_1(X1,X2)=k5_yellow_1(X1,k7_funcop_1(X1,X2))), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', d5_yellow_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/c09480fdc7674914b89145475b43abf1.E_input', l13_xboole_0)).
fof(c_0_10, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>k6_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0))))), inference(assume_negation,[status(cth)],[t27_yellow_1])).
fof(c_0_11, negated_conjecture, (l1_orders_2(esk1_0)&k6_yellow_1(k1_xboole_0,esk1_0)!=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
fof(c_0_12, plain, ![X6]:k2_tarski(X6,X6)=k1_tarski(X6), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_13, plain, ![X8, X9]:k1_enumset1(X8,X8,X9)=k2_tarski(X8,X9), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_14, plain, ![X15]:(~v1_relat_1(X15)|(~v4_relat_1(X15,k1_xboole_0)|(~v1_funct_1(X15)|(~v1_partfun1(X15,k1_xboole_0)|~v1_yellow_1(X15))))|k5_yellow_1(k1_xboole_0,X15)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_yellow_1])])).
cnf(c_0_15, negated_conjecture, (k6_yellow_1(k1_xboole_0,esk1_0)!=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_tarski(k1_xboole_0),k6_partfun1(k1_tarski(k1_xboole_0)))|~v1_relat_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v1_partfun1(X1,k1_xboole_0)|~v1_yellow_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_19, plain, ![X16]:v1_xboole_0(k2_funcop_1(k1_xboole_0,X16)), inference(variable_rename,[status(thm)],[fc2_funcop_1])).
fof(c_0_20, plain, ![X10, X11]:k7_funcop_1(X10,X11)=k2_funcop_1(X10,X11), inference(variable_rename,[status(thm)],[redefinition_k7_funcop_1])).
fof(c_0_21, plain, ![X7]:(~v1_relat_1(X7)|(~v4_relat_1(X7,k1_xboole_0)|(~v1_funct_1(X7)|~v1_partfun1(X7,k1_xboole_0)))|X7=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t134_pboole])])).
fof(c_0_22, plain, ![X12]:(v1_relat_1(esk2_1(X12))&(v4_relat_1(esk2_1(X12),X12)&(v1_funct_1(esk2_1(X12))&(v1_partfun1(esk2_1(X12),X12)&v1_yellow_1(esk2_1(X12)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_yellow_1])])).
cnf(c_0_23, negated_conjecture, (k6_yellow_1(k1_xboole_0,esk1_0)!=g1_orders_2(k1_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0),k6_partfun1(k1_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16]), c_0_16]), c_0_17]), c_0_17])).
cnf(c_0_24, plain, (k5_yellow_1(k1_xboole_0,X1)=g1_orders_2(k1_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0),k6_partfun1(k1_enumset1(k1_xboole_0,k1_xboole_0,k1_xboole_0)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_yellow_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_partfun1(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_16]), c_0_16]), c_0_17]), c_0_17])).
fof(c_0_25, plain, ![X4, X5]:(~l1_orders_2(X5)|k6_yellow_1(X4,X5)=k5_yellow_1(X4,k7_funcop_1(X4,X5))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_yellow_1])])).
fof(c_0_26, plain, ![X14]:(~v1_xboole_0(X14)|X14=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_27, plain, (v1_xboole_0(k2_funcop_1(k1_xboole_0,X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, plain, (k7_funcop_1(X1,X2)=k2_funcop_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (X1=k1_xboole_0|~v1_relat_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v1_partfun1(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (v1_partfun1(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (v1_funct_1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (v4_relat_1(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (v1_relat_1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_34, negated_conjecture, (k5_yellow_1(k1_xboole_0,X1)!=k6_yellow_1(k1_xboole_0,esk1_0)|~v1_yellow_1(X1)|~v1_partfun1(X1,k1_xboole_0)|~v1_funct_1(X1)|~v4_relat_1(X1,k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_35, plain, (k6_yellow_1(X2,X1)=k5_yellow_1(X2,k7_funcop_1(X2,X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_37, plain, (v1_xboole_0(k7_funcop_1(k1_xboole_0,X1))), inference(rw,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_38, plain, (v1_yellow_1(esk2_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_39, plain, (esk2_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32]), c_0_33])])).
cnf(c_0_40, negated_conjecture, (k6_yellow_1(k1_xboole_0,X1)!=k6_yellow_1(k1_xboole_0,esk1_0)|~v1_yellow_1(k7_funcop_1(k1_xboole_0,X1))|~v1_partfun1(k7_funcop_1(k1_xboole_0,X1),k1_xboole_0)|~v1_funct_1(k7_funcop_1(k1_xboole_0,X1))|~v4_relat_1(k7_funcop_1(k1_xboole_0,X1),k1_xboole_0)|~v1_relat_1(k7_funcop_1(k1_xboole_0,X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41, plain, (k7_funcop_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42, plain, (v1_yellow_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43, plain, (v1_partfun1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_30, c_0_39])).
cnf(c_0_44, plain, (v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_31, c_0_39])).
cnf(c_0_45, plain, (v4_relat_1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_32, c_0_39])).
cnf(c_0_46, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_33, c_0_39])).
cnf(c_0_47, negated_conjecture, (k6_yellow_1(k1_xboole_0,X1)!=k6_yellow_1(k1_xboole_0,esk1_0)|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_41]), c_0_43]), c_0_41]), c_0_44]), c_0_41]), c_0_45]), c_0_41]), c_0_46])])).
cnf(c_0_48, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_49, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_47]), c_0_48])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 50
# Proof object clause steps            : 29
# Proof object formula steps           : 21
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 26
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 12
# Processed clauses                    : 36
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 35
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 14
# ...of the previous two non-trivial   : 15
# Contextual simplify-reflections      : 0
# Paramodulations                      : 13
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 6
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 20
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 50
# Rec. Clause-clause subsumption calls : 2
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1121

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.000 s
# Total time               : 0.009 s
# Maximum resident set size: 28856 pages
