# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t31_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2))))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t31_relat_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t1_xboole_1)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t44_xboole_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', l51_zfmisc_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t36_xboole_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t43_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t3_xboole_0)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t7_tarski)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', d6_relat_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', commutativity_k2_tarski)).
fof(t136_zfmisc_1, axiom, ![X1]:?[X2]:(r2_hidden(X1,X2)&(![X3, X4]:((r2_hidden(X3,X2)&r1_tarski(X4,X3))=>r2_hidden(X4,X2))&(![X3]:(r2_hidden(X3,X2)=>r2_hidden(k1_zfmisc_1(X3),X2))&![X3]:~((r1_tarski(X3,X2)&(~(r2_tarski(X3,X2))&~(r2_hidden(X3,X2)))))))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t136_zfmisc_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t8_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t7_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', d10_xboole_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t2_xboole_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t65_xboole_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', d5_relat_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t10_xboole_1)).
fof(t13_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k1_relat_1(k2_xboole_0(X1,X2))=k2_xboole_0(k1_relat_1(X1),k1_relat_1(X2)))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t13_relat_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t39_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t1_boole)).
fof(t28_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', t28_relat_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', redefinition_k6_subset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/c1f6406d06da44aabe9a59bf328c90ac.E_input', d3_tarski)).
fof(c_0_24, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t31_relat_1])).
fof(c_0_25, plain, ![X35, X36, X37]:(~r1_tarski(X35,X36)|~r1_tarski(X36,X37)|r1_tarski(X35,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_26, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
fof(c_0_27, plain, ![X20, X21, X22]:(~r1_tarski(k4_xboole_0(X20,X21),X22)|r1_tarski(X20,k2_xboole_0(X21,X22))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
fof(c_0_28, plain, ![X10, X11]:k3_tarski(k2_tarski(X10,X11))=k2_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
cnf(c_0_29, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X38, X39]:r1_tarski(k4_xboole_0(X38,X39),X38), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_32, plain, ![X5, X6, X7]:(~r1_tarski(X5,k2_xboole_0(X6,X7))|r1_tarski(k4_xboole_0(X5,X6),X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_33, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_37, plain, ![X25, X26, X28, X29, X30]:(((r2_hidden(esk3_2(X25,X26),X25)|r1_xboole_0(X25,X26))&(r2_hidden(esk3_2(X25,X26),X26)|r1_xboole_0(X25,X26)))&(~r2_hidden(X30,X28)|~r2_hidden(X30,X29)|~r1_xboole_0(X28,X29))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_38, plain, ![X40, X41, X43]:((r2_hidden(esk4_2(X40,X41),X41)|~r2_hidden(X40,X41))&(~r2_hidden(X43,X41)|~r2_hidden(X43,esk4_2(X40,X41))|~r2_hidden(X40,X41))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
fof(c_0_39, plain, ![X61]:(~v1_relat_1(X61)|k3_relat_1(X61)=k2_xboole_0(k1_relat_1(X61),k2_relat_1(X61))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
cnf(c_0_40, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X2,X3)))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_42, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_43, plain, ![X62, X63]:k2_tarski(X62,X63)=k2_tarski(X63,X62), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_44, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, plain, (r2_hidden(esk4_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_46, plain, ![X44, X46, X47, X48, X49]:(r2_hidden(X44,esk5_1(X44))&((~r2_hidden(X46,esk5_1(X44))|~r1_tarski(X47,X46)|r2_hidden(X47,esk5_1(X44)))&((~r2_hidden(X48,esk5_1(X44))|r2_hidden(k1_zfmisc_1(X48),esk5_1(X44)))&(~r1_tarski(X49,esk5_1(X44))|(r2_tarski(X49,esk5_1(X44))|r2_hidden(X49,esk5_1(X44))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_zfmisc_1])])])])])).
fof(c_0_47, plain, ![X64, X65, X66]:(~r1_tarski(X64,X65)|~r1_tarski(X66,X65)|r1_tarski(k2_xboole_0(X64,X66),X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_48, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_49, plain, ![X12, X13]:r1_tarski(X12,k2_xboole_0(X12,X13)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_50, plain, ![X14, X15]:(((r1_tarski(X14,X15)|X14!=X15)&(r1_tarski(X15,X14)|X14!=X15))&(~r1_tarski(X14,X15)|~r1_tarski(X15,X14)|X14=X15)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_51, plain, ![X16]:r1_tarski(k1_xboole_0,X16), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_52, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_29, c_0_36])).
cnf(c_0_53, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[c_0_40, c_0_34])).
cnf(c_0_54, negated_conjecture, (r1_tarski(esk1_0,k3_tarski(k2_tarski(X1,esk2_0)))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_55, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_56, plain, (~r2_hidden(esk4_2(X1,X2),X3)|~r2_hidden(X1,X2)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_57, plain, (r2_hidden(X1,esk5_1(X1))), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_58, plain, ![X31]:r1_xboole_0(X31,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_59, plain, ![X50, X51, X52, X54, X55, X56, X57, X59]:(((~r2_hidden(X52,X51)|r2_hidden(k4_tarski(esk6_3(X50,X51,X52),X52),X50)|X51!=k2_relat_1(X50))&(~r2_hidden(k4_tarski(X55,X54),X50)|r2_hidden(X54,X51)|X51!=k2_relat_1(X50)))&((~r2_hidden(esk7_2(X56,X57),X57)|~r2_hidden(k4_tarski(X59,esk7_2(X56,X57)),X56)|X57=k2_relat_1(X56))&(r2_hidden(esk7_2(X56,X57),X57)|r2_hidden(k4_tarski(esk8_2(X56,X57),esk7_2(X56,X57)),X56)|X57=k2_relat_1(X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_60, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_61, plain, (k3_relat_1(X1)=k3_tarski(k2_tarski(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_48, c_0_34])).
fof(c_0_62, plain, ![X32, X33, X34]:(~r1_tarski(X32,X33)|r1_tarski(X32,k2_xboole_0(X34,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_63, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_64, plain, ![X75, X76]:(~v1_relat_1(X75)|(~v1_relat_1(X76)|k1_relat_1(k2_xboole_0(X75,X76))=k2_xboole_0(k1_relat_1(X75),k1_relat_1(X76)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_relat_1])])])).
fof(c_0_65, plain, ![X17, X18]:k2_xboole_0(X17,k4_xboole_0(X18,X17))=k2_xboole_0(X17,X18), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_66, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_67, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_68, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k2_tarski(X2,k4_xboole_0(X3,X4))))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_69, negated_conjecture, (r1_tarski(esk1_0,k3_tarski(k2_tarski(esk2_0,X1)))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
fof(c_0_70, plain, ![X19]:k2_xboole_0(X19,k1_xboole_0)=X19, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_71, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(esk5_1(esk4_2(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_72, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_73, plain, (r2_hidden(k4_tarski(esk6_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_74, plain, (r1_tarski(k3_tarski(k2_tarski(X1,X3)),X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_60, c_0_34])).
cnf(c_0_75, plain, (k3_tarski(k2_tarski(k2_relat_1(X1),k1_relat_1(X1)))=k3_relat_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_61, c_0_55])).
cnf(c_0_76, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_77, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_63, c_0_34])).
cnf(c_0_78, plain, (k1_relat_1(k2_xboole_0(X1,X2))=k2_xboole_0(k1_relat_1(X1),k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_79, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_80, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_81, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_82, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_83, plain, ![X73, X74]:(~v1_relat_1(X73)|(~v1_relat_1(X74)|r1_tarski(k6_subset_1(k2_relat_1(X73),k2_relat_1(X74)),k2_relat_1(k6_subset_1(X73,X74))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_relat_1])])])).
fof(c_0_84, plain, ![X8, X9]:k6_subset_1(X8,X9)=k4_xboole_0(X8,X9), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_85, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_86, plain, (r2_hidden(k4_tarski(esk6_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_73])).
fof(c_0_87, plain, ![X67, X68, X69, X70, X71]:((~r1_tarski(X67,X68)|(~r2_hidden(X69,X67)|r2_hidden(X69,X68)))&((r2_hidden(esk9_2(X70,X71),X70)|r1_tarski(X70,X71))&(~r2_hidden(esk9_2(X70,X71),X71)|r1_tarski(X70,X71)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_88, negated_conjecture, (~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_89, plain, (r1_tarski(k3_relat_1(X1),X2)|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_90, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_91, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_76, c_0_34])).
cnf(c_0_92, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X2,X1)))), inference(spm,[status(thm)],[c_0_77, c_0_55])).
cnf(c_0_93, plain, (k1_relat_1(k3_tarski(k2_tarski(X1,X2)))=k3_tarski(k2_tarski(k1_relat_1(X1),k1_relat_1(X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_34]), c_0_34])).
cnf(c_0_94, plain, (k3_tarski(k2_tarski(X1,k4_xboole_0(X2,X1)))=k3_tarski(k2_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_34]), c_0_34])).
cnf(c_0_95, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_96, plain, (k3_tarski(k2_tarski(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_82, c_0_34])).
cnf(c_0_97, plain, (r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_98, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_99, plain, (~r2_hidden(X1,k2_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_100, plain, (r2_hidden(esk9_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_101, negated_conjecture, (~r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))|~r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90])])).
cnf(c_0_102, plain, (r1_tarski(X1,k3_relat_1(X2))|~v1_relat_1(X2)|~r1_tarski(X1,k1_relat_1(X2))), inference(spm,[status(thm)],[c_0_91, c_0_75])).
cnf(c_0_103, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_104, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(k3_tarski(k2_tarski(X2,X1))))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_105, negated_conjecture, (k3_tarski(k2_tarski(esk2_0,esk1_0))=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])).
cnf(c_0_106, plain, (r1_tarski(k4_xboole_0(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k4_xboole_0(X1,X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_98]), c_0_98])).
cnf(c_0_107, plain, (r1_tarski(k2_relat_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_108, negated_conjecture, (~r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))|~r1_tarski(k1_relat_1(esk1_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103])])).
cnf(c_0_109, negated_conjecture, (r1_tarski(k1_relat_1(esk1_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_90]), c_0_103])])).
cnf(c_0_110, plain, (r1_tarski(k2_relat_1(X1),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_77, c_0_75])).
cnf(c_0_111, plain, (r1_tarski(k2_relat_1(X1),k3_tarski(k2_tarski(k2_relat_1(X2),k2_relat_1(k4_xboole_0(X1,X2)))))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_41, c_0_106])).
cnf(c_0_112, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_80, c_0_107])).
cnf(c_0_113, negated_conjecture, (~r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_109])])).
cnf(c_0_114, plain, (r1_tarski(X1,k3_relat_1(X2))|~v1_relat_1(X2)|~r1_tarski(X1,k2_relat_1(X2))), inference(spm,[status(thm)],[c_0_29, c_0_110])).
cnf(c_0_115, negated_conjecture, (r1_tarski(k2_relat_1(esk1_0),k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_95]), c_0_112]), c_0_96]), c_0_103]), c_0_90])])).
cnf(c_0_116, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_103]), c_0_115])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 117
# Proof object clause steps            : 68
# Proof object formula steps           : 49
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 38
# Processed clauses                    : 1134
# ...of these trivial                  : 5
# ...subsumed                          : 702
# ...remaining for further processing  : 427
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 17
# Generated clauses                    : 3701
# ...of the previous two non-trivial   : 3258
# Contextual simplify-reflections      : 8
# Paramodulations                      : 3697
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 363
#    Positive orientable unit clauses  : 32
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 317
# Current number of unprocessed clauses: 2159
# ...number of literals in the above   : 7042
# Current number of archived formulas  : 0
# Current number of archived clauses   : 62
# Clause-clause subsumption calls (NU) : 31208
# Rec. Clause-clause subsumption calls : 19553
# Non-unit clause-clause subsumptions  : 588
# Unit Clause-clause subsumption calls : 1514
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 86
# BW rewrite match successes           : 19
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 43523

# -------------------------------------------------
# User time                : 0.043 s
# System time              : 0.000 s
# Total time               : 0.043 s
# Maximum resident set size: 28856 pages
