# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/c29f7d71c77040658bf94a3f9c343a8f.E_input', d21_lattice3)).
fof(t41_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((r2_hidden(X2,X3)&r4_lattice3(X1,X2,X3))=>k15_lattice3(X1,X3)=X2))), file('data/MPTP/proofs/c29f7d71c77040658bf94a3f9c343a8f.E_input', t41_lattice3)).
fof(t26_lattices, axiom, ![X1]:((~(v2_struct_0(X1))&(v4_lattices(X1)&l2_lattices(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattices(X1,X2,X3)&r1_lattices(X1,X3,X2))=>X2=X3)))), file('data/MPTP/proofs/c29f7d71c77040658bf94a3f9c343a8f.E_input', t26_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/c29f7d71c77040658bf94a3f9c343a8f.E_input', cc1_lattices)).
fof(t38_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r2_hidden(X2,X3)=>(r3_lattices(X1,X2,k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),X2))))), file('data/MPTP/proofs/c29f7d71c77040658bf94a3f9c343a8f.E_input', t38_lattice3)).
fof(dt_l3_lattices, axiom, ![X1]:(l3_lattices(X1)=>(l1_lattices(X1)&l2_lattices(X1))), file('data/MPTP/proofs/c29f7d71c77040658bf94a3f9c343a8f.E_input', dt_l3_lattices)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/c29f7d71c77040658bf94a3f9c343a8f.E_input', redefinition_r3_lattices)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/c29f7d71c77040658bf94a3f9c343a8f.E_input', dt_k15_lattice3)).
fof(c_0_8, plain, ![X5, X6, X7, X8]:(((r4_lattice3(X5,X7,X6)|X7!=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5)))&(~m1_subset_1(X8,u1_struct_0(X5))|(~r4_lattice3(X5,X8,X6)|r1_lattices(X5,X7,X8))|X7!=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5))))&((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r4_lattice3(X5,X7,X6)|X7=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5)))&((r4_lattice3(X5,esk1_3(X5,X6,X7),X6)|~r4_lattice3(X5,X7,X6)|X7=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5)))&(~r1_lattices(X5,X7,esk1_3(X5,X6,X7))|~r4_lattice3(X5,X7,X6)|X7=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((r2_hidden(X2,X3)&r4_lattice3(X1,X2,X3))=>k15_lattice3(X1,X3)=X2)))), inference(assume_negation,[status(cth)],[t41_lattice3])).
cnf(c_0_10, plain, (r1_lattices(X2,X4,X1)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r4_lattice3(X2,X1,X3)|X4!=k15_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk2_0)&(v10_lattices(esk2_0)&(v4_lattice3(esk2_0)&l3_lattices(esk2_0))))&(m1_subset_1(esk3_0,u1_struct_0(esk2_0))&((r2_hidden(esk3_0,esk4_0)&r4_lattice3(esk2_0,esk3_0,esk4_0))&k15_lattice3(esk2_0,esk4_0)!=esk3_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
cnf(c_0_12, plain, (v2_struct_0(X2)|r1_lattices(X2,X4,X1)|X4!=k15_lattice3(X2,X3)|~l3_lattices(X2)|~v10_lattices(X2)|~v4_lattice3(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(cn,[status(thm)],[c_0_10])).
cnf(c_0_13, negated_conjecture, (r4_lattice3(esk2_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (v4_lattice3(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (v10_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (l3_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_19, plain, ![X13, X14, X15]:(v2_struct_0(X13)|(~v4_lattices(X13)|~l2_lattices(X13))|(~m1_subset_1(X14,u1_struct_0(X13))|(~m1_subset_1(X15,u1_struct_0(X13))|(~r1_lattices(X13,X14,X15)|~r1_lattices(X13,X15,X14)|X14=X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_lattices])])])])).
cnf(c_0_20, negated_conjecture, (r1_lattices(esk2_0,X1,esk3_0)|X1!=k15_lattice3(esk2_0,esk4_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_21, plain, (v2_struct_0(X1)|X2=X3|~v4_lattices(X1)|~l2_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_lattices(X1,X2,X3)|~r1_lattices(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_22, negated_conjecture, (r1_lattices(esk2_0,k15_lattice3(esk2_0,esk4_0),esk3_0)|~m1_subset_1(k15_lattice3(esk2_0,esk4_0),u1_struct_0(esk2_0))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_23, negated_conjecture, (k15_lattice3(esk2_0,esk4_0)!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_24, plain, ![X16]:((~v2_struct_0(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v4_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v5_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v6_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v7_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&((v8_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16))&(v9_lattices(X16)|(v2_struct_0(X16)|~v10_lattices(X16))|~l3_lattices(X16)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_25, plain, ![X17, X18, X19]:((r3_lattices(X17,X18,k15_lattice3(X17,X19))|~r2_hidden(X18,X19)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v10_lattices(X17)|(~v4_lattice3(X17)|~l3_lattices(X17)))))&(r3_lattices(X17,k16_lattice3(X17,X19),X18)|~r2_hidden(X18,X19)|~m1_subset_1(X18,u1_struct_0(X17))|(v2_struct_0(X17)|(~v10_lattices(X17)|(~v4_lattice3(X17)|~l3_lattices(X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t38_lattice3])])])])])).
cnf(c_0_26, negated_conjecture, (~l2_lattices(esk2_0)|~v4_lattices(esk2_0)|~r1_lattices(esk2_0,esk3_0,k15_lattice3(esk2_0,esk4_0))|~m1_subset_1(k15_lattice3(esk2_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_14])]), c_0_23]), c_0_18])).
cnf(c_0_27, plain, (v4_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_28, plain, ![X20]:((l1_lattices(X20)|~l3_lattices(X20))&(l2_lattices(X20)|~l3_lattices(X20))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l3_lattices])])])).
fof(c_0_29, plain, ![X21, X22, X23]:((~r3_lattices(X21,X22,X23)|r1_lattices(X21,X22,X23)|(v2_struct_0(X21)|(~v6_lattices(X21)|(~v8_lattices(X21)|(~v9_lattices(X21)|~l3_lattices(X21))))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))&(~r1_lattices(X21,X22,X23)|r3_lattices(X21,X22,X23)|(v2_struct_0(X21)|(~v6_lattices(X21)|(~v8_lattices(X21)|(~v9_lattices(X21)|~l3_lattices(X21))))|(~m1_subset_1(X22,u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
cnf(c_0_30, plain, (r3_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_32, plain, ![X24, X25]:(v2_struct_0(X24)|~l3_lattices(X24)|m1_subset_1(k15_lattice3(X24,X25),u1_struct_0(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_33, negated_conjecture, (~l2_lattices(esk2_0)|~r1_lattices(esk2_0,esk3_0,k15_lattice3(esk2_0,esk4_0))|~m1_subset_1(k15_lattice3(esk2_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_34, plain, (l2_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (r1_lattices(X1,X2,X3)|v2_struct_0(X1)|~r3_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (r3_lattices(X1,esk3_0,k15_lattice3(X1,esk4_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_39, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_40, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_41, negated_conjecture, (~r1_lattices(esk2_0,esk3_0,k15_lattice3(esk2_0,esk4_0))|~m1_subset_1(k15_lattice3(esk2_0,esk4_0),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_17])])).
cnf(c_0_42, negated_conjecture, (r1_lattices(X1,esk3_0,k15_lattice3(X1,esk4_0))|v2_struct_0(X1)|~m1_subset_1(esk3_0,u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_43, negated_conjecture, (~m1_subset_1(k15_lattice3(esk2_0,esk4_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_14]), c_0_15]), c_0_16]), c_0_17])]), c_0_18])).
cnf(c_0_44, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_37]), c_0_17])]), c_0_18]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 28
# Proof object formula steps           : 17
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 27
# Processed clauses                    : 65
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 65
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 17
# ...of the previous two non-trivial   : 15
# Contextual simplify-reflections      : 6
# Paramodulations                      : 15
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 36
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 27
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 39
# Current number of archived formulas  : 0
# Current number of archived clauses   : 29
# Clause-clause subsumption calls (NU) : 1306
# Rec. Clause-clause subsumption calls : 123
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3189

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28696 pages
