# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B02_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', t70_xboole_1)).
fof(l27_zfmisc_1, axiom, ![X1, X2]:(~(r2_hidden(X1,X2))=>r1_xboole_0(k1_tarski(X1),X2)), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', l27_zfmisc_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', t12_xboole_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', t7_ordinal1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', l32_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', d10_xboole_0)).
fof(t4_boole, axiom, ![X1]:k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', t4_boole)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', redefinition_k6_domain_1)).
fof(d1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>(v1_zfmisc_1(X1)<=>?[X2]:(m1_subset_1(X2,X1)&X1=k6_domain_1(X1,X2)))), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', d1_tex_2)).
fof(t1_tex_2, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2))), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', t1_tex_2)).
fof(t68_xboole_1, axiom, ![X1, X2, X3]:(~(v1_xboole_0(X3))=>~((r1_tarski(X3,X1)&(r1_tarski(X3,X2)&r1_xboole_0(X1,X2))))), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', t68_xboole_1)).
fof(t43_zfmisc_1, axiom, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))), file('data/MPTP/proofs/c326206cdab943b091acbb142830ab28.E_input', t43_zfmisc_1)).
fof(c_0_12, plain, ![X25, X26, X27]:((r1_xboole_0(X25,k2_xboole_0(X26,X27))|(~r1_xboole_0(X25,X26)|~r1_xboole_0(X25,X27)))&((r1_xboole_0(X25,X26)|~r1_xboole_0(X25,k2_xboole_0(X26,X27)))&(r1_xboole_0(X25,X27)|~r1_xboole_0(X25,k2_xboole_0(X26,X27))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
fof(c_0_13, plain, ![X6, X7]:(r2_hidden(X6,X7)|r1_xboole_0(k1_tarski(X6),X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l27_zfmisc_1])])])).
fof(c_0_14, plain, ![X15, X16]:(~r1_tarski(X15,X16)|k2_xboole_0(X15,X16)=X16), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_15, plain, ![X22, X23]:(~r2_hidden(X22,X23)|~r1_tarski(X23,X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_16, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|r1_xboole_0(k1_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_20, plain, ![X10, X11]:((k4_xboole_0(X10,X11)!=k1_xboole_0|r1_tarski(X10,X11))&(~r1_tarski(X10,X11)|k4_xboole_0(X10,X11)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_21, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_22, plain, (r1_xboole_0(k1_tarski(X1),X2)|r2_hidden(X1,k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_23, plain, ![X17, X18]:(((r1_tarski(X17,X18)|X17!=X18)&(r1_tarski(X18,X17)|X17!=X18))&(~r1_tarski(X17,X18)|~r1_tarski(X18,X17)|X17=X18)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_24, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,X3)|~r1_tarski(X2,X3)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X24]:k4_xboole_0(k1_xboole_0,X24)=k1_xboole_0, inference(variable_rename,[status(thm)],[t4_boole])).
cnf(c_0_27, plain, (r1_xboole_0(k1_tarski(X1),X2)|~r1_tarski(k2_xboole_0(X3,X2),X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X2,X3)!=k1_xboole_0|~r1_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, plain, (r1_xboole_0(k1_tarski(X1),X2)|~r1_tarski(X2,X1)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_27, c_0_19])).
cnf(c_0_32, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_28])).
fof(c_0_33, plain, ![X8, X9]:(v1_xboole_0(X8)|~m1_subset_1(X9,X8)|k6_domain_1(X8,X9)=k1_tarski(X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_34, plain, ![X28, X30]:(((m1_subset_1(esk3_1(X28),X28)|~v1_zfmisc_1(X28)|v1_xboole_0(X28))&(X28=k6_domain_1(X28,esk3_1(X28))|~v1_zfmisc_1(X28)|v1_xboole_0(X28)))&(~m1_subset_1(X30,X28)|X28!=k6_domain_1(X28,X30)|v1_zfmisc_1(X28)|v1_xboole_0(X28))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tex_2])])])])])])).
cnf(c_0_35, plain, (r1_xboole_0(X1,k1_xboole_0)|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (r1_xboole_0(k1_tarski(X1),X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_37, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((~(v1_xboole_0(X2))&v1_zfmisc_1(X2))=>(r1_tarski(X1,X2)=>X1=X2)))), inference(assume_negation,[status(cth)],[t1_tex_2])).
cnf(c_0_38, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (m1_subset_1(esk3_1(X1),X1)|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_40, plain, ![X19, X20, X21]:(v1_xboole_0(X21)|(~r1_tarski(X21,X19)|(~r1_tarski(X21,X20)|~r1_xboole_0(X19,X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t68_xboole_1])])])).
cnf(c_0_41, plain, (r1_xboole_0(k1_tarski(X1),k1_xboole_0)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_42, negated_conjecture, (~v1_xboole_0(esk1_0)&((~v1_xboole_0(esk2_0)&v1_zfmisc_1(esk2_0))&(r1_tarski(esk1_0,esk2_0)&esk1_0!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_37])])])])).
cnf(c_0_43, plain, (X1=k6_domain_1(X1,esk3_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (k6_domain_1(X1,esk3_1(X1))=k1_tarski(esk3_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_45, plain, ![X12, X13, X14]:((((X13=k1_tarski(X12)|X13=k1_xboole_0|X13=k1_tarski(X12)|k1_tarski(X12)!=k2_xboole_0(X13,X14))&(X14=k1_xboole_0|X13=k1_xboole_0|X13=k1_tarski(X12)|k1_tarski(X12)!=k2_xboole_0(X13,X14)))&((X13=k1_tarski(X12)|X14=k1_tarski(X12)|X13=k1_tarski(X12)|k1_tarski(X12)!=k2_xboole_0(X13,X14))&(X14=k1_xboole_0|X14=k1_tarski(X12)|X13=k1_tarski(X12)|k1_tarski(X12)!=k2_xboole_0(X13,X14))))&(((X13=k1_tarski(X12)|X13=k1_xboole_0|X14=k1_tarski(X12)|k1_tarski(X12)!=k2_xboole_0(X13,X14))&(X14=k1_xboole_0|X13=k1_xboole_0|X14=k1_tarski(X12)|k1_tarski(X12)!=k2_xboole_0(X13,X14)))&((X13=k1_tarski(X12)|X14=k1_tarski(X12)|X14=k1_tarski(X12)|k1_tarski(X12)!=k2_xboole_0(X13,X14))&(X14=k1_xboole_0|X14=k1_tarski(X12)|X14=k1_tarski(X12)|k1_tarski(X12)!=k2_xboole_0(X13,X14))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_zfmisc_1])])])).
cnf(c_0_46, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (r1_xboole_0(k1_tarski(X1),k1_xboole_0)), inference(spm,[status(thm)],[c_0_41, c_0_32])).
cnf(c_0_48, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_49, plain, (k1_tarski(esk3_1(X1))=X1|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_50, negated_conjecture, (v1_zfmisc_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, plain, (X1=k1_tarski(X2)|X1=k1_xboole_0|X1=k1_tarski(X2)|k1_tarski(X2)!=k2_xboole_0(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k1_tarski(X2))|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, negated_conjecture, (k1_tarski(esk3_1(esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_54, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|k1_tarski(X2)!=k2_xboole_0(X1,X3)), inference(cn,[status(thm)],[c_0_51])).
cnf(c_0_55, negated_conjecture, (v1_xboole_0(X1)|~r1_tarski(X1,esk2_0)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_56, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_57, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_58, plain, (X1=k1_tarski(X2)|X1=k1_xboole_0|k1_tarski(X2)!=X3|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_54, c_0_19])).
cnf(c_0_59, negated_conjecture, (~r1_tarski(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_60, negated_conjecture, (k1_tarski(X1)=esk1_0|esk1_0=k1_xboole_0|k1_tarski(X1)!=esk2_0), inference(spm,[status(thm)],[c_0_58, c_0_56])).
cnf(c_0_61, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_62, negated_conjecture, (k4_xboole_0(esk1_0,k1_xboole_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_59, c_0_25])).
cnf(c_0_63, negated_conjecture, (esk1_0=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_53]), c_0_61])).
cnf(c_0_64, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_63]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 65
# Proof object clause steps            : 40
# Proof object formula steps           : 25
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 19
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 173
# ...of these trivial                  : 2
# ...subsumed                          : 32
# ...remaining for further processing  : 139
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 9
# Backward-rewritten                   : 41
# Generated clauses                    : 214
# ...of the previous two non-trivial   : 186
# Contextual simplify-reflections      : 0
# Paramodulations                      : 212
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 63
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 45
# Current number of unprocessed clauses: 60
# ...number of literals in the above   : 179
# Current number of archived formulas  : 0
# Current number of archived clauses   : 74
# Clause-clause subsumption calls (NU) : 858
# Rec. Clause-clause subsumption calls : 597
# Non-unit clause-clause subsumptions  : 30
# Unit Clause-clause subsumption calls : 326
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 17
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3889

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.004 s
# Total time               : 0.012 s
# Maximum resident set size: 28800 pages
