# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t7_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', commutativity_k2_xboole_0)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t134_zfmisc_1)).
fof(t120_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))&k2_zfmisc_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k2_zfmisc_1(X3,X1),k2_zfmisc_1(X3,X2))), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t120_zfmisc_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t12_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', commutativity_k3_xboole_0)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t123_zfmisc_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t28_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t4_xboole_0)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', d7_xboole_0)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', d2_zfmisc_1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', idempotence_k2_xboole_0)).
fof(t117_zfmisc_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&((r1_tarski(k2_zfmisc_1(X2,X1),k2_zfmisc_1(X3,X1))|r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3)))&~(r1_tarski(X2,X3))))), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t117_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', d10_xboole_0)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/c341ce1bef2c4b029db5398f51014482.E_input', t7_xboole_0)).
fof(c_0_16, plain, ![X7, X8]:r1_tarski(X7,k2_xboole_0(X7,X8)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_17, plain, ![X41, X42]:k2_xboole_0(X41,X42)=k2_xboole_0(X42,X41), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_18, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
cnf(c_0_19, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_21, plain, ![X45, X46, X47]:(k2_zfmisc_1(k2_xboole_0(X45,X46),X47)=k2_xboole_0(k2_zfmisc_1(X45,X47),k2_zfmisc_1(X46,X47))&k2_zfmisc_1(X47,k2_xboole_0(X45,X46))=k2_xboole_0(k2_zfmisc_1(X47,X45),k2_zfmisc_1(X47,X46))), inference(variable_rename,[status(thm)],[t120_zfmisc_1])).
fof(c_0_22, plain, ![X48, X49]:(~r1_tarski(X48,X49)|k2_xboole_0(X48,X49)=X49), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_23, plain, ![X43, X44]:r1_tarski(k3_xboole_0(X43,X44),X43), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_24, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk1_0!=esk3_0|esk2_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
cnf(c_0_25, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, plain, (k2_zfmisc_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, plain, ![X39, X40]:k3_xboole_0(X39,X40)=k3_xboole_0(X40,X39), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_31, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k2_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_20])).
fof(c_0_33, plain, ![X50, X51, X52, X53]:k2_zfmisc_1(k3_xboole_0(X50,X51),k3_xboole_0(X52,X53))=k3_xboole_0(k2_zfmisc_1(X50,X52),k2_zfmisc_1(X51,X53)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
fof(c_0_34, plain, ![X10, X11]:(~r1_tarski(X10,X11)|k3_xboole_0(X10,X11)=X10), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_35, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))=k2_zfmisc_1(esk1_0,k2_xboole_0(X1,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_29]), c_0_20])).
cnf(c_0_36, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (r1_tarski(k2_zfmisc_1(X1,k3_xboole_0(X2,X3)),k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,k2_xboole_0(X1,esk2_0)))), inference(spm,[status(thm)],[c_0_19, c_0_35])).
fof(c_0_41, plain, ![X16, X17, X19, X20, X21]:((r1_xboole_0(X16,X17)|r2_hidden(esk5_2(X16,X17),k3_xboole_0(X16,X17)))&(~r2_hidden(X21,k3_xboole_0(X19,X20))|~r1_xboole_0(X19,X20))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_42, plain, (k2_zfmisc_1(k2_xboole_0(X1,X2),X3)=k2_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_43, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_28, c_0_36])).
cnf(c_0_44, plain, (r1_tarski(k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4)),k2_zfmisc_1(k3_xboole_0(X1,X3),X2))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,k2_xboole_0(X1,esk2_0)))=k2_zfmisc_1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_46, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_47, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(spm,[status(thm)],[c_0_39, c_0_19])).
fof(c_0_48, plain, ![X54, X55]:((~r1_xboole_0(X54,X55)|k3_xboole_0(X54,X55)=k1_xboole_0)&(k3_xboole_0(X54,X55)!=k1_xboole_0|r1_xboole_0(X54,X55))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_49, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(k2_xboole_0(X3,X1),X2))), inference(spm,[status(thm)],[c_0_25, c_0_42])).
cnf(c_0_50, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_43]), c_0_20])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(k3_xboole_0(esk3_0,esk1_0),esk4_0))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_52, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,k2_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, plain, (r1_xboole_0(X1,X2)|k3_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_54, plain, ![X22, X23, X24, X25, X28, X29, X30, X31, X32, X33, X35, X36]:((((r2_hidden(esk6_4(X22,X23,X24,X25),X22)|~r2_hidden(X25,X24)|X24!=k2_zfmisc_1(X22,X23))&((r2_hidden(esk7_4(X22,X23,X24,X25),X23)|~r2_hidden(X25,X24)|X24!=k2_zfmisc_1(X22,X23))&(X25=k4_tarski(esk6_4(X22,X23,X24,X25),esk7_4(X22,X23,X24,X25))|~r2_hidden(X25,X24)|X24!=k2_zfmisc_1(X22,X23))))&(~r2_hidden(X29,X22)|(~r2_hidden(X30,X23)|X28!=k4_tarski(X29,X30))|r2_hidden(X28,X24)|X24!=k2_zfmisc_1(X22,X23)))&((~r2_hidden(esk8_3(X31,X32,X33),X33)|(~r2_hidden(X35,X31)|(~r2_hidden(X36,X32)|esk8_3(X31,X32,X33)!=k4_tarski(X35,X36)))|X33=k2_zfmisc_1(X31,X32))&((r2_hidden(esk9_3(X31,X32,X33),X31)|r2_hidden(esk8_3(X31,X32,X33),X33)|X33=k2_zfmisc_1(X31,X32))&((r2_hidden(esk10_3(X31,X32,X33),X32)|r2_hidden(esk8_3(X31,X32,X33),X33)|X33=k2_zfmisc_1(X31,X32))&(esk8_3(X31,X32,X33)=k4_tarski(esk9_3(X31,X32,X33),esk10_3(X31,X32,X33))|r2_hidden(esk8_3(X31,X32,X33),X33)|X33=k2_zfmisc_1(X31,X32)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_55, plain, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_56, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,esk1_0),esk4_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_51]), c_0_42]), c_0_32])).
cnf(c_0_57, plain, (~r2_hidden(X1,k1_xboole_0)), inference(er,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_47])])).
cnf(c_0_58, plain, (r2_hidden(esk10_3(X1,X2,X3),X2)|r2_hidden(esk8_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_59, plain, ![X9]:k2_xboole_0(X9,X9)=X9, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_60, plain, (r2_hidden(esk9_3(X1,X2,X3),X1)|r2_hidden(esk8_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,esk4_0))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_62, plain, (X1=k2_zfmisc_1(X2,k1_xboole_0)|r2_hidden(esk8_3(X2,k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_63, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_64, plain, (X1=k2_zfmisc_1(k1_xboole_0,X2)|r2_hidden(esk8_3(k1_xboole_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_57, c_0_60])).
fof(c_0_65, plain, ![X56, X57, X58]:((~r1_tarski(k2_zfmisc_1(X57,X56),k2_zfmisc_1(X58,X56))|r1_tarski(X57,X58)|X56=k1_xboole_0)&(~r1_tarski(k2_zfmisc_1(X56,X57),k2_zfmisc_1(X56,X58))|r1_tarski(X57,X58)|X56=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t117_zfmisc_1])])])])).
cnf(c_0_66, negated_conjecture, (k2_zfmisc_1(esk1_0,k2_xboole_0(esk4_0,esk2_0))=k2_zfmisc_1(esk1_0,esk4_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_61]), c_0_35])).
cnf(c_0_67, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_57, c_0_62])).
cnf(c_0_68, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_19, c_0_63])).
cnf(c_0_69, plain, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,X2),X3),k2_zfmisc_1(X1,X3))), inference(spm,[status(thm)],[c_0_49, c_0_32])).
cnf(c_0_70, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_57, c_0_64])).
cnf(c_0_71, plain, (r1_tarski(X2,X3)|X1=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_72, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,k3_xboole_0(k2_xboole_0(esk4_0,esk2_0),X1)),k2_zfmisc_1(esk1_0,esk4_0))), inference(spm,[status(thm)],[c_0_37, c_0_66])).
cnf(c_0_73, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_74, plain, (r1_tarski(k3_xboole_0(k1_xboole_0,k2_zfmisc_1(X1,X2)),k2_zfmisc_1(k3_xboole_0(X1,X3),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_67]), c_0_36])).
cnf(c_0_75, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_39, c_0_68])).
cnf(c_0_76, plain, (r1_tarski(k2_zfmisc_1(k3_xboole_0(k1_xboole_0,X1),X2),k1_xboole_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
fof(c_0_77, plain, ![X61, X62]:(((r1_tarski(X61,X62)|X61!=X62)&(r1_tarski(X62,X61)|X61!=X62))&(~r1_tarski(X61,X62)|~r1_tarski(X62,X61)|X61=X62)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_78, negated_conjecture, (r1_tarski(k3_xboole_0(k2_xboole_0(esk4_0,esk2_0),X1),esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])).
cnf(c_0_79, plain, (r1_tarski(k1_xboole_0,k2_zfmisc_1(k3_xboole_0(k1_xboole_0,X1),X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_70]), c_0_75])).
cnf(c_0_80, plain, (k2_xboole_0(k1_xboole_0,k2_zfmisc_1(k3_xboole_0(k1_xboole_0,X1),X2))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_76]), c_0_20])).
cnf(c_0_81, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_82, negated_conjecture, (r1_tarski(k2_xboole_0(esk4_0,esk2_0),esk4_0)), inference(spm,[status(thm)],[c_0_78, c_0_47])).
cnf(c_0_83, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(k2_zfmisc_1(esk1_0,X1),k2_zfmisc_1(esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_29]), c_0_73])).
cnf(c_0_84, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_55, c_0_29])).
cnf(c_0_85, plain, (k2_zfmisc_1(k3_xboole_0(k1_xboole_0,X1),X2)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_79]), c_0_80])).
cnf(c_0_86, plain, (r2_hidden(esk6_4(X1,X2,X3,X4),X1)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_87, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_39, c_0_25])).
cnf(c_0_88, negated_conjecture, (k2_xboole_0(esk4_0,esk2_0)=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_19])])).
cnf(c_0_89, negated_conjecture, (r1_tarski(k1_xboole_0,esk2_0)|~r1_tarski(k1_xboole_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_83, c_0_67])).
cnf(c_0_90, negated_conjecture, (r1_tarski(k1_xboole_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_91, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_46, c_0_75])).
cnf(c_0_92, plain, (r2_hidden(esk6_4(X1,X2,k2_zfmisc_1(X1,X2),X3),X1)|~r2_hidden(X3,k2_zfmisc_1(X1,X2))), inference(er,[status(thm)],[c_0_86])).
cnf(c_0_93, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_94, negated_conjecture, (k2_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk2_0))=k2_zfmisc_1(k2_xboole_0(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_42, c_0_29])).
cnf(c_0_95, negated_conjecture, (k3_xboole_0(esk4_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_36])).
cnf(c_0_96, negated_conjecture, (r1_tarski(k1_xboole_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_97, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_98, plain, (r1_tarski(X1,X3)|X2=k1_xboole_0|~r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_99, plain, (~r2_hidden(X1,k2_zfmisc_1(X2,X3))|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_100, plain, (r2_hidden(esk5_2(X1,X1),X1)|r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_93, c_0_75])).
cnf(c_0_101, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(k2_xboole_0(esk1_0,X1),esk2_0))=k2_zfmisc_1(esk3_0,esk4_0)), inference(spm,[status(thm)],[c_0_47, c_0_94])).
cnf(c_0_102, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(X2,esk2_0))=k2_zfmisc_1(k3_xboole_0(X1,X2),esk2_0)), inference(spm,[status(thm)],[c_0_38, c_0_95])).
cnf(c_0_103, negated_conjecture, (~r1_tarski(esk2_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_96]), c_0_97])).
cnf(c_0_104, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0|r1_tarski(X2,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_85]), c_0_85]), c_0_68])])).
cnf(c_0_105, negated_conjecture, (r1_tarski(X1,esk1_0)|~r1_tarski(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_29]), c_0_97])).
cnf(c_0_106, negated_conjecture, (r1_tarski(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(X1,esk4_0))), inference(spm,[status(thm)],[c_0_31, c_0_88])).
cnf(c_0_107, negated_conjecture, (~r2_hidden(X1,k2_zfmisc_1(esk3_0,esk4_0))|~r1_xboole_0(k3_xboole_0(esk3_0,esk1_0),k3_xboole_0(esk3_0,esk1_0))), inference(spm,[status(thm)],[c_0_99, c_0_56])).
cnf(c_0_108, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_46, c_0_100])).
fof(c_0_109, plain, ![X59]:(X59=k1_xboole_0|r2_hidden(esk11_1(X59),X59)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_110, negated_conjecture, (r1_tarski(esk1_0,X1)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_29]), c_0_97])).
cnf(c_0_111, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,k2_xboole_0(esk1_0,X1)),esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_112, negated_conjecture, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_113, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_114, negated_conjecture, (~r2_hidden(X1,k2_zfmisc_1(esk3_0,esk4_0))|~r1_xboole_0(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_115, plain, (X1=k1_xboole_0|r2_hidden(esk11_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_116, negated_conjecture, (r1_tarski(k1_xboole_0,esk1_0)|~r1_tarski(k1_xboole_0,k2_zfmisc_1(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_105, c_0_70])).
cnf(c_0_117, negated_conjecture, (r1_tarski(esk1_0,k3_xboole_0(esk3_0,k2_xboole_0(esk1_0,X1)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_68])])).
cnf(c_0_118, negated_conjecture, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_50, c_0_112])).
cnf(c_0_119, negated_conjecture, (k3_xboole_0(esk3_0,esk1_0)=esk3_0), inference(spm,[status(thm)],[c_0_39, c_0_113])).
cnf(c_0_120, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)=k1_xboole_0|~r1_xboole_0(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_121, negated_conjecture, (r1_tarski(k1_xboole_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_90])])).
cnf(c_0_122, plain, (r1_tarski(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,k2_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_19, c_0_26])).
cnf(c_0_123, negated_conjecture, (esk1_0=esk3_0|~r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_81, c_0_113])).
cnf(c_0_124, negated_conjecture, (r1_tarski(esk1_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])).
cnf(c_0_125, negated_conjecture, (r1_tarski(esk1_0,k1_xboole_0)|~r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_110, c_0_70])).
cnf(c_0_126, negated_conjecture, (k2_zfmisc_1(esk3_0,esk4_0)=k1_xboole_0|k3_xboole_0(esk3_0,esk1_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_120, c_0_53])).
cnf(c_0_127, negated_conjecture, (~r1_tarski(esk1_0,k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_121]), c_0_73])).
cnf(c_0_128, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,k2_xboole_0(esk2_0,X1)))), inference(spm,[status(thm)],[c_0_122, c_0_29])).
cnf(c_0_129, negated_conjecture, (esk1_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_124])])).
cnf(c_0_130, negated_conjecture, (k3_xboole_0(esk3_0,esk1_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_68])]), c_0_127])).
cnf(c_0_131, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk3_0,k2_xboole_0(esk2_0,X1)))), inference(rw,[status(thm)],[c_0_128, c_0_129])).
cnf(c_0_132, negated_conjecture, (k1_xboole_0!=esk3_0), inference(rw,[status(thm)],[c_0_130, c_0_119])).
cnf(c_0_133, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(k2_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_81, c_0_25])).
cnf(c_0_134, negated_conjecture, (r1_tarski(esk4_0,k2_xboole_0(esk2_0,X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_131]), c_0_132])).
cnf(c_0_135, negated_conjecture, (esk1_0!=esk3_0|esk2_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_136, negated_conjecture, (esk2_0=esk4_0|~r1_tarski(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_133, c_0_88])).
cnf(c_0_137, negated_conjecture, (r1_tarski(esk4_0,esk2_0)), inference(spm,[status(thm)],[c_0_134, c_0_118])).
cnf(c_0_138, negated_conjecture, (esk2_0!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_135, c_0_129])])).
cnf(c_0_139, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_136, c_0_137])]), c_0_138]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 140
# Proof object clause steps            : 107
# Proof object formula steps           : 33
# Proof object conjectures             : 57
# Proof object clause conjectures      : 54
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 74
# Proof object simplifying inferences  : 46
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 1726
# ...of these trivial                  : 341
# ...subsumed                          : 853
# ...remaining for further processing  : 532
# Other redundant clauses eliminated   : 19
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 302
# Generated clauses                    : 10442
# ...of the previous two non-trivial   : 6686
# Contextual simplify-reflections      : 1
# Paramodulations                      : 10420
# Factorizations                       : 0
# Equation resolutions                 : 19
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 181
#    Positive orientable unit clauses  : 82
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 87
# Current number of unprocessed clauses: 4581
# ...number of literals in the above   : 9151
# Current number of archived formulas  : 0
# Current number of archived clauses   : 341
# Clause-clause subsumption calls (NU) : 5657
# Rec. Clause-clause subsumption calls : 4079
# Non-unit clause-clause subsumptions  : 534
# Unit Clause-clause subsumption calls : 1420
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 499
# BW rewrite match successes           : 111
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 108211

# -------------------------------------------------
# User time                : 0.068 s
# System time              : 0.003 s
# Total time               : 0.071 s
# Maximum resident set size: 28648 pages
