# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/c392b3e5149a4f80bc869e670073a36a.E_input', d5_xboole_0)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/c392b3e5149a4f80bc869e670073a36a.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/c392b3e5149a4f80bc869e670073a36a.E_input', t70_enumset1)).
fof(l35_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/c392b3e5149a4f80bc869e670073a36a.E_input', l35_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/c392b3e5149a4f80bc869e670073a36a.E_input', t69_enumset1)).
fof(t75_zfmisc_1, conjecture, ![X1, X2, X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))), file('data/MPTP/proofs/c392b3e5149a4f80bc869e670073a36a.E_input', t75_zfmisc_1)).
fof(t41_zfmisc_1, axiom, ![X1, X2]:~((~(X1=k1_tarski(X2))&(~(X1=k1_xboole_0)&![X3]:~((r2_hidden(X3,X1)&~(X3=X2)))))), file('data/MPTP/proofs/c392b3e5149a4f80bc869e670073a36a.E_input', t41_zfmisc_1)).
fof(c_0_7, plain, ![X25, X26, X27, X28, X29, X30, X31, X32]:((((r2_hidden(X28,X25)|~r2_hidden(X28,X27)|X27!=k4_xboole_0(X25,X26))&(~r2_hidden(X28,X26)|~r2_hidden(X28,X27)|X27!=k4_xboole_0(X25,X26)))&(~r2_hidden(X29,X25)|r2_hidden(X29,X26)|r2_hidden(X29,X27)|X27!=k4_xboole_0(X25,X26)))&((~r2_hidden(esk6_3(X30,X31,X32),X32)|(~r2_hidden(esk6_3(X30,X31,X32),X30)|r2_hidden(esk6_3(X30,X31,X32),X31))|X32=k4_xboole_0(X30,X31))&((r2_hidden(esk6_3(X30,X31,X32),X30)|r2_hidden(esk6_3(X30,X31,X32),X32)|X32=k4_xboole_0(X30,X31))&(~r2_hidden(esk6_3(X30,X31,X32),X31)|r2_hidden(esk6_3(X30,X31,X32),X32)|X32=k4_xboole_0(X30,X31))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_8, plain, ![X9, X10, X11, X12, X13, X14, X15, X16]:(((~r2_hidden(X12,X11)|(X12=X9|X12=X10)|X11!=k2_tarski(X9,X10))&((X13!=X9|r2_hidden(X13,X11)|X11!=k2_tarski(X9,X10))&(X13!=X10|r2_hidden(X13,X11)|X11!=k2_tarski(X9,X10))))&(((esk4_3(X14,X15,X16)!=X14|~r2_hidden(esk4_3(X14,X15,X16),X16)|X16=k2_tarski(X14,X15))&(esk4_3(X14,X15,X16)!=X15|~r2_hidden(esk4_3(X14,X15,X16),X16)|X16=k2_tarski(X14,X15)))&(r2_hidden(esk4_3(X14,X15,X16),X16)|(esk4_3(X14,X15,X16)=X14|esk4_3(X14,X15,X16)=X15)|X16=k2_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_9, plain, ![X18, X19]:k1_enumset1(X18,X18,X19)=k2_tarski(X18,X19), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_10, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk6_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r2_hidden(esk6_3(X1,X2,X3),X1)|r2_hidden(esk6_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_13, plain, ![X20, X21]:((k4_xboole_0(k1_tarski(X20),X21)!=k1_xboole_0|r2_hidden(X20,X21))&(~r2_hidden(X20,X21)|k4_xboole_0(k1_tarski(X20),X21)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l35_zfmisc_1])])).
fof(c_0_14, plain, ![X8]:k2_tarski(X8,X8)=k1_tarski(X8), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_15, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_16, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk6_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_19, plain, (k4_xboole_0(k1_tarski(X1),X2)=k1_xboole_0|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22, plain, (k4_xboole_0(X1,X2)=k4_xboole_0(X3,X3)|~r2_hidden(esk6_3(X3,X3,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, plain, (k4_xboole_0(k1_enumset1(X1,X1,X1),X2)=k1_xboole_0|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_19, c_0_20]), c_0_16])).
cnf(c_0_24, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X3,X1)), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (k4_xboole_0(X1,X1)=k1_xboole_0|~r2_hidden(esk6_3(X1,X1,k1_xboole_0),X2)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[c_0_24])).
fof(c_0_27, negated_conjecture, ~(![X1, X2, X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3)))))))), inference(assume_negation,[status(cth)],[t75_zfmisc_1])).
cnf(c_0_28, plain, (k4_xboole_0(X1,X1)=k1_xboole_0|~r2_hidden(X2,k1_enumset1(X3,X3,esk6_3(X1,X1,k1_xboole_0)))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_29, negated_conjecture, (((esk1_0!=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk2_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&(esk1_0!=k2_tarski(esk2_0,esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0))))&(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0|(esk1_0=k1_xboole_0|(esk1_0=k1_tarski(esk2_0)|(esk1_0=k1_tarski(esk3_0)|esk1_0=k2_tarski(esk2_0,esk3_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_27])])])])])).
cnf(c_0_30, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_31, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_28, c_0_26])).
cnf(c_0_32, negated_conjecture, (esk1_0!=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33, plain, (r2_hidden(X1,X2)|X3!=k1_xboole_0|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34, negated_conjecture, (esk1_0!=k1_xboole_0|k4_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_32, c_0_16])).
cnf(c_0_35, plain, (k4_xboole_0(X1,X2)=X1|r2_hidden(esk6_3(X1,X2,X1),X1)), inference(ef,[status(thm)],[c_0_12])).
fof(c_0_36, plain, ![X22, X23]:((r2_hidden(esk5_2(X22,X23),X22)|X22=k1_xboole_0|X22=k1_tarski(X23))&(esk5_2(X22,X23)!=X23|X22=k1_xboole_0|X22=k1_tarski(X23))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t41_zfmisc_1])])])])])).
cnf(c_0_37, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_31]), c_0_33])).
cnf(c_0_38, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)|~r2_hidden(X2,X3)|~r2_hidden(X4,X3)), inference(spm,[status(thm)],[c_0_10, c_0_23])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk6_3(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0),esk1_0),esk1_0)|k1_xboole_0!=esk1_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (r2_hidden(esk5_2(X1,X2),X1)|X1=k1_xboole_0|X1=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41, plain, (X1=k4_xboole_0(X2,X3)|r2_hidden(esk6_3(X2,X3,X1),X2)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_37, c_0_12])).
cnf(c_0_42, negated_conjecture, (k1_xboole_0!=esk1_0|~r2_hidden(esk6_3(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0),esk1_0),X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43, plain, (X1=k1_xboole_0|X1=k1_enumset1(X2,X2,X2)|r2_hidden(esk5_2(X1,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_20]), c_0_16])).
cnf(c_0_44, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk6_3(X1,X2,k1_xboole_0),X1)), inference(er,[status(thm)],[c_0_41])).
cnf(c_0_45, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_46, negated_conjecture, (k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0|esk1_0=k1_xboole_0|esk1_0=k1_tarski(esk2_0)|esk1_0=k1_tarski(esk3_0)|esk1_0=k2_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_47, negated_conjecture, (k1_xboole_0!=esk1_0|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_42, c_0_39])).
cnf(c_0_48, negated_conjecture, (esk1_0!=k1_tarski(esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_49, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|X1=k1_xboole_0|r2_hidden(esk5_2(X1,X3),X1)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_23, c_0_43])).
cnf(c_0_50, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_37, c_0_44])).
cnf(c_0_51, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_52, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_53, negated_conjecture, (esk1_0=k1_xboole_0|esk1_0=k1_enumset1(esk2_0,esk2_0,esk2_0)|esk1_0=k1_enumset1(esk2_0,esk2_0,esk3_0)|esk1_0=k1_enumset1(esk3_0,esk3_0,esk3_0)|k4_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_20]), c_0_20]), c_0_16]), c_0_16]), c_0_16]), c_0_16])).
cnf(c_0_54, negated_conjecture, (k1_xboole_0!=esk1_0), inference(spm,[status(thm)],[c_0_47, c_0_39])).
cnf(c_0_55, negated_conjecture, (esk1_0!=k2_tarski(esk2_0,esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_56, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|esk4_3(X1,X2,X3)=X1|esk4_3(X1,X2,X3)=X2|X3=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_57, negated_conjecture, (esk1_0!=k1_enumset1(esk3_0,esk3_0,esk3_0)|k4_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_20]), c_0_16]), c_0_16])).
cnf(c_0_58, plain, (k4_xboole_0(X1,k1_enumset1(X2,X2,X3))=k1_xboole_0|r2_hidden(esk5_2(X1,X3),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_26]), c_0_50])).
cnf(c_0_59, plain, (X1=X4|X1=X3|X2!=k1_enumset1(X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_51, c_0_16])).
cnf(c_0_60, plain, (r2_hidden(X1,X2)|k4_xboole_0(X3,X2)!=k1_xboole_0|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_37, c_0_52])).
cnf(c_0_61, negated_conjecture, (k4_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))=k1_xboole_0|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0), inference(sr,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, negated_conjecture, (esk1_0!=k1_enumset1(esk2_0,esk2_0,esk3_0)|k4_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_16]), c_0_16])).
cnf(c_0_63, plain, (X3=k1_enumset1(X1,X1,X2)|esk4_3(X1,X2,X3)=X2|esk4_3(X1,X2,X3)=X1|r2_hidden(esk4_3(X1,X2,X3),X3)), inference(rw,[status(thm)],[c_0_56, c_0_16])).
cnf(c_0_64, negated_conjecture, (r2_hidden(esk5_2(esk1_0,esk3_0),esk1_0)|k1_enumset1(esk3_0,esk3_0,esk3_0)!=esk1_0), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_65, plain, (X1=X2|X1=X3|~r2_hidden(X1,k1_enumset1(X2,X2,X3))), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_66, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|r2_hidden(X1,k1_enumset1(esk2_0,esk2_0,esk3_0))|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_67, negated_conjecture, (esk4_3(esk2_0,esk3_0,X1)=esk2_0|esk4_3(esk2_0,esk3_0,X1)=esk3_0|r2_hidden(esk4_3(esk2_0,esk3_0,X1),X1)|k4_xboole_0(esk1_0,X1)!=k1_xboole_0|X1!=esk1_0), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_68, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk5_2(esk1_0,esk3_0),esk1_0)|r2_hidden(esk5_2(X1,esk3_0),X1)|X1!=esk1_0), inference(spm,[status(thm)],[c_0_64, c_0_43])).
cnf(c_0_69, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|X1=esk3_0|X1=esk2_0|~r2_hidden(X1,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_59])).
cnf(c_0_70, negated_conjecture, (esk4_3(esk2_0,esk3_0,esk1_0)=esk3_0|esk4_3(esk2_0,esk3_0,esk1_0)=esk2_0|r2_hidden(esk4_3(esk2_0,esk3_0,esk1_0),esk1_0)), inference(spm,[status(thm)],[c_0_67, c_0_31])).
cnf(c_0_71, negated_conjecture, (r2_hidden(esk5_2(esk1_0,esk3_0),esk1_0)), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_68]), c_0_54])).
cnf(c_0_72, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_73, negated_conjecture, (esk4_3(esk2_0,esk3_0,esk1_0)=esk2_0|esk4_3(esk2_0,esk3_0,esk1_0)=esk3_0|k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_74, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|esk5_2(esk1_0,esk3_0)=esk2_0|esk5_2(esk1_0,esk3_0)=esk3_0), inference(spm,[status(thm)],[c_0_69, c_0_71])).
cnf(c_0_75, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X2,X4)), inference(rw,[status(thm)],[c_0_72, c_0_16])).
cnf(c_0_76, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|esk4_3(esk2_0,esk3_0,esk1_0)=esk3_0|esk4_3(esk2_0,esk3_0,esk1_0)=esk2_0|k4_xboole_0(esk1_0,X1)=k1_xboole_0|~r2_hidden(esk3_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_73])).
cnf(c_0_77, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|esk5_2(esk1_0,esk3_0)=esk3_0|esk5_2(esk1_0,esk3_0)=esk2_0|k4_xboole_0(esk1_0,X1)=k1_xboole_0|~r2_hidden(esk3_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_74])).
cnf(c_0_78, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X1,X1,X3)), inference(er,[status(thm)],[c_0_75])).
cnf(c_0_79, negated_conjecture, (esk4_3(esk2_0,esk3_0,esk1_0)=esk2_0|esk4_3(esk2_0,esk3_0,esk1_0)=esk3_0|k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_76]), c_0_26])]), c_0_73])).
cnf(c_0_80, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|esk5_2(esk1_0,esk3_0)=esk2_0|esk5_2(esk1_0,esk3_0)=esk3_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_77]), c_0_26])]), c_0_74])).
cnf(c_0_81, negated_conjecture, (esk1_0!=k1_tarski(esk2_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_82, plain, (r2_hidden(X1,k1_enumset1(X1,X1,X2))), inference(er,[status(thm)],[c_0_78])).
cnf(c_0_83, plain, (X3=k2_tarski(X1,X2)|esk4_3(X1,X2,X3)!=X2|~r2_hidden(esk4_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_84, negated_conjecture, (esk4_3(esk2_0,esk3_0,esk1_0)=esk3_0|esk4_3(esk2_0,esk3_0,esk1_0)=esk2_0|X1=esk2_0|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_65, c_0_79])).
cnf(c_0_85, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|esk5_2(X1,X2)!=X2), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_86, negated_conjecture, (esk5_2(esk1_0,esk3_0)=esk3_0|esk5_2(esk1_0,esk3_0)=esk2_0|X1=esk2_0|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_65, c_0_80])).
cnf(c_0_87, negated_conjecture, (esk1_0!=k1_enumset1(esk2_0,esk2_0,esk2_0)|k4_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_20]), c_0_16]), c_0_16])).
cnf(c_0_88, plain, (k4_xboole_0(X1,k1_enumset1(X2,X2,X3))=k1_xboole_0|r2_hidden(esk5_2(X1,X2),X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_82]), c_0_50])).
cnf(c_0_89, plain, (X3=k1_enumset1(X1,X1,X2)|esk4_3(X1,X2,X3)!=X2|~r2_hidden(esk4_3(X1,X2,X3),X3)), inference(rw,[status(thm)],[c_0_83, c_0_16])).
cnf(c_0_90, negated_conjecture, (esk4_3(esk2_0,esk3_0,esk1_0)=esk2_0|esk4_3(esk2_0,esk3_0,esk1_0)=esk3_0), inference(spm,[status(thm)],[c_0_84, c_0_70])).
cnf(c_0_91, plain, (X1=k1_xboole_0|X1=k1_enumset1(X2,X2,X2)|esk5_2(X1,X2)!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_20]), c_0_16])).
cnf(c_0_92, negated_conjecture, (esk5_2(esk1_0,esk3_0)=esk2_0|esk5_2(esk1_0,esk3_0)=esk3_0), inference(spm,[status(thm)],[c_0_86, c_0_71])).
cnf(c_0_93, negated_conjecture, (r2_hidden(esk5_2(esk1_0,esk2_0),esk1_0)|k1_enumset1(esk2_0,esk2_0,esk2_0)!=esk1_0), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_94, plain, (X3=k2_tarski(X1,X2)|esk4_3(X1,X2,X3)!=X1|~r2_hidden(esk4_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_95, negated_conjecture, (esk4_3(esk2_0,esk3_0,esk1_0)=esk2_0|k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|~r2_hidden(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_96, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|esk5_2(esk1_0,esk3_0)=esk2_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_54])).
cnf(c_0_97, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk5_2(esk1_0,esk2_0),esk1_0)|r2_hidden(esk5_2(X1,esk2_0),X1)|X1!=esk1_0), inference(spm,[status(thm)],[c_0_93, c_0_43])).
cnf(c_0_98, plain, (X3=k1_enumset1(X1,X1,X2)|esk4_3(X1,X2,X3)!=X1|~r2_hidden(esk4_3(X1,X2,X3),X3)), inference(rw,[status(thm)],[c_0_94, c_0_16])).
cnf(c_0_99, negated_conjecture, (esk4_3(esk2_0,esk3_0,esk1_0)=esk2_0|~r2_hidden(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_95]), c_0_31])])).
cnf(c_0_100, negated_conjecture, (esk5_2(esk1_0,esk3_0)=esk2_0|k4_xboole_0(esk1_0,X1)=k1_xboole_0|~r2_hidden(esk3_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_96])).
cnf(c_0_101, negated_conjecture, (r2_hidden(esk5_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_97]), c_0_54])).
cnf(c_0_102, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|~r2_hidden(esk2_0,esk1_0)|~r2_hidden(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_103, negated_conjecture, (esk5_2(esk1_0,esk3_0)=esk2_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_100]), c_0_26])]), c_0_96])).
cnf(c_0_104, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|esk5_2(esk1_0,esk2_0)=esk2_0|esk5_2(esk1_0,esk2_0)=esk3_0), inference(spm,[status(thm)],[c_0_69, c_0_101])).
cnf(c_0_105, negated_conjecture, (~r2_hidden(esk2_0,esk1_0)|~r2_hidden(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_102]), c_0_31])])).
cnf(c_0_106, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(rw,[status(thm)],[c_0_71, c_0_103])).
cnf(c_0_107, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|esk5_2(esk1_0,esk2_0)=esk3_0|esk5_2(esk1_0,esk2_0)=esk2_0|r2_hidden(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_82, c_0_104])).
cnf(c_0_108, negated_conjecture, (~r2_hidden(esk3_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_106])])).
cnf(c_0_109, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|esk5_2(esk1_0,esk2_0)=esk2_0|esk5_2(esk1_0,esk2_0)=esk3_0), inference(sr,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_110, plain, (X1=X2|r2_hidden(esk5_2(X3,X2),X3)|~r2_hidden(X1,X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_43])]), c_0_37])).
cnf(c_0_111, negated_conjecture, (esk5_2(esk1_0,esk2_0)=esk3_0|esk5_2(esk1_0,esk2_0)=esk2_0|k4_xboole_0(esk1_0,X1)=k1_xboole_0|~r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_109])).
cnf(c_0_112, negated_conjecture, (esk5_2(esk1_0,esk2_0)=X1|r2_hidden(esk5_2(esk1_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_110, c_0_101])).
cnf(c_0_113, negated_conjecture, (esk5_2(esk1_0,esk2_0)=esk2_0|esk5_2(esk1_0,esk2_0)=esk3_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_111]), c_0_82])]), c_0_109])).
cnf(c_0_114, negated_conjecture, (esk5_2(esk1_0,esk2_0)=esk2_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_108])).
cnf(c_0_115, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_114]), c_0_54])).
cnf(c_0_116, negated_conjecture, (k4_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_87, c_0_115])])).
cnf(c_0_117, negated_conjecture, (k4_xboole_0(esk1_0,X1)=k1_xboole_0|~r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_115])).
cnf(c_0_118, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_82])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 119
# Proof object clause steps            : 104
# Proof object formula steps           : 15
# Proof object conjectures             : 58
# Proof object clause conjectures      : 55
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 62
# Proof object simplifying inferences  : 65
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 21
# Processed clauses                    : 2858
# ...of these trivial                  : 35
# ...subsumed                          : 2083
# ...remaining for further processing  : 740
# Other redundant clauses eliminated   : 1728
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 166
# Backward-rewritten                   : 141
# Generated clauses                    : 73355
# ...of the previous two non-trivial   : 67883
# Contextual simplify-reflections      : 71
# Paramodulations                      : 71278
# Factorizations                       : 225
# Equation resolutions                 : 1843
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 401
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 374
# Current number of unprocessed clauses: 63974
# ...number of literals in the above   : 465136
# Current number of archived formulas  : 0
# Current number of archived clauses   : 339
# Clause-clause subsumption calls (NU) : 42104
# Rec. Clause-clause subsumption calls : 14514
# Non-unit clause-clause subsumptions  : 1651
# Unit Clause-clause subsumption calls : 894
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 240
# BW rewrite match successes           : 34
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1369388

# -------------------------------------------------
# User time                : 1.305 s
# System time              : 0.017 s
# Total time               : 1.321 s
# Maximum resident set size: 57520 pages
