# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', d5_xboole_0)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', redefinition_k6_subset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', d3_tarski)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', t7_ordinal1)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', d12_waybel_0)).
fof(t9_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>~(r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3)))))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', t9_waybel_0)).
fof(t29_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1)))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', t29_yellow_6)).
fof(dt_k4_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&l1_waybel_0(X2,X1))=>m1_subset_1(k4_yellow_6(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', dt_k4_yellow_6)).
fof(dt_l1_waybel_0, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', dt_l1_waybel_0)).
fof(rc4_waybel_0, axiom, ![X1]:(l1_struct_0(X1)=>?[X2]:(l1_waybel_0(X2,X1)&v6_waybel_0(X2,X1))), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', rc4_waybel_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/c3cf293e60a04f4a9c6685213950aba7.E_input', dt_l1_orders_2)).
fof(c_0_11, plain, ![X32, X33, X34, X35, X36, X37, X38, X39]:((((r2_hidden(X35,X32)|~r2_hidden(X35,X34)|X34!=k4_xboole_0(X32,X33))&(~r2_hidden(X35,X33)|~r2_hidden(X35,X34)|X34!=k4_xboole_0(X32,X33)))&(~r2_hidden(X36,X32)|r2_hidden(X36,X33)|r2_hidden(X36,X34)|X34!=k4_xboole_0(X32,X33)))&((~r2_hidden(esk6_3(X37,X38,X39),X39)|(~r2_hidden(esk6_3(X37,X38,X39),X37)|r2_hidden(esk6_3(X37,X38,X39),X38))|X39=k4_xboole_0(X37,X38))&((r2_hidden(esk6_3(X37,X38,X39),X37)|r2_hidden(esk6_3(X37,X38,X39),X39)|X39=k4_xboole_0(X37,X38))&(~r2_hidden(esk6_3(X37,X38,X39),X38)|r2_hidden(esk6_3(X37,X38,X39),X39)|X39=k4_xboole_0(X37,X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_12, plain, ![X15, X16]:k6_subset_1(X15,X16)=k4_xboole_0(X15,X16), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_13, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_14, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (X3!=k6_subset_1(X4,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_17, plain, ![X9, X10, X11, X12, X13]:((~r1_tarski(X9,X10)|(~r2_hidden(X11,X9)|r2_hidden(X11,X10)))&((r2_hidden(esk3_2(X12,X13),X12)|r1_tarski(X12,X13))&(~r2_hidden(esk3_2(X12,X13),X13)|r1_tarski(X12,X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|X3!=k6_subset_1(X2,X4)|~r2_hidden(X1,X3)), inference(rw,[status(thm)],[c_0_15, c_0_14])).
cnf(c_0_19, plain, (~r2_hidden(X1,k6_subset_1(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k6_subset_1(X2,X3))), inference(er,[status(thm)],[c_0_18])).
fof(c_0_22, plain, ![X28, X29]:(~r2_hidden(X28,X29)|~r1_tarski(X29,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_23, plain, (r1_tarski(k6_subset_1(X1,X2),X3)|~r2_hidden(esk3_2(k6_subset_1(X1,X2),X3),X2)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_24, plain, (r2_hidden(esk3_2(k6_subset_1(X1,X2),X3),X1)|r1_tarski(k6_subset_1(X1,X2),X3)), inference(spm,[status(thm)],[c_0_21, c_0_20])).
cnf(c_0_25, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (r1_tarski(k6_subset_1(X1,X1),X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_27, plain, ![X20, X21, X22, X23, X25, X27]:(((m1_subset_1(esk4_4(X20,X21,X22,X23),u1_struct_0(X21))|~m1_subset_1(X23,u1_struct_0(X21))|~r2_waybel_0(X20,X21,X22)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20)))&((r1_orders_2(X21,X23,esk4_4(X20,X21,X22,X23))|~m1_subset_1(X23,u1_struct_0(X21))|~r2_waybel_0(X20,X21,X22)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20)))&(r2_hidden(k2_waybel_0(X20,X21,esk4_4(X20,X21,X22,X23)),X22)|~m1_subset_1(X23,u1_struct_0(X21))|~r2_waybel_0(X20,X21,X22)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20)))))&((m1_subset_1(esk5_3(X20,X21,X25),u1_struct_0(X21))|r2_waybel_0(X20,X21,X25)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20)))&(~m1_subset_1(X27,u1_struct_0(X21))|(~r1_orders_2(X21,esk5_3(X20,X21,X25),X27)|~r2_hidden(k2_waybel_0(X20,X21,X27),X25))|r2_waybel_0(X20,X21,X25)|(v2_struct_0(X21)|~l1_waybel_0(X21,X20))|(v2_struct_0(X20)|~l1_struct_0(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
cnf(c_0_28, plain, (~r2_hidden(X1,k6_subset_1(X2,X2))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29, plain, (r2_hidden(k2_waybel_0(X1,X2,esk4_4(X1,X2,X3,X4)),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X2))|~r2_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_30, plain, ![X17, X18, X19]:((~r1_waybel_0(X17,X18,X19)|~r2_waybel_0(X17,X18,k6_subset_1(u1_struct_0(X17),X19))|(v2_struct_0(X18)|~l1_waybel_0(X18,X17))|(v2_struct_0(X17)|~l1_struct_0(X17)))&(r2_waybel_0(X17,X18,k6_subset_1(u1_struct_0(X17),X19))|r1_waybel_0(X17,X18,X19)|(v2_struct_0(X18)|~l1_waybel_0(X18,X17))|(v2_struct_0(X17)|~l1_struct_0(X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_waybel_0])])])])])).
fof(c_0_31, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>r1_waybel_0(X1,X2,u1_struct_0(X1))))), inference(assume_negation,[status(cth)],[t29_yellow_6])).
cnf(c_0_32, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_waybel_0(X1,X2,k6_subset_1(X4,X4))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33, plain, (r2_waybel_0(X1,X2,k6_subset_1(u1_struct_0(X1),X3))|r1_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_34, plain, ![X30, X31]:(v2_struct_0(X30)|~l1_struct_0(X30)|~l1_waybel_0(X31,X30)|m1_subset_1(k4_yellow_6(X30,X31),u1_struct_0(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_yellow_6])])])).
fof(c_0_35, plain, ![X41, X42]:(~l1_struct_0(X41)|(~l1_waybel_0(X42,X41)|l1_orders_2(X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])).
fof(c_0_36, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_31])])])])).
cnf(c_0_37, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_38, plain, (v2_struct_0(X1)|m1_subset_1(k4_yellow_6(X1,X2),u1_struct_0(X1))|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_39, plain, ![X43]:((l1_waybel_0(esk7_1(X43),X43)|~l1_struct_0(X43))&(v6_waybel_0(esk7_1(X43),X43)|~l1_struct_0(X43))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[rc4_waybel_0])])])])).
fof(c_0_40, plain, ![X8]:(~l1_orders_2(X8)|l1_struct_0(X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_41, plain, (l1_orders_2(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_43, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~l1_waybel_0(X2,X1)|~l1_waybel_0(X3,X2)|~l1_struct_0(X1)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (l1_waybel_0(esk7_1(X1),X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_47, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])])).
cnf(c_0_48, negated_conjecture, (~r1_waybel_0(esk1_0,esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, plain, (r1_waybel_0(X1,X2,u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)|~l1_struct_0(X2)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, negated_conjecture, (l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_52, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_53, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_42]), c_0_43]), c_0_50])]), c_0_51]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 54
# Proof object clause steps            : 31
# Proof object formula steps           : 23
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 93
# ...of these trivial                  : 0
# ...subsumed                          : 9
# ...remaining for further processing  : 84
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 2
# Generated clauses                    : 85
# ...of the previous two non-trivial   : 70
# Contextual simplify-reflections      : 0
# Paramodulations                      : 80
# Factorizations                       : 2
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 49
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 35
# ...number of literals in the above   : 127
# Current number of archived formulas  : 0
# Current number of archived clauses   : 33
# Clause-clause subsumption calls (NU) : 959
# Rec. Clause-clause subsumption calls : 152
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 45
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3475

# -------------------------------------------------
# User time                : 0.003 s
# System time              : 0.008 s
# Total time               : 0.011 s
# Maximum resident set size: 28760 pages
