# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t38_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r2_hidden(X2,X3)=>(r3_lattices(X1,X2,k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),X2))))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', t38_lattice3)).
fof(redefinition_r3_lattices, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v6_lattices(X1)&(v8_lattices(X1)&(v9_lattices(X1)&l3_lattices(X1)))))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_lattices(X1,X2,X3)<=>r1_lattices(X1,X2,X3))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', redefinition_r3_lattices)).
fof(cc1_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v10_lattices(X1))=>(~(v2_struct_0(X1))&(v4_lattices(X1)&(v5_lattices(X1)&(v6_lattices(X1)&(v7_lattices(X1)&(v8_lattices(X1)&v9_lattices(X1))))))))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', cc1_lattices)).
fof(t34_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(X2=k16_lattice3(X1,X3)<=>(r3_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_lattice3(X1,X4,X3)=>r3_lattices(X1,X4,X2))))))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', t34_lattice3)).
fof(dt_k16_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', dt_k16_lattice3)).
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', d21_lattice3)).
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', d16_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', dt_k15_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/c3f710a11aa54769869fcca064923622.E_input', d17_lattice3)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r2_hidden(X2,X3)=>(r3_lattices(X1,X2,k15_lattice3(X1,X3))&r3_lattices(X1,k16_lattice3(X1,X3),X2)))))), inference(assume_negation,[status(cth)],[t38_lattice3])).
fof(c_0_10, plain, ![X33, X34, X35]:((~r3_lattices(X33,X34,X35)|r1_lattices(X33,X34,X35)|(v2_struct_0(X33)|(~v6_lattices(X33)|(~v8_lattices(X33)|(~v9_lattices(X33)|~l3_lattices(X33))))|(~m1_subset_1(X34,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33)))))&(~r1_lattices(X33,X34,X35)|r3_lattices(X33,X34,X35)|(v2_struct_0(X33)|(~v6_lattices(X33)|(~v8_lattices(X33)|(~v9_lattices(X33)|~l3_lattices(X33))))|(~m1_subset_1(X34,u1_struct_0(X33))|~m1_subset_1(X35,u1_struct_0(X33)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_lattices])])])])).
fof(c_0_11, plain, ![X36]:((~v2_struct_0(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v4_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v5_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v6_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v7_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&((v8_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36))&(v9_lattices(X36)|(v2_struct_0(X36)|~v10_lattices(X36))|~l3_lattices(X36)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_lattices])])])])).
fof(c_0_12, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&(v4_lattice3(esk1_0)&l3_lattices(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(r2_hidden(esk2_0,esk3_0)&(~r3_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~r3_lattices(esk1_0,k16_lattice3(esk1_0,esk3_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
cnf(c_0_13, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~v6_lattices(X1)|~v8_lattices(X1)|~v9_lattices(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (v9_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, plain, (v6_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (v8_lattices(X1)|v2_struct_0(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, plain, ![X16, X17, X18, X19, X20]:(((r3_lattice3(X16,X17,X18)|X17!=k16_lattice3(X16,X18)|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16)))))&(~m1_subset_1(X19,u1_struct_0(X16))|(~r3_lattice3(X16,X19,X18)|r3_lattices(X16,X19,X17))|X17!=k16_lattice3(X16,X18)|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16))))))&((m1_subset_1(esk5_3(X16,X17,X20),u1_struct_0(X16))|~r3_lattice3(X16,X17,X20)|X17=k16_lattice3(X16,X20)|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16)))))&((r3_lattice3(X16,esk5_3(X16,X17,X20),X20)|~r3_lattice3(X16,X17,X20)|X17=k16_lattice3(X16,X20)|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16)))))&(~r3_lattices(X16,esk5_3(X16,X17,X20),X17)|~r3_lattice3(X16,X17,X20)|X17=k16_lattice3(X16,X20)|~m1_subset_1(X17,u1_struct_0(X16))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v4_lattice3(X16)|~l3_lattices(X16)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_lattice3])])])])])])])).
fof(c_0_18, plain, ![X37, X38]:(v2_struct_0(X37)|~l3_lattices(X37)|m1_subset_1(k16_lattice3(X37,X38),u1_struct_0(X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k16_lattice3])])])).
fof(c_0_19, plain, ![X22, X23, X24, X25]:(((r4_lattice3(X22,X24,X23)|X24!=k15_lattice3(X22,X23)|~m1_subset_1(X24,u1_struct_0(X22))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22))))|(v2_struct_0(X22)|~l3_lattices(X22)))&(~m1_subset_1(X25,u1_struct_0(X22))|(~r4_lattice3(X22,X25,X23)|r1_lattices(X22,X24,X25))|X24!=k15_lattice3(X22,X23)|~m1_subset_1(X24,u1_struct_0(X22))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22))))|(v2_struct_0(X22)|~l3_lattices(X22))))&((m1_subset_1(esk6_3(X22,X23,X24),u1_struct_0(X22))|~r4_lattice3(X22,X24,X23)|X24=k15_lattice3(X22,X23)|~m1_subset_1(X24,u1_struct_0(X22))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22))))|(v2_struct_0(X22)|~l3_lattices(X22)))&((r4_lattice3(X22,esk6_3(X22,X23,X24),X23)|~r4_lattice3(X22,X24,X23)|X24=k15_lattice3(X22,X23)|~m1_subset_1(X24,u1_struct_0(X22))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22))))|(v2_struct_0(X22)|~l3_lattices(X22)))&(~r1_lattices(X22,X24,esk6_3(X22,X23,X24))|~r4_lattice3(X22,X24,X23)|X24=k15_lattice3(X22,X23)|~m1_subset_1(X24,u1_struct_0(X22))|(v2_struct_0(X22)|(~v10_lattices(X22)|(~v4_lattice3(X22)|~l3_lattices(X22))))|(v2_struct_0(X22)|~l3_lattices(X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
cnf(c_0_20, negated_conjecture, (~r3_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~r3_lattices(esk1_0,k16_lattice3(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (r3_lattices(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_25, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_26, plain, ![X27, X28, X29, X30, X31]:((~r3_lattice3(X27,X28,X29)|(~m1_subset_1(X30,u1_struct_0(X27))|(~r2_hidden(X30,X29)|r1_lattices(X27,X28,X30)))|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l3_lattices(X27)))&((m1_subset_1(esk7_3(X27,X28,X31),u1_struct_0(X27))|r3_lattice3(X27,X28,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l3_lattices(X27)))&((r2_hidden(esk7_3(X27,X28,X31),X31)|r3_lattice3(X27,X28,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l3_lattices(X27)))&(~r1_lattices(X27,X28,esk7_3(X27,X28,X31))|r3_lattice3(X27,X28,X31)|~m1_subset_1(X28,u1_struct_0(X27))|(v2_struct_0(X27)|~l3_lattices(X27)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
cnf(c_0_27, plain, (r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|X2!=k16_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (v2_struct_0(X1)|m1_subset_1(k16_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X1)|X2!=k15_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_30, plain, ![X14, X15]:(v2_struct_0(X14)|~l3_lattices(X14)|m1_subset_1(k15_lattice3(X14,X15),u1_struct_0(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
cnf(c_0_31, negated_conjecture, (~r1_lattices(esk1_0,k16_lattice3(esk1_0,esk3_0),esk2_0)|~r3_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~m1_subset_1(k16_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23]), c_0_24])]), c_0_25])).
cnf(c_0_32, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (r3_lattice3(X1,k16_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_27]), c_0_28])).
fof(c_0_34, plain, ![X8, X9, X10, X11, X12]:((~r4_lattice3(X8,X9,X10)|(~m1_subset_1(X11,u1_struct_0(X8))|(~r2_hidden(X11,X10)|r1_lattices(X8,X11,X9)))|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l3_lattices(X8)))&((m1_subset_1(esk4_3(X8,X9,X12),u1_struct_0(X8))|r4_lattice3(X8,X9,X12)|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l3_lattices(X8)))&((r2_hidden(esk4_3(X8,X9,X12),X12)|r4_lattice3(X8,X9,X12)|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l3_lattices(X8)))&(~r1_lattices(X8,esk4_3(X8,X9,X12),X9)|r4_lattice3(X8,X9,X12)|~m1_subset_1(X9,u1_struct_0(X8))|(v2_struct_0(X8)|~l3_lattices(X8)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_35, plain, (v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, negated_conjecture, (~r1_lattices(esk1_0,k16_lattice3(esk1_0,esk3_0),esk2_0)|~r1_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~m1_subset_1(k16_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_21]), c_0_22]), c_0_23]), c_0_24])]), c_0_25])).
cnf(c_0_38, plain, (r1_lattices(X1,k16_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r2_hidden(X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_28])).
cnf(c_0_39, negated_conjecture, (r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_40, negated_conjecture, (v4_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_41, plain, (r1_lattices(X1,X4,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_35]), c_0_36])).
cnf(c_0_43, negated_conjecture, (~r1_lattices(esk1_0,esk2_0,k15_lattice3(esk1_0,esk3_0))|~m1_subset_1(k16_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_22]), c_0_23]), c_0_40]), c_0_24])]), c_0_25])).
cnf(c_0_44, plain, (r1_lattices(X1,X2,k15_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)|~v4_lattice3(X1)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_36])).
cnf(c_0_45, negated_conjecture, (~m1_subset_1(k16_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_39]), c_0_22]), c_0_23]), c_0_40]), c_0_24])]), c_0_25])).
cnf(c_0_46, negated_conjecture, (~m1_subset_1(k15_lattice3(esk1_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_28]), c_0_23])]), c_0_25])).
cnf(c_0_47, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_36]), c_0_23])]), c_0_25]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 29
# Proof object formula steps           : 19
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 39
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 36
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 35
# Processed clauses                    : 82
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 82
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 24
# ...of the previous two non-trivial   : 22
# Contextual simplify-reflections      : 8
# Paramodulations                      : 20
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 34
# Current number of unprocessed clauses: 10
# ...number of literals in the above   : 93
# Current number of archived formulas  : 0
# Current number of archived clauses   : 37
# Clause-clause subsumption calls (NU) : 1339
# Rec. Clause-clause subsumption calls : 52
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4539

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28684 pages
