# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/c4c2515d33df4e32aae30b1a5cd2261b.E_input', t1_xboole_1)).
fof(t119_zfmisc_1, axiom, ![X1, X2, X3, X4]:((r1_tarski(X1,X2)&r1_tarski(X3,X4))=>r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), file('data/MPTP/proofs/c4c2515d33df4e32aae30b1a5cd2261b.E_input', t119_zfmisc_1)).
fof(t32_finset_1, axiom, ![X1, X2, X3]:~((v1_finset_1(X1)&(r1_tarski(X1,k2_zfmisc_1(X2,X3))&![X4, X5]:~((v1_finset_1(X4)&(r1_tarski(X4,X2)&(v1_finset_1(X5)&(r1_tarski(X5,X3)&r1_tarski(X1,k2_zfmisc_1(X4,X5)))))))))), file('data/MPTP/proofs/c4c2515d33df4e32aae30b1a5cd2261b.E_input', t32_finset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/c4c2515d33df4e32aae30b1a5cd2261b.E_input', d10_xboole_0)).
fof(t33_finset_1, conjecture, ![X1, X2, X3]:~((v1_finset_1(X1)&(r1_tarski(X1,k2_zfmisc_1(X2,X3))&![X4]:~((v1_finset_1(X4)&(r1_tarski(X4,X2)&r1_tarski(X1,k2_zfmisc_1(X4,X3)))))))), file('data/MPTP/proofs/c4c2515d33df4e32aae30b1a5cd2261b.E_input', t33_finset_1)).
fof(c_0_5, plain, ![X6, X7, X8]:(~r1_tarski(X6,X7)|~r1_tarski(X7,X8)|r1_tarski(X6,X8)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_6, plain, ![X9, X10, X11, X12]:(~r1_tarski(X9,X10)|~r1_tarski(X11,X12)|r1_tarski(k2_zfmisc_1(X9,X11),k2_zfmisc_1(X10,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t119_zfmisc_1])])).
cnf(c_0_7, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (r1_tarski(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))|~r1_tarski(X1,X2)|~r1_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, plain, ![X13, X14, X15]:((v1_finset_1(esk1_3(X13,X14,X15))|~r1_tarski(X13,k2_zfmisc_1(X14,X15))|~v1_finset_1(X13))&((r1_tarski(esk1_3(X13,X14,X15),X14)|~r1_tarski(X13,k2_zfmisc_1(X14,X15))|~v1_finset_1(X13))&((v1_finset_1(esk2_3(X13,X14,X15))|~r1_tarski(X13,k2_zfmisc_1(X14,X15))|~v1_finset_1(X13))&((r1_tarski(esk2_3(X13,X14,X15),X15)|~r1_tarski(X13,k2_zfmisc_1(X14,X15))|~v1_finset_1(X13))&(r1_tarski(X13,k2_zfmisc_1(esk1_3(X13,X14,X15),esk2_3(X13,X14,X15)))|~r1_tarski(X13,k2_zfmisc_1(X14,X15))|~v1_finset_1(X13)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_finset_1])])])])])])).
cnf(c_0_10, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(X1,k2_zfmisc_1(X4,X5))|~r1_tarski(X5,X3)|~r1_tarski(X4,X2)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_11, plain, (r1_tarski(X1,k2_zfmisc_1(esk1_3(X1,X2,X3),esk2_3(X1,X2,X3)))|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, plain, ![X18, X19]:(((r1_tarski(X18,X19)|X18!=X19)&(r1_tarski(X19,X18)|X18!=X19))&(~r1_tarski(X18,X19)|~r1_tarski(X19,X18)|X18=X19)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_13, negated_conjecture, ~(![X1, X2, X3]:~((v1_finset_1(X1)&(r1_tarski(X1,k2_zfmisc_1(X2,X3))&![X4]:~((v1_finset_1(X4)&(r1_tarski(X4,X2)&r1_tarski(X1,k2_zfmisc_1(X4,X3))))))))), inference(assume_negation,[status(cth)],[t33_finset_1])).
cnf(c_0_14, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_finset_1(X1)|~r1_tarski(esk2_3(X1,X4,X5),X3)|~r1_tarski(esk1_3(X1,X4,X5),X2)|~r1_tarski(X1,k2_zfmisc_1(X4,X5))), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (r1_tarski(esk2_3(X1,X2,X3),X3)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, negated_conjecture, ![X23]:(v1_finset_1(esk3_0)&(r1_tarski(esk3_0,k2_zfmisc_1(esk4_0,esk5_0))&(~v1_finset_1(X23)|(~r1_tarski(X23,esk4_0)|~r1_tarski(esk3_0,k2_zfmisc_1(X23,esk5_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_13])])])])).
cnf(c_0_18, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_finset_1(X1)|~r1_tarski(esk1_3(X1,X4,X3),X2)|~r1_tarski(X1,k2_zfmisc_1(X4,X3))), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_20, negated_conjecture, (~v1_finset_1(X1)|~r1_tarski(X1,esk4_0)|~r1_tarski(esk3_0,k2_zfmisc_1(X1,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (r1_tarski(X1,k2_zfmisc_1(esk1_3(X1,X2,X3),X3))|~v1_finset_1(X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, negated_conjecture, (v1_finset_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (~v1_finset_1(esk1_3(esk3_0,X1,esk5_0))|~r1_tarski(esk1_3(esk3_0,X1,esk5_0),esk4_0)|~r1_tarski(esk3_0,k2_zfmisc_1(X1,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22])])).
cnf(c_0_24, plain, (v1_finset_1(esk1_3(X1,X2,X3))|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, negated_conjecture, (~r1_tarski(esk1_3(esk3_0,X1,esk5_0),esk4_0)|~r1_tarski(esk3_0,k2_zfmisc_1(X1,esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_22])])).
cnf(c_0_26, plain, (r1_tarski(esk1_3(X1,X2,X3),X2)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_27, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk4_0,esk5_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), c_0_22])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 18
# Proof object formula steps           : 11
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 13
# Processed clauses                    : 164
# ...of these trivial                  : 0
# ...subsumed                          : 41
# ...remaining for further processing  : 123
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 608
# ...of the previous two non-trivial   : 603
# Contextual simplify-reflections      : 2
# Paramodulations                      : 606
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 106
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 102
# Current number of unprocessed clauses: 464
# ...number of literals in the above   : 2692
# Current number of archived formulas  : 0
# Current number of archived clauses   : 15
# Clause-clause subsumption calls (NU) : 4596
# Rec. Clause-clause subsumption calls : 1500
# Non-unit clause-clause subsumptions  : 45
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14271

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.003 s
# Total time               : 0.019 s
# Maximum resident set size: 28672 pages
