# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t20_mcart_1, conjecture, ![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1)))), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t20_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t7_mcart_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t71_enumset1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t152_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t113_zfmisc_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t69_enumset1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', l54_zfmisc_1)).
fof(t35_zfmisc_1, axiom, ![X1, X2]:k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2)), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t35_zfmisc_1)).
fof(t9_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(((r2_hidden(X3,X1)|r2_hidden(X4,X1))&X2=k4_tarski(X3,X4))))))), file('data/MPTP/proofs/c5554555c06b424fa698292ce9daf30b.E_input', t9_mcart_1)).
fof(c_0_12, negated_conjecture, ~(![X1]:(?[X2, X3]:X1=k4_tarski(X2,X3)=>(~(X1=k1_mcart_1(X1))&~(X1=k2_mcart_1(X1))))), inference(assume_negation,[status(cth)],[t20_mcart_1])).
fof(c_0_13, plain, ![X44, X45]:(k1_mcart_1(k4_tarski(X44,X45))=X44&k2_mcart_1(k4_tarski(X44,X45))=X45), inference(variable_rename,[status(thm)],[t7_mcart_1])).
fof(c_0_14, negated_conjecture, (esk1_0=k4_tarski(esk2_0,esk3_0)&(esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])).
cnf(c_0_15, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (esk1_0=k4_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (esk2_0=k1_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
fof(c_0_18, plain, ![X14, X15, X16, X17, X18, X19, X20, X21]:(((~r2_hidden(X17,X16)|(X17=X14|X17=X15)|X16!=k2_tarski(X14,X15))&((X18!=X14|r2_hidden(X18,X16)|X16!=k2_tarski(X14,X15))&(X18!=X15|r2_hidden(X18,X16)|X16!=k2_tarski(X14,X15))))&(((esk4_3(X19,X20,X21)!=X19|~r2_hidden(esk4_3(X19,X20,X21),X21)|X21=k2_tarski(X19,X20))&(esk4_3(X19,X20,X21)!=X20|~r2_hidden(esk4_3(X19,X20,X21),X21)|X21=k2_tarski(X19,X20)))&(r2_hidden(esk4_3(X19,X20,X21),X21)|(esk4_3(X19,X20,X21)=X19|esk4_3(X19,X20,X21)=X20)|X21=k2_tarski(X19,X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_19, plain, ![X31, X32]:k1_enumset1(X31,X31,X32)=k2_tarski(X31,X32), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_20, plain, ![X10, X11, X12]:k2_enumset1(X10,X10,X11,X12)=k1_enumset1(X10,X11,X12), inference(variable_rename,[status(thm)],[t71_enumset1])).
cnf(c_0_21, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),esk3_0)=esk1_0), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_23, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X29, X30]:~r2_hidden(X29,k2_zfmisc_1(X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
fof(c_0_27, plain, ![X27, X28]:((k2_zfmisc_1(X27,X28)!=k1_xboole_0|(X27=k1_xboole_0|X28=k1_xboole_0))&((X27!=k1_xboole_0|k2_zfmisc_1(X27,X28)=k1_xboole_0)&(X28!=k1_xboole_0|k2_zfmisc_1(X27,X28)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_28, plain, ![X37, X38, X39, X40, X41, X42]:(((~r2_hidden(X39,X38)|X39=X37|X38!=k1_tarski(X37))&(X40!=X37|r2_hidden(X40,X38)|X38!=k1_tarski(X37)))&((~r2_hidden(esk6_2(X41,X42),X42)|esk6_2(X41,X42)!=X41|X42=k1_tarski(X41))&(r2_hidden(esk6_2(X41,X42),X42)|esk6_2(X41,X42)=X41|X42=k1_tarski(X41)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_29, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_30, negated_conjecture, (esk3_0=k2_mcart_1(esk1_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_31, plain, ![X23, X24, X25, X26]:(((r2_hidden(X23,X25)|~r2_hidden(k4_tarski(X23,X24),k2_zfmisc_1(X25,X26)))&(r2_hidden(X24,X26)|~r2_hidden(k4_tarski(X23,X24),k2_zfmisc_1(X25,X26))))&(~r2_hidden(X23,X25)|~r2_hidden(X24,X26)|r2_hidden(k4_tarski(X23,X24),k2_zfmisc_1(X25,X26)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_32, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X2,X2,X2,X4)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_33, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (r2_hidden(esk6_2(X1,X2),X2)|esk6_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),k2_mcart_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_22, c_0_30])).
cnf(c_0_38, negated_conjecture, (esk1_0=k1_mcart_1(esk1_0)|esk1_0=k2_mcart_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_39, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_40, plain, (r2_hidden(X1,X2)|X2!=k2_enumset1(X1,X1,X1,X3)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_41, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_42, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_43, plain, (esk6_2(X1,X2)=X1|X2=k2_enumset1(X1,X1,X1,X1)|r2_hidden(esk6_2(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36]), c_0_24]), c_0_25])).
cnf(c_0_44, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_45, negated_conjecture, (k4_tarski(k1_mcart_1(esk1_0),esk1_0)=esk1_0|k1_mcart_1(esk1_0)=esk1_0), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(X1,X2))|~r2_hidden(k2_mcart_1(esk1_0),X2)|~r2_hidden(k1_mcart_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_39, c_0_37])).
cnf(c_0_47, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X2))), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_48, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_enumset1(X4,X4,X4,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_24]), c_0_25])).
cnf(c_0_49, plain, (X2=k1_tarski(X1)|~r2_hidden(esk6_2(X1,X2),X2)|esk6_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_50, plain, (k1_xboole_0=k2_enumset1(X1,X1,X1,X1)|esk6_2(X1,k1_xboole_0)=X1|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_51, negated_conjecture, (k1_mcart_1(esk1_0)=esk1_0|r2_hidden(esk1_0,X1)|~r2_hidden(esk1_0,k2_zfmisc_1(X2,X1))), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_52, negated_conjecture, (r2_hidden(esk1_0,k2_zfmisc_1(X1,k2_enumset1(k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),X2)))|~r2_hidden(k1_mcart_1(esk1_0),X1)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, plain, (r2_hidden(X1,X2)|X2!=k2_enumset1(X3,X3,X3,X1)), inference(er,[status(thm)],[c_0_48])).
fof(c_0_54, plain, ![X8, X9]:k2_zfmisc_1(k1_tarski(X8),k1_tarski(X9))=k1_tarski(k4_tarski(X8,X9)), inference(variable_rename,[status(thm)],[t35_zfmisc_1])).
cnf(c_0_55, plain, (X2=k2_enumset1(X1,X1,X1,X1)|esk6_2(X1,X2)!=X1|~r2_hidden(esk6_2(X1,X2),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_36]), c_0_24]), c_0_25])).
cnf(c_0_56, plain, (k1_xboole_0=k2_enumset1(X1,X1,X1,X1)|esk6_2(X1,k1_xboole_0)=X1), inference(er,[status(thm)],[c_0_50])).
fof(c_0_57, plain, ![X33, X35, X36]:((r2_hidden(esk5_1(X33),X33)|X33=k1_xboole_0)&((~r2_hidden(X35,X33)|esk5_1(X33)!=k4_tarski(X35,X36)|X33=k1_xboole_0)&(~r2_hidden(X36,X33)|esk5_1(X33)!=k4_tarski(X35,X36)|X33=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t9_mcart_1])])])])])])).
cnf(c_0_58, negated_conjecture, (k1_mcart_1(esk1_0)=esk1_0|r2_hidden(esk1_0,k2_enumset1(k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),X1))|~r2_hidden(k1_mcart_1(esk1_0),X2)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_59, plain, (r2_hidden(X1,k2_enumset1(X2,X2,X2,X1))), inference(er,[status(thm)],[c_0_53])).
cnf(c_0_60, plain, (k2_zfmisc_1(k1_tarski(X1),k1_tarski(X2))=k1_tarski(k4_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_61, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_62, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_42])).
cnf(c_0_63, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk5_1(X2)!=k4_tarski(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64, negated_conjecture, (k1_mcart_1(esk1_0)=esk1_0|r2_hidden(esk1_0,k2_enumset1(k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),X1))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_66, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_67, plain, (k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X2,X2))=k2_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_36]), c_0_36]), c_0_36]), c_0_24]), c_0_24]), c_0_24]), c_0_25]), c_0_25]), c_0_25])).
cnf(c_0_68, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(k1_mcart_1(esk1_0),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_61]), c_0_62])).
cnf(c_0_69, negated_conjecture, (k2_enumset1(k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),X1)=k1_xboole_0|k1_mcart_1(esk1_0)=esk1_0|esk5_1(k2_enumset1(k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),X1))!=k4_tarski(X2,esk1_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_70, plain, (X1=X3|X2!=k2_enumset1(X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_36]), c_0_24]), c_0_25])).
cnf(c_0_71, plain, (k2_enumset1(X1,X1,X1,X1)=k1_xboole_0|k2_enumset1(X2,X2,X2,X2)=k1_xboole_0|k2_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_72, negated_conjecture, (k2_enumset1(k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),k1_mcart_1(esk1_0),X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_47])).
cnf(c_0_73, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,X1)=k1_xboole_0|k1_mcart_1(esk1_0)=esk1_0|esk5_1(k2_enumset1(esk1_0,esk1_0,esk1_0,X1))!=k4_tarski(X2,esk1_0)), inference(spm,[status(thm)],[c_0_69, c_0_38])).
cnf(c_0_74, plain, (X1=X2|~r2_hidden(X1,k2_enumset1(X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_70])).
cnf(c_0_75, plain, (r2_hidden(esk5_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_76, negated_conjecture, (k2_enumset1(k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),k2_mcart_1(esk1_0),k2_mcart_1(esk1_0))=k1_xboole_0|k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_37]), c_0_72])).
cnf(c_0_77, plain, (X2=k1_xboole_0|~r2_hidden(X1,X2)|esk5_1(X2)!=k4_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_78, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,X1)=k1_xboole_0|k1_mcart_1(esk1_0)=esk1_0|esk5_1(k2_enumset1(esk1_0,esk1_0,esk1_0,X1))!=esk1_0), inference(spm,[status(thm)],[c_0_73, c_0_45])).
cnf(c_0_79, plain, (esk5_1(k2_enumset1(X1,X1,X1,X1))=X1|k2_enumset1(X1,X1,X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_80, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_76]), c_0_62])).
cnf(c_0_81, plain, (X1=k1_xboole_0|k4_tarski(esk5_1(X1),X2)!=esk5_1(X1)), inference(spm,[status(thm)],[c_0_77, c_0_75])).
cnf(c_0_82, negated_conjecture, (k1_mcart_1(esk1_0)=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_80])).
cnf(c_0_83, plain, (k2_enumset1(X1,X1,X1,X1)=k1_xboole_0|k4_tarski(X1,X2)!=X1), inference(spm,[status(thm)],[c_0_81, c_0_79])).
cnf(c_0_84, negated_conjecture, (k4_tarski(esk1_0,k2_mcart_1(esk1_0))=esk1_0), inference(rw,[status(thm)],[c_0_37, c_0_82])).
cnf(c_0_85, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_80]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 86
# Proof object clause steps            : 61
# Proof object formula steps           : 25
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 25
# Processed clauses                    : 216
# ...of these trivial                  : 2
# ...subsumed                          : 87
# ...remaining for further processing  : 127
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 6
# Backward-rewritten                   : 27
# Generated clauses                    : 303
# ...of the previous two non-trivial   : 271
# Contextual simplify-reflections      : 1
# Paramodulations                      : 288
# Factorizations                       : 2
# Equation resolutions                 : 13
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 67
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 49
# Current number of unprocessed clauses: 98
# ...number of literals in the above   : 297
# Current number of archived formulas  : 0
# Current number of archived clauses   : 60
# Clause-clause subsumption calls (NU) : 631
# Rec. Clause-clause subsumption calls : 515
# Non-unit clause-clause subsumptions  : 43
# Unit Clause-clause subsumption calls : 160
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5667

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28632 pages
