# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t33_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2)))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t33_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t69_enumset1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t70_enumset1)).
fof(connectedness_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', connectedness_r1_ordinal1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t75_enumset1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t7_ordinal1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', redefinition_r1_ordinal1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', d3_xboole_0)).
fof(t26_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t26_ordinal1)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t29_ordinal1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', l1_zfmisc_1)).
fof(t24_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>~((~(r2_hidden(X1,X2))&(~(X1=X2)&~(r2_hidden(X2,X1))))))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', t24_ordinal1)).
fof(d6_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:(X9=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)<=>![X10]:(r2_hidden(X10,X9)<=>~((~(X10=X1)&(~(X10=X2)&(~(X10=X3)&(~(X10=X4)&(~(X10=X5)&(~(X10=X6)&(~(X10=X7)&~(X10=X8))))))))))), file('data/MPTP/proofs/c57dfc54c0a04fab8262df0a726a0f62.E_input', d6_enumset1)).
fof(c_0_19, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2))))), inference(assume_negation,[status(cth)],[t33_ordinal1])).
fof(c_0_20, plain, ![X60]:k1_ordinal1(X60)=k2_xboole_0(X60,k1_tarski(X60)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_21, plain, ![X86]:k2_tarski(X86,X86)=k1_tarski(X86), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_22, plain, ![X14, X15]:k3_tarski(k2_tarski(X14,X15))=k2_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_23, plain, ![X30, X31]:k1_enumset1(X30,X30,X31)=k2_tarski(X30,X31), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_24, plain, ![X84, X85]:(~v3_ordinal1(X84)|~v3_ordinal1(X85)|(r1_ordinal1(X84,X85)|r1_ordinal1(X85,X84))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_ordinal1])])).
fof(c_0_25, negated_conjecture, (v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&((~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k1_ordinal1(esk1_0),esk2_0))&(r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k1_ordinal1(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_19])])])).
cnf(c_0_26, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, ![X20, X21, X22]:k2_enumset1(X20,X20,X21,X22)=k1_enumset1(X20,X21,X22), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_31, plain, ![X54, X55, X56, X57]:k3_enumset1(X54,X54,X55,X56,X57)=k2_enumset1(X54,X55,X56,X57), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_32, plain, ![X25, X26, X27, X28, X29]:k4_enumset1(X25,X25,X26,X27,X28,X29)=k3_enumset1(X25,X26,X27,X28,X29), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_33, plain, ![X48, X49, X50, X51, X52, X53]:k5_enumset1(X48,X48,X49,X50,X51,X52,X53)=k4_enumset1(X48,X49,X50,X51,X52,X53), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_34, plain, ![X32, X33, X34, X35, X36, X37, X38]:k6_enumset1(X32,X32,X33,X34,X35,X36,X37,X38)=k5_enumset1(X32,X33,X34,X35,X36,X37,X38), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_35, plain, ![X23, X24]:(~r2_hidden(X23,X24)|~r1_tarski(X24,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_36, plain, ![X18, X19]:((~r1_ordinal1(X18,X19)|r1_tarski(X18,X19)|(~v3_ordinal1(X18)|~v3_ordinal1(X19)))&(~r1_tarski(X18,X19)|r1_ordinal1(X18,X19)|(~v3_ordinal1(X18)|~v3_ordinal1(X19)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_37, plain, (r1_ordinal1(X1,X2)|r1_ordinal1(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_39, plain, ![X39, X40, X41, X42, X43, X44, X45, X46]:(((~r2_hidden(X42,X41)|(r2_hidden(X42,X39)|r2_hidden(X42,X40))|X41!=k2_xboole_0(X39,X40))&((~r2_hidden(X43,X39)|r2_hidden(X43,X41)|X41!=k2_xboole_0(X39,X40))&(~r2_hidden(X43,X40)|r2_hidden(X43,X41)|X41!=k2_xboole_0(X39,X40))))&(((~r2_hidden(esk3_3(X44,X45,X46),X44)|~r2_hidden(esk3_3(X44,X45,X46),X46)|X46=k2_xboole_0(X44,X45))&(~r2_hidden(esk3_3(X44,X45,X46),X45)|~r2_hidden(esk3_3(X44,X45,X46),X46)|X46=k2_xboole_0(X44,X45)))&(r2_hidden(esk3_3(X44,X45,X46),X46)|(r2_hidden(esk3_3(X44,X45,X46),X44)|r2_hidden(esk3_3(X44,X45,X46),X45))|X46=k2_xboole_0(X44,X45)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_40, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_41, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_42, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_43, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_45, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_46, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_47, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_48, plain, ![X82, X83]:(~v3_ordinal1(X82)|(~v3_ordinal1(X83)|(r1_ordinal1(X82,X83)|r2_hidden(X83,X82)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])).
fof(c_0_49, plain, ![X13]:(~v3_ordinal1(X13)|v3_ordinal1(k1_ordinal1(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
cnf(c_0_50, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_51, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_52, negated_conjecture, (r1_ordinal1(X1,esk2_0)|r1_ordinal1(esk2_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_53, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_54, plain, ![X16, X17]:((~r1_tarski(k1_tarski(X16),X17)|r2_hidden(X16,X17))&(~r2_hidden(X16,X17)|r1_tarski(k1_tarski(X16),X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
cnf(c_0_55, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_56, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_29]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47])).
cnf(c_0_57, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, plain, (~r1_ordinal1(X1,X2)|~r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_60, negated_conjecture, (r1_ordinal1(esk2_0,esk1_0)|r1_ordinal1(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_61, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
fof(c_0_62, plain, ![X58, X59]:(~v3_ordinal1(X58)|(~v3_ordinal1(X59)|(r2_hidden(X58,X59)|(X58=X59|r2_hidden(X59,X58))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_ordinal1])])])])).
cnf(c_0_63, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_64, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_42]), c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk2_0,k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))|~r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_38])])).
cnf(c_0_66, plain, (v3_ordinal1(k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1))))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_41]), c_0_29]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47])).
cnf(c_0_67, negated_conjecture, (r1_ordinal1(esk1_0,esk2_0)|~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_53]), c_0_38])])).
cnf(c_0_68, plain, (r1_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_27]), c_0_29]), c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_69, plain, (r2_hidden(X1,X2)|X1=X2|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_41]), c_0_29]), c_0_42]), c_0_43]), c_0_43]), c_0_44]), c_0_44]), c_0_45]), c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47])).
cnf(c_0_71, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k3_tarski(k6_enumset1(X3,X3,X3,X3,X3,X3,X3,X2)))), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_72, negated_conjecture, (r2_hidden(esk2_0,k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))|~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_53])])).
cnf(c_0_73, negated_conjecture, (~r2_hidden(esk2_0,esk1_0)|~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_67]), c_0_38]), c_0_53])])).
cnf(c_0_74, plain, (~r2_hidden(X1,k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X2))|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_50, c_0_68])).
cnf(c_0_75, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_76, negated_conjecture, (X1=esk2_0|r2_hidden(X1,esk2_0)|r2_hidden(esk2_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_69, c_0_38])).
cnf(c_0_77, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|~r2_hidden(esk2_0,k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))|~v3_ordinal1(k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_70]), c_0_38])])).
cnf(c_0_78, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_74])).
cnf(c_0_79, plain, (r2_hidden(X1,X3)|X3!=k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_42]), c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_80, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,esk1_0)|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_76, c_0_53])).
cnf(c_0_81, negated_conjecture, (~r2_hidden(esk2_0,k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_66]), c_0_53])]), c_0_78])).
cnf(c_0_82, plain, (r2_hidden(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_83, negated_conjecture, (esk2_0=esk1_0|r2_hidden(esk2_0,esk1_0)), inference(sr,[status(thm)],[c_0_80, c_0_78])).
cnf(c_0_84, negated_conjecture, (~r2_hidden(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_85, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_86, plain, ![X61, X62, X63, X64, X65, X66, X67, X68, X69, X70, X71, X72, X73, X74, X75, X76, X77, X78, X79, X80]:(((~r2_hidden(X70,X69)|(X70=X61|(X70=X62|(X70=X63|(X70=X64|(X70=X65|(X70=X66|(X70=X67|X70=X68)))))))|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))&((X71!=X61|r2_hidden(X71,X69)|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))&((X71!=X62|r2_hidden(X71,X69)|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))&((X71!=X63|r2_hidden(X71,X69)|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))&((X71!=X64|r2_hidden(X71,X69)|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))&((X71!=X65|r2_hidden(X71,X69)|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))&((X71!=X66|r2_hidden(X71,X69)|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))&((X71!=X67|r2_hidden(X71,X69)|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))&(X71!=X68|r2_hidden(X71,X69)|X69!=k6_enumset1(X61,X62,X63,X64,X65,X66,X67,X68))))))))))&(((esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)!=X72|~r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79))&((esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)!=X73|~r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79))&((esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)!=X74|~r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79))&((esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)!=X75|~r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79))&((esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)!=X76|~r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79))&((esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)!=X77|~r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79))&((esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)!=X78|~r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79))&(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)!=X79|~r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79)))))))))&(r2_hidden(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80),X80)|(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)=X72|(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)=X73|(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)=X74|(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)=X75|(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)=X76|(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)=X77|(esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)=X78|esk4_9(X72,X73,X74,X75,X76,X77,X78,X79,X80)=X79)))))))|X80=k6_enumset1(X72,X73,X74,X75,X76,X77,X78,X79)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_enumset1])])])])])])])).
cnf(c_0_87, negated_conjecture, (esk2_0=esk1_0), inference(sr,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_88, plain, (r2_hidden(X1,X3)|X3!=k3_tarski(k6_enumset1(X4,X4,X4,X4,X4,X4,X4,X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_42]), c_0_43]), c_0_44]), c_0_45]), c_0_46]), c_0_47])).
cnf(c_0_89, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k6_enumset1(X2,X4,X5,X6,X7,X8,X9,X10)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_90, negated_conjecture, (~r2_hidden(esk1_0,k3_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))))), inference(rw,[status(thm)],[c_0_81, c_0_87])).
cnf(c_0_91, plain, (r2_hidden(X1,k3_tarski(k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_88])).
cnf(c_0_92, plain, (r2_hidden(X1,k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_89])])).
cnf(c_0_93, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 94
# Proof object clause steps            : 55
# Proof object formula steps           : 39
# Proof object conjectures             : 25
# Proof object clause conjectures      : 22
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 93
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 46
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 37
# Processed clauses                    : 123
# ...of these trivial                  : 0
# ...subsumed                          : 15
# ...remaining for further processing  : 108
# Other redundant clauses eliminated   : 20
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 14
# Generated clauses                    : 86
# ...of the previous two non-trivial   : 87
# Contextual simplify-reflections      : 2
# Paramodulations                      : 71
# Factorizations                       : 0
# Equation resolutions                 : 20
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 36
# ...number of literals in the above   : 108
# Current number of archived formulas  : 0
# Current number of archived clauses   : 64
# Clause-clause subsumption calls (NU) : 142
# Rec. Clause-clause subsumption calls : 101
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 79
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 57
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4077

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28736 pages
