# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t64_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0)), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t64_relat_1)).
fof(s3_funct_1__e2_25__funct_1, axiom, ![X1]:?[X2]:((v1_relat_1(X2)&v1_funct_1(X2))&(k1_relat_1(X2)=X1&![X3]:(r2_hidden(X3,X1)=>k1_funct_1(X2,X3)=k1_xboole_0))), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', s3_funct_1__e2_25__funct_1)).
fof(t18_funct_1, conjecture, ![X1, X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k1_relat_1(X3)&r1_tarski(k2_relat_1(X3),X1)))))), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t18_funct_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', d1_tarski)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t70_enumset1)).
fof(t15_funct_1, axiom, ![X1]:(~(X1=k1_xboole_0)=>![X2]:?[X3]:((v1_relat_1(X3)&v1_funct_1(X3))&(k1_relat_1(X3)=X1&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t15_funct_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t2_xboole_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', d3_tarski)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t4_xboole_0)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t2_boole)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t60_relat_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', t65_xboole_1)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/c5b15d28e836472698ec43b04b350b31.E_input', d4_xboole_0)).
fof(c_0_14, plain, ![X5]:((k1_relat_1(X5)!=k1_xboole_0|X5=k1_xboole_0|~v1_relat_1(X5))&(k2_relat_1(X5)!=k1_xboole_0|X5=k1_xboole_0|~v1_relat_1(X5))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
fof(c_0_15, plain, ![X6, X8]:((v1_relat_1(esk1_1(X6))&v1_funct_1(esk1_1(X6)))&(k1_relat_1(esk1_1(X6))=X6&(~r2_hidden(X8,X6)|k1_funct_1(esk1_1(X6),X8)=k1_xboole_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s3_funct_1__e2_25__funct_1])])])])).
fof(c_0_16, negated_conjecture, ~(![X1, X2]:~((~((X1=k1_xboole_0&~(X2=k1_xboole_0)))&![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>~((X2=k1_relat_1(X3)&r1_tarski(k2_relat_1(X3),X1))))))), inference(assume_negation,[status(cth)],[t18_funct_1])).
cnf(c_0_17, plain, (X1=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (v1_relat_1(esk1_1(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (k1_relat_1(esk1_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, ![X12, X13, X14, X15, X16, X17]:(((~r2_hidden(X14,X13)|X14=X12|X13!=k1_tarski(X12))&(X15!=X12|r2_hidden(X15,X13)|X13!=k1_tarski(X12)))&((~r2_hidden(esk4_2(X16,X17),X17)|esk4_2(X16,X17)!=X16|X17=k1_tarski(X16))&(r2_hidden(esk4_2(X16,X17),X17)|esk4_2(X16,X17)=X16|X17=k1_tarski(X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_21, plain, ![X19, X20]:k1_setfam_1(k2_tarski(X19,X20))=k3_xboole_0(X19,X20), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_22, plain, ![X21, X22]:k1_enumset1(X21,X21,X22)=k2_tarski(X21,X22), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_23, negated_conjecture, ![X11]:((esk2_0!=k1_xboole_0|esk3_0=k1_xboole_0)&(~v1_relat_1(X11)|~v1_funct_1(X11)|(esk3_0!=k1_relat_1(X11)|~r1_tarski(k2_relat_1(X11),esk2_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])])).
fof(c_0_24, plain, ![X23, X24]:(((v1_relat_1(esk5_2(X23,X24))|X23=k1_xboole_0)&(v1_funct_1(esk5_2(X23,X24))|X23=k1_xboole_0))&((k1_relat_1(esk5_2(X23,X24))=X23|X23=k1_xboole_0)&(k2_relat_1(esk5_2(X23,X24))=k1_tarski(X24)|X23=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_funct_1])])])])])])).
fof(c_0_25, plain, ![X26]:r1_tarski(k1_xboole_0,X26), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_26, plain, (v1_funct_1(esk1_1(X1))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, plain, (esk1_1(k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])])).
cnf(c_0_28, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_29, plain, ![X27, X28, X29, X30, X31]:((~r1_tarski(X27,X28)|(~r2_hidden(X29,X27)|r2_hidden(X29,X28)))&((r2_hidden(esk6_2(X30,X31),X30)|r1_tarski(X30,X31))&(~r2_hidden(esk6_2(X30,X31),X31)|r1_tarski(X30,X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_30, plain, ![X33, X34, X36, X37, X38]:((r1_xboole_0(X33,X34)|r2_hidden(esk7_2(X33,X34),k3_xboole_0(X33,X34)))&(~r2_hidden(X38,k3_xboole_0(X36,X37))|~r1_xboole_0(X36,X37))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
cnf(c_0_31, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_33, plain, ![X39]:k3_xboole_0(X39,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_34, negated_conjecture, (~v1_relat_1(X1)|~v1_funct_1(X1)|esk3_0!=k1_relat_1(X1)|~r1_tarski(k2_relat_1(X1),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_35, plain, (k2_relat_1(esk5_2(X1,X2))=k1_tarski(X2)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (v1_relat_1(esk5_2(X1,X2))|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (v1_funct_1(esk5_2(X1,X2))|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_38, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_39, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
cnf(c_0_40, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_41, plain, (v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_42, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_18, c_0_27])).
cnf(c_0_43, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_28])).
cnf(c_0_44, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_45, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_46, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_47, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_48, plain, ![X40]:r1_xboole_0(X40,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_49, plain, ![X41, X42, X43, X44, X45, X46, X47, X48]:((((r2_hidden(X44,X41)|~r2_hidden(X44,X43)|X43!=k3_xboole_0(X41,X42))&(r2_hidden(X44,X42)|~r2_hidden(X44,X43)|X43!=k3_xboole_0(X41,X42)))&(~r2_hidden(X45,X41)|~r2_hidden(X45,X42)|r2_hidden(X45,X43)|X43!=k3_xboole_0(X41,X42)))&((~r2_hidden(esk8_3(X46,X47,X48),X48)|(~r2_hidden(esk8_3(X46,X47,X48),X46)|~r2_hidden(esk8_3(X46,X47,X48),X47))|X48=k3_xboole_0(X46,X47))&((r2_hidden(esk8_3(X46,X47,X48),X46)|r2_hidden(esk8_3(X46,X47,X48),X48)|X48=k3_xboole_0(X46,X47))&(r2_hidden(esk8_3(X46,X47,X48),X47)|r2_hidden(esk8_3(X46,X47,X48),X48)|X48=k3_xboole_0(X46,X47))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_50, negated_conjecture, (X1=k1_xboole_0|k1_relat_1(esk5_2(X1,X2))!=esk3_0|~r1_tarski(k1_tarski(X2),esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36]), c_0_37])).
cnf(c_0_51, plain, (k1_relat_1(esk5_2(X1,X2))=X1|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_52, negated_conjecture, (esk3_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_38]), c_0_39]), c_0_40]), c_0_41]), c_0_42])])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_54, plain, (esk6_2(k1_tarski(X1),X2)=X1|r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_55, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k1_setfam_1(k1_enumset1(X2,X2,X3)))), inference(rw,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56, plain, (k1_setfam_1(k1_enumset1(X1,X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_47, c_0_46])).
cnf(c_0_57, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (r2_hidden(esk8_3(X1,X2,X3),X2)|r2_hidden(esk8_3(X1,X2,X3),X3)|X3=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_59, negated_conjecture, (~r1_tarski(k1_tarski(X1),esk2_0)), inference(sr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51])]), c_0_52])).
cnf(c_0_60, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])])).
cnf(c_0_62, plain, (X3=k1_setfam_1(k1_enumset1(X1,X1,X2))|r2_hidden(esk8_3(X1,X2,X3),X3)|r2_hidden(esk8_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_58, c_0_46])).
cnf(c_0_63, negated_conjecture, (~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_64, plain, (X1=k1_xboole_0|r2_hidden(esk8_3(X2,k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_56])).
cnf(c_0_65, negated_conjecture, (esk3_0=k1_xboole_0|esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_66, negated_conjecture, (esk2_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_67, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_66])]), c_0_52]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 68
# Proof object clause steps            : 40
# Proof object formula steps           : 28
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 22
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 32
# Processed clauses                    : 98
# ...of these trivial                  : 0
# ...subsumed                          : 9
# ...remaining for further processing  : 89
# Other redundant clauses eliminated   : 9
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 7
# Generated clauses                    : 143
# ...of the previous two non-trivial   : 129
# Contextual simplify-reflections      : 3
# Paramodulations                      : 133
# Factorizations                       : 2
# Equation resolutions                 : 9
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 44
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 28
# Current number of unprocessed clauses: 89
# ...number of literals in the above   : 327
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 189
# Rec. Clause-clause subsumption calls : 152
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 72
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3598

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.011 s
# Maximum resident set size: 26236 pages
