# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t5_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X3)=>r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3))&((r2_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)=>r2_lattice3(X1,X2,X3))&((r1_lattice3(X1,X2,X3)=>r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3))&(r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)=>r1_lattice3(X1,X2,X3))))))), file('data/MPTP/proofs/c7be82471f144aafb291995726aa7461.E_input', t5_yellow_0)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/c7be82471f144aafb291995726aa7461.E_input', dt_k2_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/c7be82471f144aafb291995726aa7461.E_input', d10_yellow_0)).
fof(t52_yellow_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r2_yellow_0(X1,X2)|r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))=>k2_yellow_0(X1,X2)=k2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))), file('data/MPTP/proofs/c7be82471f144aafb291995726aa7461.E_input', t52_yellow_0)).
fof(t50_yellow_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r1_yellow_0(X1,X2)=>r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&((r1_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r1_yellow_0(X1,X2))&((r2_yellow_0(X1,X2)=>r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))&(r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))=>r2_yellow_0(X1,X2)))))), file('data/MPTP/proofs/c7be82471f144aafb291995726aa7461.E_input', t50_yellow_0)).
fof(c_0_5, plain, ![X5, X6, X7]:((~r2_lattice3(X5,X6,X7)|r2_lattice3(X5,k3_xboole_0(X6,u1_struct_0(X5)),X7)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))&((~r2_lattice3(X5,k3_xboole_0(X6,u1_struct_0(X5)),X7)|r2_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))&((~r1_lattice3(X5,X6,X7)|r1_lattice3(X5,k3_xboole_0(X6,u1_struct_0(X5)),X7)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))&(~r1_lattice3(X5,k3_xboole_0(X6,u1_struct_0(X5)),X7)|r1_lattice3(X5,X6,X7)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|~l1_orders_2(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_yellow_0])])])])])).
fof(c_0_6, plain, ![X8, X9]:(~l1_orders_2(X8)|m1_subset_1(k2_yellow_0(X8,X9),u1_struct_0(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
fof(c_0_7, plain, ![X10, X11, X12, X13]:(((r1_lattice3(X10,X11,X12)|X12!=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&(~m1_subset_1(X13,u1_struct_0(X10))|(~r1_lattice3(X10,X11,X13)|r1_orders_2(X10,X13,X12))|X12!=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10)))&((m1_subset_1(esk1_3(X10,X11,X12),u1_struct_0(X10))|~r1_lattice3(X10,X11,X12)|X12=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&((r1_lattice3(X10,X11,esk1_3(X10,X11,X12))|~r1_lattice3(X10,X11,X12)|X12=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))&(~r1_orders_2(X10,esk1_3(X10,X11,X12),X12)|~r1_lattice3(X10,X11,X12)|X12=k2_yellow_0(X10,X11)|~r2_yellow_0(X10,X11)|~m1_subset_1(X12,u1_struct_0(X10))|~l1_orders_2(X10))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
fof(c_0_8, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:((r2_yellow_0(X1,X2)|r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1))))=>k2_yellow_0(X1,X2)=k2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))))), inference(assume_negation,[status(cth)],[t52_yellow_0])).
cnf(c_0_9, plain, (r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)|v2_struct_0(X1)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (r1_lattice3(X1,X2,X3)|X3!=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_12, plain, ![X17, X18]:((~r1_yellow_0(X17,X18)|r1_yellow_0(X17,k3_xboole_0(X18,u1_struct_0(X17)))|(v2_struct_0(X17)|~l1_orders_2(X17)))&((~r1_yellow_0(X17,k3_xboole_0(X18,u1_struct_0(X17)))|r1_yellow_0(X17,X18)|(v2_struct_0(X17)|~l1_orders_2(X17)))&((~r2_yellow_0(X17,X18)|r2_yellow_0(X17,k3_xboole_0(X18,u1_struct_0(X17)))|(v2_struct_0(X17)|~l1_orders_2(X17)))&(~r2_yellow_0(X17,k3_xboole_0(X18,u1_struct_0(X17)))|r2_yellow_0(X17,X18)|(v2_struct_0(X17)|~l1_orders_2(X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_yellow_0])])])])])).
fof(c_0_13, negated_conjecture, ((~v2_struct_0(esk2_0)&l1_orders_2(esk2_0))&((r2_yellow_0(esk2_0,esk3_0)|r2_yellow_0(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0))))&k2_yellow_0(esk2_0,esk3_0)!=k2_yellow_0(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
cnf(c_0_14, plain, (r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),k2_yellow_0(X1,X3))|v2_struct_0(X1)|~r1_lattice3(X1,X2,k2_yellow_0(X1,X3))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_15, plain, (r1_lattice3(X1,X2,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_11]), c_0_10])).
cnf(c_0_16, plain, (r2_yellow_0(X1,X2)|v2_struct_0(X1)|~r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, negated_conjecture, (r2_yellow_0(esk2_0,esk3_0)|r2_yellow_0(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21, plain, (r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),k2_yellow_0(X1,X2))|v2_struct_0(X1)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22, negated_conjecture, (r2_yellow_0(esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_23, plain, (r2_yellow_0(X1,k3_xboole_0(X2,u1_struct_0(X1)))|v2_struct_0(X1)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_24, plain, (k2_yellow_0(X1,X2)=k2_yellow_0(X1,X3)|m1_subset_1(esk1_3(X1,X3,k2_yellow_0(X1,X2)),u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,k2_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_10])).
cnf(c_0_25, negated_conjecture, (r1_lattice3(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)),k2_yellow_0(esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_18])]), c_0_19])).
cnf(c_0_26, negated_conjecture, (r2_yellow_0(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_22]), c_0_18])]), c_0_19])).
cnf(c_0_27, negated_conjecture, (k2_yellow_0(esk2_0,esk3_0)!=k2_yellow_0(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_28, plain, (r1_lattice3(X1,X2,esk1_3(X1,X2,X3))|X3=k2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_29, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30, plain, (r1_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattice3(X1,k3_xboole_0(X2,u1_struct_0(X1)),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk1_3(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)),k2_yellow_0(esk2_0,esk3_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26]), c_0_18])]), c_0_27])).
cnf(c_0_32, plain, (k2_yellow_0(X1,X2)=k2_yellow_0(X1,X3)|r1_lattice3(X1,X3,esk1_3(X1,X3,k2_yellow_0(X1,X2)))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,k2_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_28, c_0_10])).
cnf(c_0_33, plain, (X3=k2_yellow_0(X1,X2)|~r1_orders_2(X1,esk1_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_34, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r2_yellow_0(X1,X3)|~r1_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_29]), c_0_10])).
cnf(c_0_35, negated_conjecture, (r1_lattice3(esk2_0,X1,esk1_3(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)),k2_yellow_0(esk2_0,esk3_0)))|~r1_lattice3(esk2_0,k3_xboole_0(X1,u1_struct_0(esk2_0)),esk1_3(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)),k2_yellow_0(esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_18])]), c_0_19])).
cnf(c_0_36, negated_conjecture, (r1_lattice3(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)),esk1_3(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)),k2_yellow_0(esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_25]), c_0_26]), c_0_18])]), c_0_27])).
cnf(c_0_37, plain, (k2_yellow_0(X1,X2)=k2_yellow_0(X1,X3)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,esk1_3(X1,X3,k2_yellow_0(X1,X2)))|~r1_lattice3(X1,X3,k2_yellow_0(X1,X2))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_10]), c_0_24])).
cnf(c_0_38, negated_conjecture, (r1_lattice3(esk2_0,esk3_0,esk1_3(esk2_0,k3_xboole_0(esk3_0,u1_struct_0(esk2_0)),k2_yellow_0(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_26]), c_0_22]), c_0_25]), c_0_18])]), c_0_27]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 29
# Proof object formula steps           : 11
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 369
# ...of these trivial                  : 3
# ...subsumed                          : 38
# ...remaining for further processing  : 328
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 1
# Generated clauses                    : 1056
# ...of the previous two non-trivial   : 933
# Contextual simplify-reflections      : 4
# Paramodulations                      : 1054
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 304
#    Positive orientable unit clauses  : 141
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 161
# Current number of unprocessed clauses: 600
# ...number of literals in the above   : 1728
# Current number of archived formulas  : 0
# Current number of archived clauses   : 24
# Clause-clause subsumption calls (NU) : 5915
# Rec. Clause-clause subsumption calls : 3486
# Non-unit clause-clause subsumptions  : 47
# Unit Clause-clause subsumption calls : 60
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5124
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 122923

# -------------------------------------------------
# User time                : 0.037 s
# System time              : 0.000 s
# Total time               : 0.037 s
# Maximum resident set size: 28548 pages
