# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', t4_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', t3_subset)).
fof(t79_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', t79_zfmisc_1)).
fof(d9_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X2,X1)<=>(r1_tarski(k2_struct_0(X2),k2_struct_0(X1))&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(r2_hidden(X3,u1_pre_topc(X2))<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(r2_hidden(X4,u1_pre_topc(X1))&X3=k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2)))))))))), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', d9_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', dt_m1_pre_topc)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', t2_subset)).
fof(t39_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', t39_pre_topc)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', dt_l1_pre_topc)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', d2_subset_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', l13_xboole_0)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/c80826cbe7b748c3b0316f20d9e05f15.E_input', t2_xboole_1)).
fof(c_0_12, plain, ![X5, X6, X7]:(~r2_hidden(X5,X6)|~m1_subset_1(X6,k1_zfmisc_1(X7))|m1_subset_1(X5,X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_13, plain, ![X8, X9]:((~m1_subset_1(X8,k1_zfmisc_1(X9))|r1_tarski(X8,X9))&(~r1_tarski(X8,X9)|m1_subset_1(X8,k1_zfmisc_1(X9)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_14, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_16, plain, ![X10, X11]:(~r1_tarski(X10,X11)|r1_tarski(k1_zfmisc_1(X10),k1_zfmisc_1(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t79_zfmisc_1])])).
fof(c_0_17, plain, ![X12, X13, X14, X16, X18]:(((r1_tarski(k2_struct_0(X13),k2_struct_0(X12))|~m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))&(((m1_subset_1(esk1_3(X12,X13,X14),k1_zfmisc_1(u1_struct_0(X12)))|~r2_hidden(X14,u1_pre_topc(X13))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))&((r2_hidden(esk1_3(X12,X13,X14),u1_pre_topc(X12))|~r2_hidden(X14,u1_pre_topc(X13))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))&(X14=k9_subset_1(u1_struct_0(X13),esk1_3(X12,X13,X14),k2_struct_0(X13))|~r2_hidden(X14,u1_pre_topc(X13))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))))&(~m1_subset_1(X16,k1_zfmisc_1(u1_struct_0(X12)))|(~r2_hidden(X16,u1_pre_topc(X12))|X14!=k9_subset_1(u1_struct_0(X13),X16,k2_struct_0(X13)))|r2_hidden(X14,u1_pre_topc(X13))|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|~m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))))&((m1_subset_1(esk2_2(X12,X13),k1_zfmisc_1(u1_struct_0(X13)))|~r1_tarski(k2_struct_0(X13),k2_struct_0(X12))|m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))&((~r2_hidden(esk2_2(X12,X13),u1_pre_topc(X13))|(~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X12)))|(~r2_hidden(X18,u1_pre_topc(X12))|esk2_2(X12,X13)!=k9_subset_1(u1_struct_0(X13),X18,k2_struct_0(X13))))|~r1_tarski(k2_struct_0(X13),k2_struct_0(X12))|m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))&((m1_subset_1(esk3_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|r2_hidden(esk2_2(X12,X13),u1_pre_topc(X13))|~r1_tarski(k2_struct_0(X13),k2_struct_0(X12))|m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))&((r2_hidden(esk3_2(X12,X13),u1_pre_topc(X12))|r2_hidden(esk2_2(X12,X13),u1_pre_topc(X13))|~r1_tarski(k2_struct_0(X13),k2_struct_0(X12))|m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))&(esk2_2(X12,X13)=k9_subset_1(u1_struct_0(X13),esk3_2(X12,X13),k2_struct_0(X13))|r2_hidden(esk2_2(X12,X13),u1_pre_topc(X13))|~r1_tarski(k2_struct_0(X13),k2_struct_0(X12))|m1_pre_topc(X13,X12)|~l1_pre_topc(X13)|~l1_pre_topc(X12))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_pre_topc])])])])])).
fof(c_0_18, plain, ![X20, X21]:(~l1_pre_topc(X20)|(~m1_pre_topc(X21,X20)|l1_pre_topc(X21))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_19, plain, (m1_subset_1(X1,X2)|~r1_tarski(X3,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20, plain, (r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_21, plain, ![X22, X23]:(~m1_subset_1(X22,X23)|(v1_xboole_0(X23)|r2_hidden(X22,X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_22, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t39_pre_topc])).
cnf(c_0_23, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, ![X27]:(~l1_struct_0(X27)|k2_struct_0(X27)=u1_struct_0(X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_26, plain, ![X28]:(~l1_pre_topc(X28)|l1_struct_0(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_27, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X3,X2)|~r2_hidden(X1,k1_zfmisc_1(X3))), inference(spm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_28, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_29, negated_conjecture, (l1_pre_topc(esk4_0)&(m1_pre_topc(esk5_0,esk4_0)&(m1_subset_1(esk6_0,k1_zfmisc_1(u1_struct_0(esk5_0)))&~m1_subset_1(esk6_0,k1_zfmisc_1(u1_struct_0(esk4_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_22])])])).
cnf(c_0_30, plain, (r1_tarski(k2_struct_0(X1),k2_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_31, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (v1_xboole_0(k1_zfmisc_1(X1))|m1_subset_1(X2,k1_zfmisc_1(X3))|~r1_tarski(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_34, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(u1_struct_0(esk5_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, plain, (r1_tarski(k2_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
fof(c_0_36, plain, ![X29, X30]:(((~m1_subset_1(X30,X29)|r2_hidden(X30,X29)|v1_xboole_0(X29))&(~r2_hidden(X30,X29)|m1_subset_1(X30,X29)|v1_xboole_0(X29)))&((~m1_subset_1(X30,X29)|v1_xboole_0(X30)|~v1_xboole_0(X29))&(~v1_xboole_0(X30)|m1_subset_1(X30,X29)|~v1_xboole_0(X29)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_37, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk5_0)))|m1_subset_1(esk6_0,k1_zfmisc_1(X1))|~r1_tarski(u1_struct_0(esk5_0),X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|~l1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_35, c_0_31])).
cnf(c_0_39, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_40, negated_conjecture, (~m1_subset_1(esk6_0,k1_zfmisc_1(u1_struct_0(esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_41, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk5_0)))|m1_subset_1(esk6_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(esk5_0)|~m1_pre_topc(esk5_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_42, negated_conjecture, (m1_pre_topc(esk5_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, negated_conjecture, (l1_pre_topc(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_44, negated_conjecture, (v1_xboole_0(esk6_0)|~v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk5_0)))), inference(spm,[status(thm)],[c_0_39, c_0_34])).
cnf(c_0_45, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk5_0)))|~l1_struct_0(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_42]), c_0_43])])).
fof(c_0_46, plain, ![X31]:(~v1_xboole_0(X31)|X31=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_47, negated_conjecture, (v1_xboole_0(esk6_0)|~l1_struct_0(esk5_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_48, negated_conjecture, (l1_pre_topc(esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_42]), c_0_43])])).
cnf(c_0_49, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_50, negated_conjecture, (v1_xboole_0(esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_32]), c_0_48])])).
fof(c_0_51, plain, ![X32]:r1_tarski(k1_xboole_0,X32), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_52, negated_conjecture, (~r1_tarski(esk6_0,u1_struct_0(esk4_0))), inference(spm,[status(thm)],[c_0_40, c_0_15])).
cnf(c_0_53, negated_conjecture, (esk6_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_55, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53]), c_0_54])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 56
# Proof object clause steps            : 31
# Proof object formula steps           : 25
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 56
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 54
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 11
# Generated clauses                    : 71
# ...of the previous two non-trivial   : 68
# Contextual simplify-reflections      : 6
# Paramodulations                      : 70
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 41
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 35
# Current number of unprocessed clauses: 36
# ...number of literals in the above   : 208
# Current number of archived formulas  : 0
# Current number of archived clauses   : 12
# Clause-clause subsumption calls (NU) : 551
# Rec. Clause-clause subsumption calls : 175
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 62
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3346

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.004 s
# Total time               : 0.010 s
# Maximum resident set size: 28576 pages
