# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t15_yellow_0, conjecture, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/c932d479228a44e6a8f9c51c695236ee.E_input', t15_yellow_0)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/c932d479228a44e6a8f9c51c695236ee.E_input', d7_yellow_0)).
fof(t30_orders_2, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r1_orders_2(X1,X2,X3)&r2_orders_2(X1,X3,X2)))))), file('data/MPTP/proofs/c932d479228a44e6a8f9c51c695236ee.E_input', t30_orders_2)).
fof(d10_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>(r1_orders_2(X1,X2,X3)&~(X2=X3)))))), file('data/MPTP/proofs/c932d479228a44e6a8f9c51c695236ee.E_input', d10_orders_2)).
fof(c_0_4, negated_conjecture, ~(![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))))))), inference(assume_negation,[status(cth)],[t15_yellow_0])).
fof(c_0_5, negated_conjecture, ![X8, X11]:((v5_orders_2(esk1_0)&l1_orders_2(esk1_0))&(((m1_subset_1(esk3_1(X8),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X8)|~m1_subset_1(X8,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0))&((r2_lattice3(esk1_0,esk2_0,esk3_1(X8))|~r2_lattice3(esk1_0,esk2_0,X8)|~m1_subset_1(X8,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0))&(~r1_orders_2(esk1_0,X8,esk3_1(X8))|~r2_lattice3(esk1_0,esk2_0,X8)|~m1_subset_1(X8,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0))))&((m1_subset_1(esk4_0,u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,esk2_0))&((r2_lattice3(esk1_0,esk2_0,esk4_0)|r1_yellow_0(esk1_0,esk2_0))&(~m1_subset_1(X11,u1_struct_0(esk1_0))|(~r2_lattice3(esk1_0,esk2_0,X11)|r1_orders_2(esk1_0,esk4_0,X11))|r1_yellow_0(esk1_0,esk2_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])])).
fof(c_0_6, plain, ![X15, X16, X18, X19, X21, X22, X25]:(((m1_subset_1(esk5_2(X15,X16),u1_struct_0(X15))|~r1_yellow_0(X15,X16)|~l1_orders_2(X15))&((r2_lattice3(X15,X16,esk5_2(X15,X16))|~r1_yellow_0(X15,X16)|~l1_orders_2(X15))&((~m1_subset_1(X18,u1_struct_0(X15))|(~r2_lattice3(X15,X16,X18)|r1_orders_2(X15,esk5_2(X15,X16),X18))|~r1_yellow_0(X15,X16)|~l1_orders_2(X15))&((m1_subset_1(esk6_3(X15,X16,X19),u1_struct_0(X15))|~r2_lattice3(X15,X16,X19)|X19=esk5_2(X15,X16)|~m1_subset_1(X19,u1_struct_0(X15))|~r1_yellow_0(X15,X16)|~l1_orders_2(X15))&((r2_lattice3(X15,X16,esk6_3(X15,X16,X19))|~r2_lattice3(X15,X16,X19)|X19=esk5_2(X15,X16)|~m1_subset_1(X19,u1_struct_0(X15))|~r1_yellow_0(X15,X16)|~l1_orders_2(X15))&(~r1_orders_2(X15,X19,esk6_3(X15,X16,X19))|~r2_lattice3(X15,X16,X19)|X19=esk5_2(X15,X16)|~m1_subset_1(X19,u1_struct_0(X15))|~r1_yellow_0(X15,X16)|~l1_orders_2(X15)))))))&(((m1_subset_1(esk8_3(X15,X21,X22),u1_struct_0(X15))|m1_subset_1(esk7_3(X15,X21,X22),u1_struct_0(X15))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))&(((r2_lattice3(X15,X21,esk8_3(X15,X21,X22))|m1_subset_1(esk7_3(X15,X21,X22),u1_struct_0(X15))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))&(~m1_subset_1(X25,u1_struct_0(X15))|(~r2_lattice3(X15,X21,X25)|r1_orders_2(X15,esk8_3(X15,X21,X22),X25))|m1_subset_1(esk7_3(X15,X21,X22),u1_struct_0(X15))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15)))&(esk8_3(X15,X21,X22)!=X22|m1_subset_1(esk7_3(X15,X21,X22),u1_struct_0(X15))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))))&(((m1_subset_1(esk8_3(X15,X21,X22),u1_struct_0(X15))|r2_lattice3(X15,X21,esk7_3(X15,X21,X22))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))&(((r2_lattice3(X15,X21,esk8_3(X15,X21,X22))|r2_lattice3(X15,X21,esk7_3(X15,X21,X22))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))&(~m1_subset_1(X25,u1_struct_0(X15))|(~r2_lattice3(X15,X21,X25)|r1_orders_2(X15,esk8_3(X15,X21,X22),X25))|r2_lattice3(X15,X21,esk7_3(X15,X21,X22))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15)))&(esk8_3(X15,X21,X22)!=X22|r2_lattice3(X15,X21,esk7_3(X15,X21,X22))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))))&((m1_subset_1(esk8_3(X15,X21,X22),u1_struct_0(X15))|~r1_orders_2(X15,X22,esk7_3(X15,X21,X22))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))&(((r2_lattice3(X15,X21,esk8_3(X15,X21,X22))|~r1_orders_2(X15,X22,esk7_3(X15,X21,X22))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))&(~m1_subset_1(X25,u1_struct_0(X15))|(~r2_lattice3(X15,X21,X25)|r1_orders_2(X15,esk8_3(X15,X21,X22),X25))|~r1_orders_2(X15,X22,esk7_3(X15,X21,X22))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15)))&(esk8_3(X15,X21,X22)!=X22|~r1_orders_2(X15,X22,esk7_3(X15,X21,X22))|~r2_lattice3(X15,X21,X22)|~m1_subset_1(X22,u1_struct_0(X15))|r1_yellow_0(X15,X21)|~l1_orders_2(X15))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
cnf(c_0_7, negated_conjecture, (~r1_orders_2(esk1_0,X1,esk3_1(X1))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (r1_orders_2(X2,esk5_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, negated_conjecture, (m1_subset_1(esk3_1(X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, negated_conjecture, (~r2_lattice3(esk1_0,X1,esk3_1(esk5_2(esk1_0,X1)))|~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,X1))|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_7, c_0_8]), c_0_9])]), c_0_10])).
cnf(c_0_12, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_1(X1))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13, negated_conjecture, (~r2_lattice3(esk1_0,esk2_0,esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_14, plain, (r2_lattice3(X1,X2,esk5_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, negated_conjecture, (~m1_subset_1(esk5_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~r1_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_9])])).
cnf(c_0_16, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_18, negated_conjecture, (~r1_yellow_0(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_9])])).
fof(c_0_19, plain, ![X12, X13, X14]:(~v5_orders_2(X12)|~l1_orders_2(X12)|(~m1_subset_1(X13,u1_struct_0(X12))|(~m1_subset_1(X14,u1_struct_0(X12))|(~r1_orders_2(X12,X13,X14)|~r2_orders_2(X12,X14,X13))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_orders_2])])])).
fof(c_0_20, plain, ![X26, X27, X28]:(((r1_orders_2(X26,X27,X28)|~r2_orders_2(X26,X27,X28)|~m1_subset_1(X28,u1_struct_0(X26))|~m1_subset_1(X27,u1_struct_0(X26))|~l1_orders_2(X26))&(X27!=X28|~r2_orders_2(X26,X27,X28)|~m1_subset_1(X28,u1_struct_0(X26))|~m1_subset_1(X27,u1_struct_0(X26))|~l1_orders_2(X26)))&(~r1_orders_2(X26,X27,X28)|X27=X28|r2_orders_2(X26,X27,X28)|~m1_subset_1(X28,u1_struct_0(X26))|~m1_subset_1(X27,u1_struct_0(X26))|~l1_orders_2(X26))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_orders_2])])])])])).
cnf(c_0_21, plain, (r1_orders_2(X2,esk8_3(X2,X3,X4),X1)|r2_lattice3(X2,X3,esk7_3(X2,X3,X4))|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, plain, (r1_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_24, negated_conjecture, (r1_orders_2(esk1_0,esk4_0,X1)|r1_yellow_0(esk1_0,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_25, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_0)|r1_yellow_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_26, plain, (~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~r2_orders_2(X1,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (X2=X3|r2_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,X1,esk4_0),X2)|r2_lattice3(esk1_0,X1,esk7_3(esk1_0,X1,esk4_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,esk4_0)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_9])])).
cnf(c_0_29, negated_conjecture, (r1_yellow_0(esk1_0,X1)|esk8_3(esk1_0,X1,esk4_0)!=esk4_0|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_22]), c_0_9])]), c_0_18])).
cnf(c_0_30, plain, (r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_31, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk4_0)), inference(sr,[status(thm)],[c_0_25, c_0_18])).
cnf(c_0_32, plain, (X1=X2|~r1_orders_2(X3,X2,X1)|~r1_orders_2(X3,X1,X2)|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(X3))|~l1_orders_2(X3)|~v5_orders_2(X3)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,X1,esk4_0),esk4_0)|r2_lattice3(esk1_0,X1,esk7_3(esk1_0,X1,esk4_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,X1,esk4_0)), inference(spm,[status(thm)],[c_0_28, c_0_22])).
cnf(c_0_35, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk4_0)!=esk4_0|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_22]), c_0_9])]), c_0_18])).
cnf(c_0_36, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|esk8_3(X1,X2,X3)!=X3|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_37, negated_conjecture, (X1=esk4_0|~r1_orders_2(esk1_0,X1,esk4_0)|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_24]), c_0_22]), c_0_9]), c_0_33])]), c_0_18])).
cnf(c_0_38, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,esk2_0,esk4_0),esk4_0)|r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_31]), c_0_18])).
cnf(c_0_39, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk4_0)!=esk4_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_31]), c_0_22]), c_0_9])]), c_0_18])).
cnf(c_0_40, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_41, plain, (r1_orders_2(X2,esk8_3(X2,X3,X4),X1)|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_orders_2(X2,X4,esk7_3(X2,X3,X4))|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_42, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk7_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_43, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_44, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk8_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_24]), c_0_22]), c_0_9])]), c_0_18])).
cnf(c_0_46, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_47, negated_conjecture, (r1_orders_2(esk1_0,esk8_3(esk1_0,X1,esk4_0),X2)|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_24]), c_0_22]), c_0_9])]), c_0_18])).
cnf(c_0_48, negated_conjecture, (r2_lattice3(esk1_0,X1,esk8_3(esk1_0,X1,esk4_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_24]), c_0_22]), c_0_9])]), c_0_18])).
cnf(c_0_49, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_31]), c_0_22]), c_0_9])]), c_0_18])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk8_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_31]), c_0_22]), c_0_9])]), c_0_18])).
cnf(c_0_51, plain, (r1_orders_2(X2,esk8_3(X2,X3,X4),X1)|m1_subset_1(esk7_3(X2,X3,X4),u1_struct_0(X2))|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_52, negated_conjecture, (r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,X1,esk4_0))|~r2_lattice3(esk1_0,X1,esk4_0)|~m1_subset_1(esk7_3(esk1_0,X1,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_47]), c_0_22])]), c_0_45]), c_0_29])).
cnf(c_0_53, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_31])]), c_0_18]), c_0_50])).
cnf(c_0_54, negated_conjecture, (esk8_3(esk1_0,X1,X2)=esk4_0|m1_subset_1(esk7_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|r1_yellow_0(esk1_0,X1)|~r2_lattice3(esk1_0,esk2_0,esk8_3(esk1_0,X1,X2))|~r2_lattice3(esk1_0,X1,esk4_0)|~r2_lattice3(esk1_0,X1,X2)|~m1_subset_1(esk8_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_51]), c_0_22]), c_0_9])])).
cnf(c_0_55, plain, (r2_lattice3(X1,X2,esk8_3(X1,X2,X3))|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_56, negated_conjecture, (~r2_lattice3(esk1_0,esk2_0,esk7_3(esk1_0,esk2_0,esk4_0))|~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_31])]), c_0_18])).
cnf(c_0_57, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk4_0|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(esk8_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_31]), c_0_9])]), c_0_18])).
cnf(c_0_58, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_59, negated_conjecture, (~m1_subset_1(esk7_3(esk1_0,esk2_0,esk4_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_49]), c_0_50])).
cnf(c_0_60, negated_conjecture, (esk8_3(esk1_0,esk2_0,X1)=esk4_0|m1_subset_1(esk7_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r2_lattice3(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_9])]), c_0_18])).
cnf(c_0_61, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_31]), c_0_22])]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 62
# Proof object clause steps            : 53
# Proof object formula steps           : 9
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 80
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 112
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 99
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 99
# ...of the previous two non-trivial   : 82
# Contextual simplify-reflections      : 12
# Paramodulations                      : 96
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 62
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 55
# Current number of unprocessed clauses: 26
# ...number of literals in the above   : 261
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 1331
# Rec. Clause-clause subsumption calls : 48
# Non-unit clause-clause subsumptions  : 18
# Unit Clause-clause subsumption calls : 36
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6976

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.003 s
# Total time               : 0.013 s
# Maximum resident set size: 28564 pages
