# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.005 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((k1_relat_1(X1)=k1_relat_1(X2)&![X3]:(r2_hidden(X3,k1_relat_1(X1))=>k1_funct_1(X1,X3)=k1_funct_1(X2,X3)))=>X1=X2))), file('data/MPTP/proofs/cacfd28d0d664aa7b53010f988b956be.E_input', t9_funct_1)).
fof(d2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(X1=X2<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X1)<=>r2_hidden(k4_tarski(X3,X4),X2))))), file('data/MPTP/proofs/cacfd28d0d664aa7b53010f988b956be.E_input', d2_relat_1)).
fof(t8_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(k4_tarski(X1,X2),X3)<=>(r2_hidden(X1,k1_relat_1(X3))&X2=k1_funct_1(X3,X1)))), file('data/MPTP/proofs/cacfd28d0d664aa7b53010f988b956be.E_input', t8_funct_1)).
fof(t20_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X2,k2_relat_1(X3))))), file('data/MPTP/proofs/cacfd28d0d664aa7b53010f988b956be.E_input', t20_relat_1)).
fof(c_0_4, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((k1_relat_1(X1)=k1_relat_1(X2)&![X3]:(r2_hidden(X3,k1_relat_1(X1))=>k1_funct_1(X1,X3)=k1_funct_1(X2,X3)))=>X1=X2)))), inference(assume_negation,[status(cth)],[t9_funct_1])).
fof(c_0_5, plain, ![X8, X9, X10, X11, X12, X13]:(((~r2_hidden(k4_tarski(X10,X11),X8)|r2_hidden(k4_tarski(X10,X11),X9)|X8!=X9|~v1_relat_1(X9)|~v1_relat_1(X8))&(~r2_hidden(k4_tarski(X12,X13),X9)|r2_hidden(k4_tarski(X12,X13),X8)|X8!=X9|~v1_relat_1(X9)|~v1_relat_1(X8)))&((~r2_hidden(k4_tarski(esk3_2(X8,X9),esk4_2(X8,X9)),X8)|~r2_hidden(k4_tarski(esk3_2(X8,X9),esk4_2(X8,X9)),X9)|X8=X9|~v1_relat_1(X9)|~v1_relat_1(X8))&(r2_hidden(k4_tarski(esk3_2(X8,X9),esk4_2(X8,X9)),X8)|r2_hidden(k4_tarski(esk3_2(X8,X9),esk4_2(X8,X9)),X9)|X8=X9|~v1_relat_1(X9)|~v1_relat_1(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_relat_1])])])])])])).
fof(c_0_6, negated_conjecture, ![X7]:((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((k1_relat_1(esk1_0)=k1_relat_1(esk2_0)&(~r2_hidden(X7,k1_relat_1(esk1_0))|k1_funct_1(esk1_0,X7)=k1_funct_1(esk2_0,X7)))&esk1_0!=esk2_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
cnf(c_0_7, plain, (r2_hidden(k4_tarski(esk3_2(X1,X2),esk4_2(X1,X2)),X1)|r2_hidden(k4_tarski(esk3_2(X1,X2),esk4_2(X1,X2)),X2)|X1=X2|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, plain, ![X16, X17, X18]:(((r2_hidden(X16,k1_relat_1(X18))|~r2_hidden(k4_tarski(X16,X17),X18)|(~v1_relat_1(X18)|~v1_funct_1(X18)))&(X17=k1_funct_1(X18,X16)|~r2_hidden(k4_tarski(X16,X17),X18)|(~v1_relat_1(X18)|~v1_funct_1(X18))))&(~r2_hidden(X16,k1_relat_1(X18))|X17!=k1_funct_1(X18,X16)|r2_hidden(k4_tarski(X16,X17),X18)|(~v1_relat_1(X18)|~v1_funct_1(X18)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_funct_1])])])).
cnf(c_0_10, negated_conjecture, (X1=esk2_0|r2_hidden(k4_tarski(esk3_2(X1,esk2_0),esk4_2(X1,esk2_0)),esk2_0)|r2_hidden(k4_tarski(esk3_2(X1,esk2_0),esk4_2(X1,esk2_0)),X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_11, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_13, plain, ![X19, X20, X21]:((r2_hidden(X19,k1_relat_1(X21))|~r2_hidden(k4_tarski(X19,X20),X21)|~v1_relat_1(X21))&(r2_hidden(X20,k2_relat_1(X21))|~r2_hidden(k4_tarski(X19,X20),X21)|~v1_relat_1(X21))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
cnf(c_0_14, plain, (X1=k1_funct_1(X2,X3)|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, negated_conjecture, (r2_hidden(k4_tarski(esk3_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),esk1_0)|r2_hidden(k4_tarski(esk3_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12])).
cnf(c_0_16, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_17, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (k1_relat_1(esk1_0)=k1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_19, plain, (r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|r2_hidden(k4_tarski(esk3_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16]), c_0_8])])).
cnf(c_0_21, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_22, negated_conjecture, (r2_hidden(k4_tarski(esk3_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),esk1_0)|r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_15]), c_0_18]), c_0_8])])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,k1_funct_1(X2,X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (k1_funct_1(esk1_0,X1)=k1_funct_1(esk2_0,X1)|~r2_hidden(X1,k1_relat_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25, negated_conjecture, (k1_funct_1(esk2_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)|k1_funct_1(esk1_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_20]), c_0_21]), c_0_11])])).
cnf(c_0_26, negated_conjecture, (r2_hidden(esk3_2(esk1_0,esk2_0),k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_22]), c_0_11])])).
cnf(c_0_27, negated_conjecture, (r2_hidden(k4_tarski(X1,k1_funct_1(esk1_0,X1)),esk2_0)|~r2_hidden(X1,k1_relat_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_18]), c_0_16]), c_0_8])])).
cnf(c_0_28, negated_conjecture, (k1_funct_1(esk1_0,esk3_2(esk1_0,esk2_0))=esk4_2(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
cnf(c_0_29, plain, (X1=X2|~r2_hidden(k4_tarski(esk3_2(X1,X2),esk4_2(X1,X2)),X1)|~r2_hidden(k4_tarski(esk3_2(X1,X2),esk4_2(X1,X2)),X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k4_tarski(esk3_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_26])])).
cnf(c_0_31, negated_conjecture, (r2_hidden(k4_tarski(esk3_2(esk1_0,esk2_0),esk4_2(esk1_0,esk2_0)),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_28]), c_0_26]), c_0_21]), c_0_11])])).
cnf(c_0_32, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_8]), c_0_11])]), c_0_12]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 24
# Proof object formula steps           : 9
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 30
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 14
# Processed clauses                    : 46
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 43
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 3
# Generated clauses                    : 36
# ...of the previous two non-trivial   : 26
# Contextual simplify-reflections      : 0
# Paramodulations                      : 35
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 4
# ...number of literals in the above   : 8
# Current number of archived formulas  : 0
# Current number of archived clauses   : 18
# Clause-clause subsumption calls (NU) : 20
# Rec. Clause-clause subsumption calls : 11
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1913

# -------------------------------------------------
# User time                : 0.003 s
# System time              : 0.003 s
# Total time               : 0.006 s
# Maximum resident set size: 28844 pages
