# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t25_yellow_0, conjecture, ![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X2=k12_lattice3(X1,X2,X3)<=>r3_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/cb837c5f675d4ae29aa0ce82e2724652.E_input', t25_yellow_0)).
fof(redefinition_k12_lattice3, axiom, ![X1, X2, X3]:(((v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)), file('data/MPTP/proofs/cb837c5f675d4ae29aa0ce82e2724652.E_input', redefinition_k12_lattice3)).
fof(redefinition_r3_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))&(m1_subset_1(X2,u1_struct_0(X1))&m1_subset_1(X3,u1_struct_0(X1))))=>(r3_orders_2(X1,X2,X3)<=>r1_orders_2(X1,X2,X3))), file('data/MPTP/proofs/cb837c5f675d4ae29aa0ce82e2724652.E_input', redefinition_r3_orders_2)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('data/MPTP/proofs/cb837c5f675d4ae29aa0ce82e2724652.E_input', cc2_lattice3)).
fof(t21_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>(v2_lattice3(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r2_yellow_0(X1,k2_tarski(X2,X3)))))), file('data/MPTP/proofs/cb837c5f675d4ae29aa0ce82e2724652.E_input', t21_yellow_0)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/cb837c5f675d4ae29aa0ce82e2724652.E_input', t19_yellow_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/cb837c5f675d4ae29aa0ce82e2724652.E_input', t24_orders_2)).
fof(c_0_7, negated_conjecture, ~(![X1]:((v3_orders_2(X1)&(v5_orders_2(X1)&(v2_lattice3(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X2=k12_lattice3(X1,X2,X3)<=>r3_orders_2(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t25_yellow_0])).
fof(c_0_8, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_9, negated_conjecture, ((v3_orders_2(esk1_0)&(v5_orders_2(esk1_0)&(v2_lattice3(esk1_0)&l1_orders_2(esk1_0))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((esk2_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r3_orders_2(esk1_0,esk2_0,esk3_0))&(esk2_0=k12_lattice3(esk1_0,esk2_0,esk3_0)|r3_orders_2(esk1_0,esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
fof(c_0_10, plain, ![X22, X23, X24]:(~v5_orders_2(X22)|(~v2_lattice3(X22)|~l1_orders_2(X22))|(~m1_subset_1(X23,u1_struct_0(X22))|~m1_subset_1(X24,u1_struct_0(X22)))|k12_lattice3(X22,X23,X24)=k11_lattice3(X22,X23,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
fof(c_0_11, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_8])).
fof(c_0_12, plain, ![X14, X15, X16]:((~r3_orders_2(X14,X15,X16)|r1_orders_2(X14,X15,X16)|(v2_struct_0(X14)|(~v3_orders_2(X14)|~l1_orders_2(X14))|(~m1_subset_1(X15,u1_struct_0(X14))|~m1_subset_1(X16,u1_struct_0(X14)))))&(~r1_orders_2(X14,X15,X16)|r3_orders_2(X14,X15,X16)|(v2_struct_0(X14)|(~v3_orders_2(X14)|~l1_orders_2(X14))|(~m1_subset_1(X15,u1_struct_0(X14))|~m1_subset_1(X16,u1_struct_0(X14)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r3_orders_2])])])])).
cnf(c_0_13, negated_conjecture, (esk2_0=k12_lattice3(esk1_0,esk2_0,esk3_0)|r3_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_18, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_20, plain, ![X26, X27, X28, X29, X30]:(((r1_orders_2(X28,X29,X27)|(X29!=k11_lattice3(X28,X27,X26)|~r2_yellow_0(X28,k2_tarski(X27,X26)))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26))&((r1_orders_2(X28,X29,X26)|(X29!=k11_lattice3(X28,X27,X26)|~r2_yellow_0(X28,k2_tarski(X27,X26)))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26))&(~m1_subset_1(X30,u1_struct_0(X28))|(~r1_orders_2(X28,X30,X27)|~r1_orders_2(X28,X30,X26)|r1_orders_2(X28,X30,X29))|(X29!=k11_lattice3(X28,X27,X26)|~r2_yellow_0(X28,k2_tarski(X27,X26)))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26))))&(((X29=k11_lattice3(X28,X27,X26)|(m1_subset_1(esk6_4(X26,X27,X28,X29),u1_struct_0(X28))|~r1_orders_2(X28,X29,X26)|~r1_orders_2(X28,X29,X27))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26))&(r2_yellow_0(X28,k2_tarski(X27,X26))|(m1_subset_1(esk6_4(X26,X27,X28,X29),u1_struct_0(X28))|~r1_orders_2(X28,X29,X26)|~r1_orders_2(X28,X29,X27))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26)))&((((X29=k11_lattice3(X28,X27,X26)|(r1_orders_2(X28,esk6_4(X26,X27,X28,X29),X27)|~r1_orders_2(X28,X29,X26)|~r1_orders_2(X28,X29,X27))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26))&(r2_yellow_0(X28,k2_tarski(X27,X26))|(r1_orders_2(X28,esk6_4(X26,X27,X28,X29),X27)|~r1_orders_2(X28,X29,X26)|~r1_orders_2(X28,X29,X27))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26)))&((X29=k11_lattice3(X28,X27,X26)|(r1_orders_2(X28,esk6_4(X26,X27,X28,X29),X26)|~r1_orders_2(X28,X29,X26)|~r1_orders_2(X28,X29,X27))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26))&(r2_yellow_0(X28,k2_tarski(X27,X26))|(r1_orders_2(X28,esk6_4(X26,X27,X28,X29),X26)|~r1_orders_2(X28,X29,X26)|~r1_orders_2(X28,X29,X27))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26))))&((X29=k11_lattice3(X28,X27,X26)|(~r1_orders_2(X28,esk6_4(X26,X27,X28,X29),X29)|~r1_orders_2(X28,X29,X26)|~r1_orders_2(X28,X29,X27))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26))&(r2_yellow_0(X28,k2_tarski(X27,X26))|(~r1_orders_2(X28,esk6_4(X26,X27,X28,X29),X29)|~r1_orders_2(X28,X29,X26)|~r1_orders_2(X28,X29,X27))|~m1_subset_1(X29,u1_struct_0(X28))|~epred1_3(X28,X27,X26)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])])).
fof(c_0_21, plain, ![X25]:(~l1_orders_2(X25)|(~v2_lattice3(X25)|~v2_struct_0(X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
cnf(c_0_22, plain, (r1_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (k11_lattice3(esk1_0,esk2_0,esk3_0)=esk2_0|r3_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_24, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, plain, (r1_orders_2(X1,X2,X3)|X2!=k11_lattice3(X1,X4,X3)|~r2_yellow_0(X1,k2_tarski(X4,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~epred1_3(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (k11_lattice3(esk1_0,esk2_0,esk3_0)=esk2_0|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_15]), c_0_16]), c_0_17]), c_0_24])])).
fof(c_0_28, plain, ![X17, X18, X19]:((~v2_lattice3(X17)|(~m1_subset_1(X18,u1_struct_0(X17))|(~m1_subset_1(X19,u1_struct_0(X17))|r2_yellow_0(X17,k2_tarski(X18,X19))))|(~v5_orders_2(X17)|~l1_orders_2(X17)))&((m1_subset_1(esk4_1(X17),u1_struct_0(X17))|v2_lattice3(X17)|(~v5_orders_2(X17)|~l1_orders_2(X17)))&((m1_subset_1(esk5_1(X17),u1_struct_0(X17))|v2_lattice3(X17)|(~v5_orders_2(X17)|~l1_orders_2(X17)))&(~r2_yellow_0(X17,k2_tarski(esk4_1(X17),esk5_1(X17)))|v2_lattice3(X17)|(~v5_orders_2(X17)|~l1_orders_2(X17)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_yellow_0])])])])])).
fof(c_0_29, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_8])).
cnf(c_0_30, negated_conjecture, (esk2_0!=k12_lattice3(esk1_0,esk2_0,esk3_0)|~r3_orders_2(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_31, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,X1)|esk2_0!=k11_lattice3(esk1_0,X2,X1)|~epred1_3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,k2_tarski(X2,X1))), inference(spm,[status(thm)],[c_0_25, c_0_16])).
cnf(c_0_32, negated_conjecture, (k11_lattice3(esk1_0,esk2_0,esk3_0)=esk2_0|r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_17]), c_0_18])])).
cnf(c_0_33, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|~v2_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_34, plain, ![X9, X10, X11]:(~v5_orders_2(X9)|~l1_orders_2(X9)|(~m1_subset_1(X10,u1_struct_0(X9))|(~m1_subset_1(X11,u1_struct_0(X9))|epred1_3(X9,X10,X11)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])).
cnf(c_0_35, negated_conjecture, (k11_lattice3(esk1_0,esk2_0,esk3_0)!=esk2_0|~r3_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_14]), c_0_15]), c_0_16]), c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_36, plain, (r3_orders_2(X1,X2,X3)|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|~epred1_3(esk1_0,esk2_0,esk3_0)|~r2_yellow_0(esk1_0,k2_tarski(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, negated_conjecture, (r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_18]), c_0_17]), c_0_19])])).
cnf(c_0_39, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_40, plain, ![X12, X13]:(v2_struct_0(X12)|(~v3_orders_2(X12)|~l1_orders_2(X12))|(~m1_subset_1(X13,u1_struct_0(X12))|r1_orders_2(X12,X13,X13))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_41, negated_conjecture, (v2_struct_0(esk1_0)|k11_lattice3(esk1_0,esk2_0,esk3_0)!=esk2_0|~r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_15]), c_0_16]), c_0_17]), c_0_24])])).
cnf(c_0_42, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)|~epred1_3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_15]), c_0_16])])).
cnf(c_0_43, negated_conjecture, (epred1_3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_19]), c_0_17])])).
cnf(c_0_44, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45, negated_conjecture, (k11_lattice3(esk1_0,esk2_0,esk3_0)!=esk2_0|~r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_41]), c_0_17]), c_0_18])])).
cnf(c_0_46, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_15]), c_0_16])])).
cnf(c_0_47, plain, (X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk6_4(X4,X3,X2,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_48, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk6_4(X4,X3,X2,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_49, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_16]), c_0_17]), c_0_24])])).
cnf(c_0_50, negated_conjecture, (k11_lattice3(esk1_0,esk2_0,esk3_0)!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])])).
cnf(c_0_51, plain, (k11_lattice3(X1,X2,X3)=X2|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_49]), c_0_17]), c_0_18])])).
cnf(c_0_53, negated_conjecture, (~epred1_3(esk1_0,esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_46]), c_0_52]), c_0_16])])).
cnf(c_0_54, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_43]), c_0_15]), c_0_16])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 36
# Proof object formula steps           : 19
# Proof object conjectures             : 28
# Proof object clause conjectures      : 25
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 54
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 90
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 88
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 8
# Generated clauses                    : 68
# ...of the previous two non-trivial   : 63
# Contextual simplify-reflections      : 0
# Paramodulations                      : 67
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 46
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 33
# Current number of unprocessed clauses: 31
# ...number of literals in the above   : 240
# Current number of archived formulas  : 0
# Current number of archived clauses   : 42
# Clause-clause subsumption calls (NU) : 361
# Rec. Clause-clause subsumption calls : 47
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4308

# -------------------------------------------------
# User time                : 0.011 s
# System time              : 0.000 s
# Total time               : 0.011 s
# Maximum resident set size: 28572 pages
