# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_orders_2(X1,X2,X3)<=>(r1_orders_2(X1,X2,X3)&~(X2=X3)))))), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', d10_orders_2)).
fof(fraenkel_a_2_1_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&l1_orders_2(X2)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_orders_2(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r2_hidden(X5,X3)=>r2_orders_2(X2,X4,X5))))))), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', fraenkel_a_2_1_orders_2)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', t4_subset)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', t2_subset)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', d4_subset_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', t5_subset)).
fof(dt_k2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', dt_k2_orders_2)).
fof(d13_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2))), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', d13_orders_2)).
fof(t5_mcart_1, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~((r2_hidden(X2,X1)&![X3, X4, X5, X6]:((r2_hidden(X3,X4)&(r2_hidden(X4,X5)&(r2_hidden(X5,X6)&r2_hidden(X6,X2))))=>r1_xboole_0(X3,X1)))))), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', t5_mcart_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', d3_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', dt_l1_orders_2)).
fof(t46_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k2_struct_0(X1))=k1_xboole_0), file('data/MPTP/proofs/cef44e0eca014e3c872f4bc5b964220d.E_input', t46_orders_2)).
fof(c_0_13, plain, ![X7, X8, X9]:(((r1_orders_2(X7,X8,X9)|~r2_orders_2(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7))&(X8!=X9|~r2_orders_2(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7)))&(~r1_orders_2(X7,X8,X9)|X8=X9|r2_orders_2(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_orders_2])])])])])).
fof(c_0_14, plain, ![X10, X11, X12, X14, X15]:(((m1_subset_1(esk1_3(X10,X11,X12),u1_struct_0(X11))|~r2_hidden(X10,a_2_1_orders_2(X11,X12))|(v2_struct_0(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|~l1_orders_2(X11))))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))))&((X10=esk1_3(X10,X11,X12)|~r2_hidden(X10,a_2_1_orders_2(X11,X12))|(v2_struct_0(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|~l1_orders_2(X11))))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))))&(~m1_subset_1(X14,u1_struct_0(X11))|(~r2_hidden(X14,X12)|r2_orders_2(X11,esk1_3(X10,X11,X12),X14))|~r2_hidden(X10,a_2_1_orders_2(X11,X12))|(v2_struct_0(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|~l1_orders_2(X11))))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))))))&((m1_subset_1(esk2_4(X10,X11,X12,X15),u1_struct_0(X11))|X10!=X15|~m1_subset_1(X15,u1_struct_0(X11))|r2_hidden(X10,a_2_1_orders_2(X11,X12))|(v2_struct_0(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|~l1_orders_2(X11))))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))))&((r2_hidden(esk2_4(X10,X11,X12,X15),X12)|X10!=X15|~m1_subset_1(X15,u1_struct_0(X11))|r2_hidden(X10,a_2_1_orders_2(X11,X12))|(v2_struct_0(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|~l1_orders_2(X11))))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))))&(~r2_orders_2(X11,X15,esk2_4(X10,X11,X12,X15))|X10!=X15|~m1_subset_1(X15,u1_struct_0(X11))|r2_hidden(X10,a_2_1_orders_2(X11,X12))|(v2_struct_0(X11)|(~v3_orders_2(X11)|(~v4_orders_2(X11)|(~v5_orders_2(X11)|~l1_orders_2(X11))))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(X11)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_orders_2])])])])])])).
fof(c_0_15, plain, ![X17, X18, X19]:(~r2_hidden(X17,X18)|~m1_subset_1(X18,k1_zfmisc_1(X19))|m1_subset_1(X17,X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_16, plain, (X1!=X2|~r2_orders_2(X3,X1,X2)|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (r2_orders_2(X2,esk1_3(X4,X2,X3),X1)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_hidden(X4,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (~r2_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (v2_struct_0(X1)|r2_orders_2(X1,esk1_3(X2,X1,X3),X4)|~r2_hidden(X2,a_2_1_orders_2(X1,X3))|~r2_hidden(X4,X3)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_21, plain, ![X30, X31]:(~m1_subset_1(X30,X31)|(v1_xboole_0(X31)|r2_hidden(X30,X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_22, plain, ![X20]:m1_subset_1(k2_subset_1(X20),k1_zfmisc_1(X20)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_23, plain, ![X21]:k2_subset_1(X21)=X21, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_24, plain, ![X22, X23, X24]:(~r2_hidden(X22,X23)|(~m1_subset_1(X23,k1_zfmisc_1(X24))|~v1_xboole_0(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_25, plain, ![X25, X26]:(v2_struct_0(X25)|(~v3_orders_2(X25)|(~v4_orders_2(X25)|(~v5_orders_2(X25)|~l1_orders_2(X25))))|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|m1_subset_1(k2_orders_2(X25,X26),k1_zfmisc_1(u1_struct_0(X25)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_orders_2])])])).
cnf(c_0_26, plain, (v2_struct_0(X1)|~r2_hidden(esk1_3(X2,X1,X3),X3)|~r2_hidden(X2,a_2_1_orders_2(X1,X3))|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_18])).
cnf(c_0_27, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, ![X27, X28]:(v2_struct_0(X27)|(~v3_orders_2(X27)|(~v4_orders_2(X27)|(~v5_orders_2(X27)|~l1_orders_2(X27))))|(~m1_subset_1(X28,k1_zfmisc_1(u1_struct_0(X27)))|k2_orders_2(X27,X28)=a_2_1_orders_2(X27,X28))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_orders_2])])])])).
cnf(c_0_31, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (v2_struct_0(X1)|m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_33, plain, ![X33, X35, X36, X37, X38]:((r2_hidden(esk4_1(X33),X33)|X33=k1_xboole_0)&(~r2_hidden(X35,X36)|(~r2_hidden(X36,X37)|(~r2_hidden(X37,X38)|~r2_hidden(X38,esk4_1(X33))))|r1_xboole_0(X35,X33)|X33=k1_xboole_0)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_mcart_1])])])])])])])).
cnf(c_0_34, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|~r2_hidden(X3,a_2_1_orders_2(X2,X1))|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk1_3(X3,X2,X1),X1)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_35, plain, (m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_36, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_37, plain, (v2_struct_0(X1)|k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,k2_orders_2(X1,X3))|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39, plain, (r2_hidden(esk4_1(X1),X1)|X1=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_orders_2(X1,u1_struct_0(X1)))|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_41, plain, (a_2_1_orders_2(X1,u1_struct_0(X1))=k2_orders_2(X1,u1_struct_0(X1))|v2_struct_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_37, c_0_36])).
fof(c_0_42, plain, ![X32]:(~l1_struct_0(X32)|k2_struct_0(X32)=u1_struct_0(X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_43, plain, ![X39]:(~l1_orders_2(X39)|l1_struct_0(X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_44, plain, (k2_orders_2(X1,X2)=k1_xboole_0|v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
fof(c_0_45, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k2_struct_0(X1))=k1_xboole_0)), inference(assume_negation,[status(cth)],[t46_orders_2])).
cnf(c_0_46, plain, (v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~r2_hidden(X2,k2_orders_2(X1,u1_struct_0(X1)))|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_48, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (k2_orders_2(X1,u1_struct_0(X1))=k1_xboole_0|v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_44, c_0_36])).
fof(c_0_50, negated_conjecture, ((~v2_struct_0(esk3_0)&(v3_orders_2(esk3_0)&(v4_orders_2(esk3_0)&(v5_orders_2(esk3_0)&l1_orders_2(esk3_0)))))&k2_orders_2(esk3_0,k2_struct_0(esk3_0))!=k1_xboole_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_45])])])])).
cnf(c_0_51, plain, (v1_xboole_0(k2_struct_0(X1))|v2_struct_0(X1)|~r2_hidden(X2,k2_orders_2(X1,k2_struct_0(X1)))|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_52, plain, (k2_orders_2(X1,k2_struct_0(X1))=k1_xboole_0|v2_struct_0(X1)|~v1_xboole_0(k2_struct_0(X1))|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_47]), c_0_48])).
cnf(c_0_53, negated_conjecture, (k2_orders_2(esk3_0,k2_struct_0(esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_54, plain, (k2_orders_2(X1,k2_struct_0(X1))=k1_xboole_0|v2_struct_0(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_39]), c_0_52])).
cnf(c_0_55, negated_conjecture, (v5_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_56, negated_conjecture, (v4_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, negated_conjecture, (v3_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_58, negated_conjecture, (l1_orders_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_60, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 61
# Proof object clause steps            : 34
# Proof object formula steps           : 27
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 25
# Processed clauses                    : 202
# ...of these trivial                  : 0
# ...subsumed                          : 82
# ...remaining for further processing  : 120
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 340
# ...of the previous two non-trivial   : 319
# Contextual simplify-reflections      : 25
# Paramodulations                      : 336
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 113
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 105
# Current number of unprocessed clauses: 142
# ...number of literals in the above   : 1283
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4
# Clause-clause subsumption calls (NU) : 6895
# Rec. Clause-clause subsumption calls : 851
# Non-unit clause-clause subsumptions  : 110
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12517

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.003 s
# Total time               : 0.018 s
# Maximum resident set size: 28824 pages
