# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d13_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,k1_relat_1(X1))&r2_hidden(k1_funct_1(X1,X4),X2))))), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', d13_funct_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t2_xboole_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', d3_tarski)).
fof(t3_xboole_1, axiom, ![X1]:(r1_tarski(X1,k1_xboole_0)=>X1=k1_xboole_0), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t3_xboole_1)).
fof(t175_funct_2, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(r1_tarski(k10_relat_1(X1,X3),X2)=>k10_relat_1(X1,X3)=k10_relat_1(k7_relat_1(X1,X2),X3))), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t175_funct_2)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t12_setfam_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t70_enumset1)).
fof(t175_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(X3,k2_xboole_0(X1,X2))=k2_xboole_0(k10_relat_1(X3,X1),k10_relat_1(X3,X2))), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t175_relat_1)).
fof(t19_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X1,X3))=>r1_tarski(X1,k3_xboole_0(X2,X3))), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t19_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t72_enumset1)).
fof(t139_funct_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>k10_relat_1(k7_relat_1(X3,X1),X2)=k3_xboole_0(X1,k10_relat_1(X3,X2))), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t139_funct_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t12_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', d10_xboole_0)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', t17_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/d16bd250360e4ce3b278d013bc0172f7.E_input', commutativity_k2_tarski)).
fof(c_0_17, plain, ![X37, X38, X39, X40, X41, X42, X43]:((((r2_hidden(X40,k1_relat_1(X37))|~r2_hidden(X40,X39)|X39!=k10_relat_1(X37,X38)|(~v1_relat_1(X37)|~v1_funct_1(X37)))&(r2_hidden(k1_funct_1(X37,X40),X38)|~r2_hidden(X40,X39)|X39!=k10_relat_1(X37,X38)|(~v1_relat_1(X37)|~v1_funct_1(X37))))&(~r2_hidden(X41,k1_relat_1(X37))|~r2_hidden(k1_funct_1(X37,X41),X38)|r2_hidden(X41,X39)|X39!=k10_relat_1(X37,X38)|(~v1_relat_1(X37)|~v1_funct_1(X37))))&((~r2_hidden(esk5_3(X37,X42,X43),X43)|(~r2_hidden(esk5_3(X37,X42,X43),k1_relat_1(X37))|~r2_hidden(k1_funct_1(X37,esk5_3(X37,X42,X43)),X42))|X43=k10_relat_1(X37,X42)|(~v1_relat_1(X37)|~v1_funct_1(X37)))&((r2_hidden(esk5_3(X37,X42,X43),k1_relat_1(X37))|r2_hidden(esk5_3(X37,X42,X43),X43)|X43=k10_relat_1(X37,X42)|(~v1_relat_1(X37)|~v1_funct_1(X37)))&(r2_hidden(k1_funct_1(X37,esk5_3(X37,X42,X43)),X42)|r2_hidden(esk5_3(X37,X42,X43),X43)|X43=k10_relat_1(X37,X42)|(~v1_relat_1(X37)|~v1_funct_1(X37)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d13_funct_1])])])])])])).
fof(c_0_18, plain, ![X28, X29]:(~r2_hidden(X28,X29)|~r1_tarski(X29,X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_19, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~r2_hidden(X2,X4)|X4!=k10_relat_1(X1,X3)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (r2_hidden(k1_funct_1(X1,X2),X3)|~r2_hidden(X2,k10_relat_1(X1,X3))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_19])).
fof(c_0_22, plain, ![X60]:r1_tarski(k1_xboole_0,X60), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_23, plain, (~r2_hidden(X1,k10_relat_1(X2,X3))|~r1_tarski(X3,k1_funct_1(X2,X1))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_25, plain, ![X11, X12, X13, X14, X15]:((~r1_tarski(X11,X12)|(~r2_hidden(X13,X11)|r2_hidden(X13,X12)))&((r2_hidden(esk4_2(X14,X15),X14)|r1_tarski(X14,X15))&(~r2_hidden(esk4_2(X14,X15),X15)|r1_tarski(X14,X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_26, plain, ![X8]:(~r1_tarski(X8,k1_xboole_0)|X8=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_xboole_1])])).
cnf(c_0_27, plain, (~r2_hidden(X1,k10_relat_1(X2,k1_xboole_0))|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_29, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:(r1_tarski(k10_relat_1(X1,X3),X2)=>k10_relat_1(X1,X3)=k10_relat_1(k7_relat_1(X1,X2),X3)))), inference(assume_negation,[status(cth)],[t175_funct_2])).
cnf(c_0_30, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_31, plain, (r1_tarski(k10_relat_1(X1,k1_xboole_0),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_32, negated_conjecture, ((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&(r1_tarski(k10_relat_1(esk1_0,esk3_0),esk2_0)&k10_relat_1(esk1_0,esk3_0)!=k10_relat_1(k7_relat_1(esk1_0,esk2_0),esk3_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])).
fof(c_0_33, plain, ![X9, X10]:k1_setfam_1(k2_tarski(X9,X10))=k3_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_34, plain, ![X47, X48]:k1_enumset1(X47,X47,X48)=k2_tarski(X47,X48), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_35, plain, ![X32, X33, X34]:(~v1_relat_1(X34)|k10_relat_1(X34,k2_xboole_0(X32,X33))=k2_xboole_0(k10_relat_1(X34,X32),k10_relat_1(X34,X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t175_relat_1])])).
cnf(c_0_36, plain, (k10_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_39, plain, ![X61, X62, X63]:(~r1_tarski(X61,X62)|~r1_tarski(X61,X63)|r1_tarski(X61,k3_xboole_0(X62,X63))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t19_xboole_1])])).
cnf(c_0_40, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_41, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_42, plain, ![X54, X55, X56]:k2_enumset1(X54,X54,X55,X56)=k1_enumset1(X54,X55,X56), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_43, plain, ![X22, X23, X24, X25]:k3_enumset1(X22,X22,X23,X24,X25)=k2_enumset1(X22,X23,X24,X25), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_44, plain, ![X64, X65, X66]:(~v1_relat_1(X66)|k10_relat_1(k7_relat_1(X66,X64),X65)=k3_xboole_0(X64,k10_relat_1(X66,X65))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t139_funct_1])])).
fof(c_0_45, plain, ![X20, X21]:(~r1_tarski(X20,X21)|k2_xboole_0(X20,X21)=X21), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_46, plain, (k10_relat_1(X1,k2_xboole_0(X2,X3))=k2_xboole_0(k10_relat_1(X1,X2),k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, negated_conjecture, (k10_relat_1(esk1_0,k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])])).
fof(c_0_48, plain, ![X49, X50]:(((r1_tarski(X49,X50)|X49!=X50)&(r1_tarski(X50,X49)|X49!=X50))&(~r1_tarski(X49,X50)|~r1_tarski(X50,X49)|X49=X50)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_49, plain, (r1_tarski(X1,k3_xboole_0(X2,X3))|~r1_tarski(X1,X2)|~r1_tarski(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k3_xboole_0(X2,k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55, negated_conjecture, (k2_xboole_0(k1_xboole_0,k10_relat_1(esk1_0,X1))=k10_relat_1(esk1_0,k2_xboole_0(k1_xboole_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_38])])).
fof(c_0_56, plain, ![X30, X31]:r1_tarski(k3_xboole_0(X30,X31),X30), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_57, plain, ![X26, X27]:k2_tarski(X26,X27)=k2_tarski(X27,X26), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_58, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (r1_tarski(X1,k1_setfam_1(k3_enumset1(X2,X2,X2,X2,X3)))|~r1_tarski(X1,X3)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_60, plain, (k10_relat_1(k7_relat_1(X1,X2),X3)=k1_setfam_1(k3_enumset1(X2,X2,X2,X2,k10_relat_1(X1,X3)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_61, negated_conjecture, (k10_relat_1(esk1_0,k2_xboole_0(k1_xboole_0,X1))=k10_relat_1(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_24])])).
cnf(c_0_62, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64, plain, (k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2))=X3|~r1_tarski(k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)),X3)|~r1_tarski(X3,X2)|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, negated_conjecture, (k1_setfam_1(k3_enumset1(X1,X1,X1,X1,k10_relat_1(esk1_0,X2)))=k10_relat_1(k7_relat_1(esk1_0,X1),k2_xboole_0(k1_xboole_0,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_38])])).
cnf(c_0_66, plain, (r1_tarski(k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_67, plain, (k3_enumset1(X1,X1,X1,X1,X2)=k3_enumset1(X2,X2,X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_41]), c_0_41]), c_0_51]), c_0_51]), c_0_52]), c_0_52])).
cnf(c_0_68, negated_conjecture, (k10_relat_1(k7_relat_1(esk1_0,X1),k2_xboole_0(k1_xboole_0,X2))=X3|~r1_tarski(k10_relat_1(k7_relat_1(esk1_0,X1),k2_xboole_0(k1_xboole_0,X2)),X3)|~r1_tarski(X3,k10_relat_1(esk1_0,X2))|~r1_tarski(X3,X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_69, negated_conjecture, (k10_relat_1(k7_relat_1(esk1_0,X1),k2_xboole_0(k1_xboole_0,X2))=k10_relat_1(k7_relat_1(esk1_0,X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_65]), c_0_38])])).
cnf(c_0_70, plain, (r1_tarski(k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)),X2)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_71, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_72, negated_conjecture, (k10_relat_1(k7_relat_1(esk1_0,X1),X2)=X3|~r1_tarski(k10_relat_1(k7_relat_1(esk1_0,X1),X2),X3)|~r1_tarski(X3,k10_relat_1(esk1_0,X2))|~r1_tarski(X3,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_69]), c_0_69])).
cnf(c_0_73, plain, (r1_tarski(k10_relat_1(k7_relat_1(X1,X2),X3),k10_relat_1(X1,X3))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_70, c_0_60])).
cnf(c_0_74, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_75, negated_conjecture, (k10_relat_1(esk1_0,esk3_0)!=k10_relat_1(k7_relat_1(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_76, negated_conjecture, (k10_relat_1(k7_relat_1(esk1_0,X1),X2)=k10_relat_1(esk1_0,X2)|~r1_tarski(k10_relat_1(esk1_0,X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74]), c_0_38])])).
cnf(c_0_77, negated_conjecture, (r1_tarski(k10_relat_1(esk1_0,esk3_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_78, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 79
# Proof object clause steps            : 44
# Proof object formula steps           : 35
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 35
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 37
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 33
# Processed clauses                    : 41169
# ...of these trivial                  : 215
# ...subsumed                          : 37153
# ...remaining for further processing  : 3801
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 361
# Backward-rewritten                   : 306
# Generated clauses                    : 538083
# ...of the previous two non-trivial   : 468740
# Contextual simplify-reflections      : 262
# Paramodulations                      : 538068
# Factorizations                       : 10
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3098
#    Positive orientable unit clauses  : 277
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 49
#    Non-unit-clauses                  : 2770
# Current number of unprocessed clauses: 424193
# ...number of literals in the above   : 1821852
# Current number of archived formulas  : 0
# Current number of archived clauses   : 702
# Clause-clause subsumption calls (NU) : 1606176
# Rec. Clause-clause subsumption calls : 741149
# Non-unit clause-clause subsumptions  : 24206
# Unit Clause-clause subsumption calls : 9597
# Rewrite failures with RHS unbound    : 5
# BW rewrite match attempts            : 4297
# BW rewrite match successes           : 128
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12700613

# -------------------------------------------------
# User time                : 7.524 s
# System time              : 0.129 s
# Total time               : 7.653 s
# Maximum resident set size: 322740 pages
