# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t70_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X7))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k6_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t70_mcart_1)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t50_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', d3_zfmisc_1)).
fof(t48_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4)))))))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t48_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', d3_mcart_1)).
fof(dt_k7_mcart_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>m1_subset_1(k7_mcart_1(X1,X2,X3,X4),X3)), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', dt_k7_mcart_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', d1_xboole_0)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t10_mcart_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t7_xboole_0)).
fof(d4_zfmisc_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', d4_zfmisc_1)).
fof(t55_mcart_1, axiom, ![X1, X2, X3, X4]:((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(X4=k1_xboole_0))))<=>~(k4_zfmisc_1(X1,X2,X3,X4)=k1_xboole_0)), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t55_mcart_1)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t70_enumset1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t113_zfmisc_1)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/d2ae7edf96ed428b8e9728a4f71d7f32.E_input', t152_zfmisc_1)).
fof(c_0_16, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X7))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k6_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t70_mcart_1])).
fof(c_0_17, plain, ![X44, X45, X46, X47]:((k5_mcart_1(X44,X45,X46,X47)=k1_mcart_1(k1_mcart_1(X47))|~m1_subset_1(X47,k3_zfmisc_1(X44,X45,X46))|X46=k1_xboole_0|X45=k1_xboole_0|X44=k1_xboole_0)&((k6_mcart_1(X44,X45,X46,X47)=k2_mcart_1(k1_mcart_1(X47))|~m1_subset_1(X47,k3_zfmisc_1(X44,X45,X46))|X46=k1_xboole_0|X45=k1_xboole_0|X44=k1_xboole_0)&(k7_mcart_1(X44,X45,X46,X47)=k2_mcart_1(X47)|~m1_subset_1(X47,k3_zfmisc_1(X44,X45,X46))|X46=k1_xboole_0|X45=k1_xboole_0|X44=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
fof(c_0_18, plain, ![X53, X54, X55]:k3_zfmisc_1(X53,X54,X55)=k2_zfmisc_1(k2_zfmisc_1(X53,X54),X55), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_19, negated_conjecture, ![X14, X15, X16]:(m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&((~m1_subset_1(X14,esk1_0)|(~m1_subset_1(X15,esk2_0)|(~m1_subset_1(X16,esk3_0)|(esk5_0!=k3_mcart_1(X14,X15,X16)|esk4_0=X15))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&esk4_0!=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])])).
fof(c_0_20, plain, ![X59, X60, X61, X62]:(X59=k1_xboole_0|(X60=k1_xboole_0|(X61=k1_xboole_0|(~m1_subset_1(X62,k3_zfmisc_1(X59,X60,X61))|X62=k3_mcart_1(k5_mcart_1(X59,X60,X61,X62),k6_mcart_1(X59,X60,X61,X62),k7_mcart_1(X59,X60,X61,X62)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t48_mcart_1])])])])).
fof(c_0_21, plain, ![X56, X57, X58]:k3_mcart_1(X56,X57,X58)=k4_tarski(k4_tarski(X56,X57),X58), inference(variable_rename,[status(thm)],[d3_mcart_1])).
cnf(c_0_22, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, plain, ![X23, X24, X25, X26]:(~m1_subset_1(X26,k3_zfmisc_1(X23,X24,X25))|m1_subset_1(k7_mcart_1(X23,X24,X25,X26),X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_mcart_1])])).
cnf(c_0_26, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k3_mcart_1(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_29, negated_conjecture, (m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_30, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_33, plain, (m1_subset_1(k7_mcart_1(X2,X3,X4,X1),X4)|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_34, plain, ![X38, X39]:(((~m1_subset_1(X39,X38)|r2_hidden(X39,X38)|v1_xboole_0(X38))&(~r2_hidden(X39,X38)|m1_subset_1(X39,X38)|v1_xboole_0(X38)))&((~m1_subset_1(X39,X38)|v1_xboole_0(X39)|~v1_xboole_0(X38))&(~v1_xboole_0(X39)|m1_subset_1(X39,X38)|~v1_xboole_0(X38)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_35, negated_conjecture, (esk4_0=X2|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|esk5_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4=k4_tarski(k4_tarski(k5_mcart_1(X1,X2,X3,X4),k6_mcart_1(X1,X2,X3,X4)),k7_mcart_1(X1,X2,X3,X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_23]), c_0_27])).
cnf(c_0_37, negated_conjecture, (k7_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)=k2_mcart_1(esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_38, plain, (m1_subset_1(k7_mcart_1(X2,X3,X4,X1),X4)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[c_0_33, c_0_23])).
fof(c_0_39, plain, ![X63, X64, X65]:((~v1_xboole_0(X63)|~r2_hidden(X64,X63))&(r2_hidden(esk7_1(X65),X65)|v1_xboole_0(X65))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_40, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, negated_conjecture, (esk4_0=X2|esk5_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_35, c_0_27])).
cnf(c_0_43, negated_conjecture, (k4_tarski(k4_tarski(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)),k2_mcart_1(esk5_0))=esk5_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_29]), c_0_37]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(k2_mcart_1(esk5_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_37]), c_0_29])])).
cnf(c_0_45, negated_conjecture, (esk4_0!=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_46, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_48, plain, ![X50, X51, X52]:((r2_hidden(k1_mcart_1(X50),X51)|~r2_hidden(X50,k2_zfmisc_1(X51,X52)))&(r2_hidden(k2_mcart_1(X50),X52)|~r2_hidden(X50,k2_zfmisc_1(X51,X52)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_49, negated_conjecture, (v1_xboole_0(esk5_0)|~v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_40, c_0_29])).
cnf(c_0_50, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_41, c_0_29])).
cnf(c_0_51, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_52, plain, (k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_53, plain, ![X67]:(X67=k1_xboole_0|r2_hidden(esk8_1(X67),X67)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_54, negated_conjecture, (~m1_subset_1(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk2_0)|~m1_subset_1(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])]), c_0_45])).
cnf(c_0_55, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_56, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, negated_conjecture, (v1_xboole_0(esk5_0)|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_51, c_0_23])).
cnf(c_0_59, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_52, c_0_23])).
cnf(c_0_60, plain, (X1=k1_xboole_0|r2_hidden(esk8_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_61, negated_conjecture, (~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk1_0)|~m1_subset_1(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk2_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_62, negated_conjecture, (v1_xboole_0(esk5_0)|r2_hidden(k1_mcart_1(esk5_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63, negated_conjecture, (k1_mcart_1(k1_mcart_1(esk5_0))=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_29]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_64, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_65, negated_conjecture, (k2_mcart_1(k1_mcart_1(esk5_0))=k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_29]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_66, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_47, c_0_60])).
fof(c_0_67, plain, ![X17, X18, X19, X20]:k4_zfmisc_1(X17,X18,X19,X20)=k2_zfmisc_1(k3_zfmisc_1(X17,X18,X19),X20), inference(variable_rename,[status(thm)],[d4_zfmisc_1])).
cnf(c_0_68, negated_conjecture, (~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk1_0)|~r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_55])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(esk5_0)|r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_62]), c_0_63])).
cnf(c_0_70, negated_conjecture, (v1_xboole_0(esk5_0)|r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_62]), c_0_65])).
cnf(c_0_71, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0|r2_hidden(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_66, c_0_50])).
fof(c_0_72, plain, ![X40, X41, X42, X43]:((X40=k1_xboole_0|(X41=k1_xboole_0|(X42=k1_xboole_0|X43=k1_xboole_0))|k4_zfmisc_1(X40,X41,X42,X43)!=k1_xboole_0)&((X40!=k1_xboole_0|k4_zfmisc_1(X40,X41,X42,X43)=k1_xboole_0)&((X41!=k1_xboole_0|k4_zfmisc_1(X40,X41,X42,X43)=k1_xboole_0)&((X42!=k1_xboole_0|k4_zfmisc_1(X40,X41,X42,X43)=k1_xboole_0)&(X43!=k1_xboole_0|k4_zfmisc_1(X40,X41,X42,X43)=k1_xboole_0))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_mcart_1])])])])).
cnf(c_0_73, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k3_zfmisc_1(X1,X2,X3),X4)), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_74, plain, ![X29, X30, X31, X32, X33, X34, X35, X36]:(((~r2_hidden(X32,X31)|(X32=X29|X32=X30)|X31!=k2_tarski(X29,X30))&((X33!=X29|r2_hidden(X33,X31)|X31!=k2_tarski(X29,X30))&(X33!=X30|r2_hidden(X33,X31)|X31!=k2_tarski(X29,X30))))&(((esk6_3(X34,X35,X36)!=X34|~r2_hidden(esk6_3(X34,X35,X36),X36)|X36=k2_tarski(X34,X35))&(esk6_3(X34,X35,X36)!=X35|~r2_hidden(esk6_3(X34,X35,X36),X36)|X36=k2_tarski(X34,X35)))&(r2_hidden(esk6_3(X34,X35,X36),X36)|(esk6_3(X34,X35,X36)=X34|esk6_3(X34,X35,X36)=X35)|X36=k2_tarski(X34,X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_75, plain, ![X48, X49]:k1_enumset1(X48,X48,X49)=k2_tarski(X48,X49), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_76, negated_conjecture, (v1_xboole_0(esk5_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70])).
cnf(c_0_77, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0|r2_hidden(k1_mcart_1(esk5_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_56, c_0_71])).
cnf(c_0_78, plain, (k4_zfmisc_1(X2,X3,X1,X4)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_79, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(rw,[status(thm)],[c_0_73, c_0_23])).
fof(c_0_80, plain, ![X27, X28]:((k2_zfmisc_1(X27,X28)!=k1_xboole_0|(X27=k1_xboole_0|X28=k1_xboole_0))&((X27!=k1_xboole_0|k2_zfmisc_1(X27,X28)=k1_xboole_0)&(X28!=k1_xboole_0|k2_zfmisc_1(X27,X28)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_81, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_82, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_83, negated_conjecture, (esk5_0=k1_xboole_0), inference(spm,[status(thm)],[c_0_66, c_0_76])).
cnf(c_0_84, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0|r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk1_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_77]), c_0_63])).
cnf(c_0_85, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0|r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0),esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_77]), c_0_65])).
cnf(c_0_86, plain, (k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X2,X3),X1),X4)=k1_xboole_0|X1!=k1_xboole_0), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_87, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_80])).
fof(c_0_88, plain, ![X21, X22]:~r2_hidden(X21,k2_zfmisc_1(X21,X22)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_89, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X4,X4,X2)), inference(rw,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_90, plain, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|X4=k1_xboole_0|k4_zfmisc_1(X1,X2,X3,X4)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_91, negated_conjecture, (~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,k1_xboole_0),esk1_0)|~r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,k1_xboole_0),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_83]), c_0_83])).
cnf(c_0_92, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0|r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,k1_xboole_0),esk1_0)), inference(rw,[status(thm)],[c_0_84, c_0_83])).
cnf(c_0_93, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0|r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,k1_xboole_0),esk2_0)), inference(rw,[status(thm)],[c_0_85, c_0_83])).
cnf(c_0_94, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_95, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_96, plain, (r2_hidden(X1,X2)|X2!=k1_enumset1(X3,X3,X1)), inference(er,[status(thm)],[c_0_89])).
cnf(c_0_97, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)!=k1_xboole_0), inference(rw,[status(thm)],[c_0_90, c_0_79])).
cnf(c_0_98, negated_conjecture, (k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93])).
cnf(c_0_99, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_94])).
cnf(c_0_100, plain, (X1!=k1_xboole_0|~r2_hidden(X2,k1_xboole_0)), inference(spm,[status(thm)],[c_0_95, c_0_87])).
cnf(c_0_101, plain, (r2_hidden(X1,k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[c_0_96])).
cnf(c_0_102, negated_conjecture, (X1=k1_xboole_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])]), c_0_30]), c_0_31]), c_0_32])).
cnf(c_0_103, plain, (~r2_hidden(X1,k1_xboole_0)), inference(er,[status(thm)],[c_0_100])).
cnf(c_0_104, plain, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102]), c_0_103]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 105
# Proof object clause steps            : 72
# Proof object formula steps           : 33
# Proof object conjectures             : 36
# Proof object clause conjectures      : 33
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 49
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 36
# Processed clauses                    : 350
# ...of these trivial                  : 16
# ...subsumed                          : 133
# ...remaining for further processing  : 201
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 22
# Backward-rewritten                   : 131
# Generated clauses                    : 958
# ...of the previous two non-trivial   : 755
# Contextual simplify-reflections      : 10
# Paramodulations                      : 927
# Factorizations                       : 8
# Equation resolutions                 : 22
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 11
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 7
# Current number of unprocessed clauses: 332
# ...number of literals in the above   : 1466
# Current number of archived formulas  : 0
# Current number of archived clauses   : 192
# Clause-clause subsumption calls (NU) : 1587
# Rec. Clause-clause subsumption calls : 1484
# Non-unit clause-clause subsumptions  : 107
# Unit Clause-clause subsumption calls : 134
# Rewrite failures with RHS unbound    : 3
# BW rewrite match attempts            : 64
# BW rewrite match successes           : 62
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12889

# -------------------------------------------------
# User time                : 0.019 s
# System time              : 0.000 s
# Total time               : 0.019 s
# Maximum resident set size: 28624 pages
