# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t2_boole)).
fof(d2_xboole_0, axiom, k1_xboole_0=o_0_0_xboole_0, file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', d2_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t28_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t36_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', d5_xboole_0)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t4_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t3_xboole_0)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', idempotence_k3_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', l32_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t41_xboole_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t39_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t17_xboole_1)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t16_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', commutativity_k2_xboole_0)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t12_xboole_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t7_xboole_0)).
fof(t32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k4_xboole_0(X2,X1)=>X1=X2), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t32_xboole_1)).
fof(t47_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/d2da7a48f45e400fa6dff876eefbb3b9.E_input', t47_xboole_1)).
fof(c_0_19, plain, ![X5]:k3_xboole_0(X5,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_20, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_21, plain, (k1_xboole_0=o_0_0_xboole_0), inference(split_conjunct,[status(thm)],[d2_xboole_0])).
fof(c_0_22, plain, ![X6, X7]:k3_xboole_0(X6,X7)=k3_xboole_0(X7,X6), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_23, plain, ![X8, X9]:(~r1_tarski(X8,X9)|k3_xboole_0(X8,X9)=X8), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_24, plain, ![X10, X11]:r1_tarski(k4_xboole_0(X10,X11),X10), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_25, plain, ![X12, X13, X14, X15, X16, X17, X18, X19]:((((r2_hidden(X15,X12)|~r2_hidden(X15,X14)|X14!=k4_xboole_0(X12,X13))&(~r2_hidden(X15,X13)|~r2_hidden(X15,X14)|X14!=k4_xboole_0(X12,X13)))&(~r2_hidden(X16,X12)|r2_hidden(X16,X13)|r2_hidden(X16,X14)|X14!=k4_xboole_0(X12,X13)))&((~r2_hidden(esk1_3(X17,X18,X19),X19)|(~r2_hidden(esk1_3(X17,X18,X19),X17)|r2_hidden(esk1_3(X17,X18,X19),X18))|X19=k4_xboole_0(X17,X18))&((r2_hidden(esk1_3(X17,X18,X19),X17)|r2_hidden(esk1_3(X17,X18,X19),X19)|X19=k4_xboole_0(X17,X18))&(~r2_hidden(esk1_3(X17,X18,X19),X18)|r2_hidden(esk1_3(X17,X18,X19),X19)|X19=k4_xboole_0(X17,X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_26, plain, (k3_xboole_0(X1,o_0_0_xboole_0)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_21])).
cnf(c_0_27, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_30, plain, ![X21, X22, X24, X25, X26]:((r1_xboole_0(X21,X22)|r2_hidden(esk2_2(X21,X22),k3_xboole_0(X21,X22)))&(~r2_hidden(X26,k3_xboole_0(X24,X25))|~r1_xboole_0(X24,X25))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_31, plain, ![X28, X29, X31, X32, X33]:(((r2_hidden(esk3_2(X28,X29),X28)|r1_xboole_0(X28,X29))&(r2_hidden(esk3_2(X28,X29),X29)|r1_xboole_0(X28,X29)))&(~r2_hidden(X33,X31)|~r2_hidden(X33,X32)|~r1_xboole_0(X31,X32))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_32, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k3_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_34, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_27])).
fof(c_0_35, plain, ![X27]:k3_xboole_0(X27,X27)=X27, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_36, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (r2_hidden(esk3_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (k4_xboole_0(o_0_0_xboole_0,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_40, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk2_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (~r1_xboole_0(o_0_0_xboole_0,X1)|~r2_hidden(X2,o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_36, c_0_33])).
cnf(c_0_43, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_44, plain, (~r2_hidden(X1,o_0_0_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_45, plain, (r1_xboole_0(X1,X1)|r2_hidden(esk2_2(X1,X1),X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_48, plain, (r1_xboole_0(o_0_0_xboole_0,o_0_0_xboole_0)|~r2_hidden(esk2_2(o_0_0_xboole_0,o_0_0_xboole_0),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_49, plain, (r1_xboole_0(X1,o_0_0_xboole_0)|r2_hidden(esk2_2(X1,o_0_0_xboole_0),o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_40, c_0_26])).
cnf(c_0_50, plain, (r1_xboole_0(o_0_0_xboole_0,X1)|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_51, plain, (r1_xboole_0(o_0_0_xboole_0,o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_52, plain, (~r1_xboole_0(X1,k4_xboole_0(X1,X2))|~r2_hidden(X3,k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_36, c_0_34])).
cnf(c_0_53, plain, (r1_xboole_0(o_0_0_xboole_0,X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_54, plain, (r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk1_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_55, plain, (r2_hidden(esk1_3(X1,X2,X3),X1)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_56, plain, ![X34, X35]:((k4_xboole_0(X34,X35)!=k1_xboole_0|r1_tarski(X34,X35))&(~r1_tarski(X34,X35)|k4_xboole_0(X34,X35)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_57, plain, ![X36, X37, X38]:k4_xboole_0(k4_xboole_0(X36,X37),X38)=k4_xboole_0(X36,k2_xboole_0(X37,X38)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_58, plain, (~r2_hidden(X1,o_0_0_xboole_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_39])).
cnf(c_0_59, plain, (X1=k4_xboole_0(X2,X2)|r2_hidden(esk1_3(X2,X2,X1),X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
fof(c_0_60, plain, ![X39, X40]:k2_xboole_0(X39,k4_xboole_0(X40,X39))=k2_xboole_0(X39,X40), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_61, plain, ![X41, X42]:r1_tarski(k3_xboole_0(X41,X42),X41), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_62, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_64, plain, (k4_xboole_0(X1,X1)=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_66, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_67, plain, ![X47, X48, X49]:k3_xboole_0(k3_xboole_0(X47,X48),X49)=k3_xboole_0(X47,k3_xboole_0(X48,X49)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_68, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=o_0_0_xboole_0), inference(rw,[status(thm)],[c_0_62, c_0_21])).
cnf(c_0_69, plain, (k4_xboole_0(X1,k2_xboole_0(X2,X1))=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
fof(c_0_70, plain, ![X43, X44]:k2_xboole_0(X43,X44)=k2_xboole_0(X44,X43), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_71, plain, ![X45, X46]:(~r1_tarski(X45,X46)|k2_xboole_0(X45,X46)=X46), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_72, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_66, c_0_27])).
cnf(c_0_73, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_74, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_29, c_0_63])).
cnf(c_0_76, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_77, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_78, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_79, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_28, c_0_74])).
cnf(c_0_80, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,X3)),k4_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_81, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_66]), c_0_76])).
cnf(c_0_82, plain, (r1_tarski(k3_xboole_0(X1,X2),k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_83, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_84, plain, (r1_tarski(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_82, c_0_79])).
cnf(c_0_85, plain, (k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_29]), c_0_76])).
cnf(c_0_86, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X2,X3)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_83])).
cnf(c_0_87, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_88, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|k4_xboole_0(X1,k2_xboole_0(X2,X3))!=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_63])).
cnf(c_0_89, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X1)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_64]), c_0_26])).
cnf(c_0_90, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),X1)), inference(spm,[status(thm)],[c_0_87, c_0_81])).
cnf(c_0_91, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_73, c_0_41])).
cnf(c_0_92, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X1),X2)), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_93, plain, (k3_xboole_0(X1,k4_xboole_0(k3_xboole_0(X1,X2),X3))=k4_xboole_0(k3_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_90]), c_0_27])).
cnf(c_0_94, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_91, c_0_27])).
cnf(c_0_95, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),X2),X1)), inference(spm,[status(thm)],[c_0_92, c_0_76])).
cnf(c_0_96, plain, (k3_xboole_0(X1,k4_xboole_0(k3_xboole_0(X2,X1),X3))=k4_xboole_0(k3_xboole_0(X2,X1),X3)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_97, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_66, c_0_73])).
cnf(c_0_98, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(k2_xboole_0(X1,X2),X3),k4_xboole_0(X2,X1)),X1)), inference(spm,[status(thm)],[c_0_95, c_0_65])).
cnf(c_0_99, plain, (k3_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X3))=k4_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_96, c_0_79])).
cnf(c_0_100, plain, (r1_tarski(o_0_0_xboole_0,X1)), inference(spm,[status(thm)],[c_0_66, c_0_26])).
fof(c_0_101, plain, ![X50]:(X50=k1_xboole_0|r2_hidden(esk4_1(X50),X50)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_102, plain, (r1_tarski(k3_xboole_0(X1,k3_xboole_0(X2,X3)),k3_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_97, c_0_94])).
cnf(c_0_103, plain, (k3_xboole_0(X1,k3_xboole_0(k4_xboole_0(X1,X2),X3))=k3_xboole_0(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_73, c_0_34])).
cnf(c_0_104, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,X3))),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_63])).
cnf(c_0_105, plain, (k2_xboole_0(o_0_0_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_77, c_0_100])).
cnf(c_0_106, plain, (X1=k1_xboole_0|r2_hidden(esk4_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_107, plain, (k2_xboole_0(X1,k3_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_72]), c_0_76])).
cnf(c_0_108, plain, (r1_tarski(k3_xboole_0(k4_xboole_0(X1,X2),X3),k3_xboole_0(X1,X3))), inference(spm,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_109, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_110, plain, (X1=o_0_0_xboole_0|r2_hidden(esk4_1(X1),X1)), inference(rw,[status(thm)],[c_0_106, c_0_21])).
cnf(c_0_111, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk3_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_38, c_0_37])).
cnf(c_0_112, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))), inference(spm,[status(thm)],[c_0_80, c_0_107])).
cnf(c_0_113, plain, (r1_tarski(k4_xboole_0(X1,X2),k3_xboole_0(X1,k2_xboole_0(X3,k4_xboole_0(X1,X2))))), inference(spm,[status(thm)],[c_0_108, c_0_79])).
cnf(c_0_114, plain, (k2_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_109]), c_0_76])).
cnf(c_0_115, plain, (k3_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_36, c_0_110])).
cnf(c_0_116, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_111, c_0_47])).
cnf(c_0_117, plain, (k3_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X1,k3_xboole_0(X3,X2)))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_28, c_0_112])).
fof(c_0_118, plain, ![X52, X53]:(k4_xboole_0(X52,X53)!=k4_xboole_0(X53,X52)|X52=X53), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_xboole_1])])).
cnf(c_0_119, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_120, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k3_xboole_0(X1,X2))), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_121, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_122, plain, (k4_xboole_0(k3_xboole_0(X1,X2),X2)=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_64]), c_0_26])).
cnf(c_0_123, plain, (X1=X2|k4_xboole_0(X1,X2)!=k4_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_124, plain, (k4_xboole_0(X1,X2)=o_0_0_xboole_0|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_119, c_0_21])).
cnf(c_0_125, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1))),o_0_0_xboole_0)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_126, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),X3)), inference(spm,[status(thm)],[c_0_88, c_0_122])).
cnf(c_0_127, plain, (X1=X2|k4_xboole_0(X1,X2)!=o_0_0_xboole_0|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_128, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X2,X1)))=o_0_0_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_125]), c_0_26])).
cnf(c_0_129, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,k2_xboole_0(X2,X3)),X2),k4_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_126, c_0_65])).
cnf(c_0_130, plain, (k3_xboole_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))=X1), inference(spm,[status(thm)],[c_0_28, c_0_84])).
cnf(c_0_131, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_29])])).
cnf(c_0_132, plain, (r1_tarski(k4_xboole_0(X1,X2),k4_xboole_0(k2_xboole_0(X3,X1),X2))), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_133, plain, (k4_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,X1)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_131]), c_0_34])).
cnf(c_0_134, plain, (~r1_xboole_0(k3_xboole_0(X1,X2),X3)|~r2_hidden(X4,k3_xboole_0(X1,k3_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_36, c_0_73])).
cnf(c_0_135, plain, (r1_tarski(k4_xboole_0(k3_xboole_0(X1,X2),X3),k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_132, c_0_107])).
cnf(c_0_136, plain, (k4_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X3,X1))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_133, c_0_86])).
cnf(c_0_137, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_134, c_0_40])).
cnf(c_0_138, plain, (r1_tarski(k3_xboole_0(X1,X2),k4_xboole_0(X2,k4_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_139, plain, (r1_xboole_0(X1,k3_xboole_0(X2,k4_xboole_0(X3,k3_xboole_0(X1,X2))))), inference(spm,[status(thm)],[c_0_137, c_0_116])).
cnf(c_0_140, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X3,X1)))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_138]), c_0_73]), c_0_34])).
cnf(c_0_141, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_139, c_0_34])).
fof(c_0_142, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t47_xboole_1])).
cnf(c_0_143, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_109]), c_0_27]), c_0_140])).
cnf(c_0_144, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X1,X2)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_115, c_0_141])).
fof(c_0_145, negated_conjecture, k4_xboole_0(esk5_0,k3_xboole_0(esk5_0,esk6_0))!=k4_xboole_0(esk5_0,esk6_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_142])])])).
cnf(c_0_146, plain, (k4_xboole_0(X1,k3_xboole_0(X2,X1))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_143]), c_0_27]), c_0_34])).
cnf(c_0_147, plain, (k3_xboole_0(X1,k4_xboole_0(X2,k3_xboole_0(X2,X1)))=o_0_0_xboole_0), inference(spm,[status(thm)],[c_0_144, c_0_27])).
cnf(c_0_148, plain, (k2_xboole_0(X1,o_0_0_xboole_0)=X1), inference(spm,[status(thm)],[c_0_76, c_0_105])).
cnf(c_0_149, negated_conjecture, (k4_xboole_0(esk5_0,k3_xboole_0(esk5_0,esk6_0))!=k4_xboole_0(esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_145])).
cnf(c_0_150, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_63]), c_0_148]), c_0_63]), c_0_76]), c_0_107])).
cnf(c_0_151, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_149, c_0_150])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 152
# Proof object clause steps            : 114
# Proof object formula steps           : 38
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 84
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 16773
# ...of these trivial                  : 4729
# ...subsumed                          : 10158
# ...remaining for further processing  : 1886
# Other redundant clauses eliminated   : 1409
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 111
# Generated clauses                    : 635442
# ...of the previous two non-trivial   : 303532
# Contextual simplify-reflections      : 2
# Paramodulations                      : 634026
# Factorizations                       : 6
# Equation resolutions                 : 1410
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1760
#    Positive orientable unit clauses  : 1325
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 426
# Current number of unprocessed clauses: 286338
# ...number of literals in the above   : 481817
# Current number of archived formulas  : 0
# Current number of archived clauses   : 123
# Clause-clause subsumption calls (NU) : 54695
# Rec. Clause-clause subsumption calls : 50807
# Non-unit clause-clause subsumptions  : 3998
# Unit Clause-clause subsumption calls : 1378
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18668
# BW rewrite match successes           : 188
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5325231

# -------------------------------------------------
# User time                : 3.298 s
# System time              : 0.080 s
# Total time               : 3.378 s
# Maximum resident set size: 223000 pages
