# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t52_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(~(r2_hidden(X4,X2))<=>r2_hidden(X4,X3)))=>X2=k3_subset_1(X1,X3)))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', t52_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', d1_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', fc1_subset_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', t100_xboole_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', redefinition_k6_subset_1)).
fof(dt_k6_subset_1, axiom, ![X1, X2]:m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1)), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', dt_k6_subset_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', t28_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', commutativity_k3_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', d1_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', d3_tarski)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', t48_xboole_1)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', t47_xboole_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', d5_subset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/d4f3295f95454757ad80ab7524e49d15.E_input', d10_xboole_0)).
fof(c_0_17, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(~(r2_hidden(X4,X2))<=>r2_hidden(X4,X3)))=>X2=k3_subset_1(X1,X3))))), inference(assume_negation,[status(cth)],[t52_subset_1])).
fof(c_0_18, plain, ![X9, X10, X11, X12, X13, X14]:(((~r2_hidden(X11,X10)|r1_tarski(X11,X9)|X10!=k1_zfmisc_1(X9))&(~r1_tarski(X12,X9)|r2_hidden(X12,X10)|X10!=k1_zfmisc_1(X9)))&((~r2_hidden(esk4_2(X13,X14),X14)|~r1_tarski(esk4_2(X13,X14),X13)|X14=k1_zfmisc_1(X13))&(r2_hidden(esk4_2(X13,X14),X14)|r1_tarski(esk4_2(X13,X14),X13)|X14=k1_zfmisc_1(X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_19, plain, ![X16, X17]:(((~m1_subset_1(X17,X16)|r2_hidden(X17,X16)|v1_xboole_0(X16))&(~r2_hidden(X17,X16)|m1_subset_1(X17,X16)|v1_xboole_0(X16)))&((~m1_subset_1(X17,X16)|v1_xboole_0(X17)|~v1_xboole_0(X16))&(~v1_xboole_0(X17)|m1_subset_1(X17,X16)|~v1_xboole_0(X16)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_20, negated_conjecture, ![X8]:(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&(((r2_hidden(X8,esk2_0)|r2_hidden(X8,esk3_0)|~m1_subset_1(X8,esk1_0))&(~r2_hidden(X8,esk3_0)|~r2_hidden(X8,esk2_0)|~m1_subset_1(X8,esk1_0)))&esk2_0!=k3_subset_1(esk1_0,esk3_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
fof(c_0_21, plain, ![X18]:~v1_xboole_0(k1_zfmisc_1(X18)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_22, plain, ![X23, X24]:k4_xboole_0(X23,X24)=k5_xboole_0(X23,k3_xboole_0(X23,X24)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_23, plain, ![X21, X22]:k6_subset_1(X21,X22)=k4_xboole_0(X21,X22), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_24, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_28, plain, ![X19, X20]:m1_subset_1(k6_subset_1(X19,X20),k1_zfmisc_1(X19)), inference(variable_rename,[status(thm)],[dt_k6_subset_1])).
cnf(c_0_29, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_30, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_31, plain, ![X25, X26]:(~r1_tarski(X25,X26)|k3_xboole_0(X25,X26)=X25), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_32, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
fof(c_0_34, plain, ![X27, X28]:k3_xboole_0(X27,X28)=k3_xboole_0(X28,X27), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_35, plain, (m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (k6_subset_1(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_39, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_40, plain, ![X29, X30, X31, X32, X33, X34, X35, X36]:((((r2_hidden(X32,X29)|~r2_hidden(X32,X31)|X31!=k4_xboole_0(X29,X30))&(~r2_hidden(X32,X30)|~r2_hidden(X32,X31)|X31!=k4_xboole_0(X29,X30)))&(~r2_hidden(X33,X29)|r2_hidden(X33,X30)|r2_hidden(X33,X31)|X31!=k4_xboole_0(X29,X30)))&((~r2_hidden(esk5_3(X34,X35,X36),X36)|(~r2_hidden(esk5_3(X34,X35,X36),X34)|r2_hidden(esk5_3(X34,X35,X36),X35))|X36=k4_xboole_0(X34,X35))&((r2_hidden(esk5_3(X34,X35,X36),X34)|r2_hidden(esk5_3(X34,X35,X36),X36)|X36=k4_xboole_0(X34,X35))&(~r2_hidden(esk5_3(X34,X35,X36),X35)|r2_hidden(esk5_3(X34,X35,X36),X36)|X36=k4_xboole_0(X34,X35))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_41, plain, (m1_subset_1(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_42, negated_conjecture, (k3_xboole_0(esk1_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_44, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_45, plain, ![X38, X39, X40]:((~v1_xboole_0(X38)|~r2_hidden(X39,X38))&(r2_hidden(esk6_1(X40),X40)|v1_xboole_0(X40))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(k5_xboole_0(esk1_0,esk2_0),k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_47, plain, ![X42, X43, X44, X45, X46, X47, X48, X49]:((((r2_hidden(X45,X42)|~r2_hidden(X45,X44)|X44!=k3_xboole_0(X42,X43))&(r2_hidden(X45,X43)|~r2_hidden(X45,X44)|X44!=k3_xboole_0(X42,X43)))&(~r2_hidden(X46,X42)|~r2_hidden(X46,X43)|r2_hidden(X46,X44)|X44!=k3_xboole_0(X42,X43)))&((~r2_hidden(esk7_3(X47,X48,X49),X49)|(~r2_hidden(esk7_3(X47,X48,X49),X47)|~r2_hidden(esk7_3(X47,X48,X49),X48))|X49=k3_xboole_0(X47,X48))&((r2_hidden(esk7_3(X47,X48,X49),X47)|r2_hidden(esk7_3(X47,X48,X49),X49)|X49=k3_xboole_0(X47,X48))&(r2_hidden(esk7_3(X47,X48,X49),X48)|r2_hidden(esk7_3(X47,X48,X49),X49)|X49=k3_xboole_0(X47,X48))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_43]), c_0_27])).
cnf(c_0_49, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_30]), c_0_36])).
cnf(c_0_50, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_51, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_52, plain, ![X51, X52, X53, X54, X55]:((~r1_tarski(X51,X52)|(~r2_hidden(X53,X51)|r2_hidden(X53,X52)))&((r2_hidden(esk8_2(X54,X55),X54)|r1_tarski(X54,X55))&(~r2_hidden(esk8_2(X54,X55),X55)|r1_tarski(X54,X55)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_53, negated_conjecture, (r2_hidden(k5_xboole_0(esk1_0,esk2_0),k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_46]), c_0_27])).
cnf(c_0_54, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_55, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_48])).
fof(c_0_57, plain, ![X57, X58]:k4_xboole_0(X57,k4_xboole_0(X57,X58))=k3_xboole_0(X57,X58), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_58, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_49])).
cnf(c_0_59, negated_conjecture, (r2_hidden(X1,esk2_0)|r2_hidden(X1,esk3_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_60, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_53])).
cnf(c_0_63, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X4!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_30]), c_0_36])).
cnf(c_0_64, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_65, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_56]), c_0_39])).
fof(c_0_66, plain, ![X59, X60]:k4_xboole_0(X59,k3_xboole_0(X59,X60))=k4_xboole_0(X59,X60), inference(variable_rename,[status(thm)],[t47_xboole_1])).
fof(c_0_67, plain, ![X61, X62]:(~m1_subset_1(X62,k1_zfmisc_1(X61))|k3_subset_1(X61,X62)=k4_xboole_0(X61,X62)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_68, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_69, negated_conjecture, (~r2_hidden(X1,k5_xboole_0(esk1_0,esk2_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_58, c_0_42])).
cnf(c_0_70, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_71, plain, (r1_tarski(X1,X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_72, negated_conjecture, (r2_hidden(X1,esk3_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_73, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,k5_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_74, negated_conjecture, (~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_75, plain, (r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_76, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_77, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_78, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_79, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_30]), c_0_30]), c_0_36]), c_0_36])).
fof(c_0_80, plain, ![X63, X64]:(((r1_tarski(X63,X64)|X63!=X64)&(r1_tarski(X64,X63)|X63!=X64))&(~r1_tarski(X63,X64)|~r1_tarski(X64,X63)|X63=X64)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_81, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),X1)|~r2_hidden(esk8_2(k5_xboole_0(esk1_0,esk2_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_82, negated_conjecture, (r1_tarski(X1,esk3_0)|r2_hidden(esk8_2(X1,esk3_0),esk2_0)|~r2_hidden(esk8_2(X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),X1)|r2_hidden(esk8_2(k5_xboole_0(esk1_0,esk2_0),X1),esk1_0)), inference(spm,[status(thm)],[c_0_73, c_0_70])).
cnf(c_0_84, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_60])).
cnf(c_0_85, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_61, c_0_38])).
cnf(c_0_86, plain, (r1_tarski(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|r2_hidden(esk8_2(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3))),X3)|~r2_hidden(esk8_2(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3))),X2)), inference(spm,[status(thm)],[c_0_71, c_0_75])).
cnf(c_0_87, negated_conjecture, (r1_tarski(esk3_0,X1)|r2_hidden(esk8_2(esk3_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_70])).
cnf(c_0_88, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_30]), c_0_30]), c_0_36]), c_0_36])).
cnf(c_0_89, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_30]), c_0_36])).
cnf(c_0_90, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_41, c_0_79])).
cnf(c_0_91, negated_conjecture, (k3_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk2_0))=k5_xboole_0(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_62]), c_0_39])).
cnf(c_0_92, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_93, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])).
cnf(c_0_94, negated_conjecture, (r1_tarski(esk3_0,X1)|~r2_hidden(esk8_2(esk3_0,X1),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_70]), c_0_85])).
cnf(c_0_95, negated_conjecture, (r1_tarski(esk3_0,k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,X1)))|r2_hidden(esk8_2(esk3_0,k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,X1))),X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_96, plain, (k3_subset_1(X1,k3_xboole_0(X1,X2))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90])])).
cnf(c_0_97, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_42]), c_0_91])).
cnf(c_0_98, negated_conjecture, (k5_xboole_0(esk1_0,esk2_0)=esk3_0|~r1_tarski(esk3_0,k5_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_99, negated_conjecture, (r1_tarski(esk3_0,k5_xboole_0(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_42])).
cnf(c_0_100, negated_conjecture, (esk2_0!=k3_subset_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_101, negated_conjecture, (k3_subset_1(esk1_0,k5_xboole_0(esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_91]), c_0_97])).
cnf(c_0_102, negated_conjecture, (k5_xboole_0(esk1_0,esk2_0)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_99])])).
cnf(c_0_103, negated_conjecture, (k3_subset_1(esk1_0,esk3_0)=k5_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_96, c_0_65])).
cnf(c_0_104, negated_conjecture, (k5_xboole_0(esk1_0,esk3_0)!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_89]), c_0_65]), c_0_43])])).
cnf(c_0_105, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102]), c_0_103]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 71
# Proof object formula steps           : 35
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 32
# Proof object simplifying inferences  : 42
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 40
# Processed clauses                    : 18416
# ...of these trivial                  : 264
# ...subsumed                          : 15808
# ...remaining for further processing  : 2344
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 148
# Backward-rewritten                   : 389
# Generated clauses                    : 251655
# ...of the previous two non-trivial   : 233617
# Contextual simplify-reflections      : 60
# Paramodulations                      : 251483
# Factorizations                       : 160
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1756
#    Positive orientable unit clauses  : 201
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 23
#    Non-unit-clauses                  : 1531
# Current number of unprocessed clauses: 214201
# ...number of literals in the above   : 847659
# Current number of archived formulas  : 0
# Current number of archived clauses   : 580
# Clause-clause subsumption calls (NU) : 585710
# Rec. Clause-clause subsumption calls : 330968
# Non-unit clause-clause subsumptions  : 10109
# Unit Clause-clause subsumption calls : 38891
# Rewrite failures with RHS unbound    : 56
# BW rewrite match attempts            : 402
# BW rewrite match successes           : 80
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5162287

# -------------------------------------------------
# User time                : 3.220 s
# System time              : 0.057 s
# Total time               : 3.276 s
# Maximum resident set size: 176768 pages
