# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_AE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t50_funct_1, conjecture, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&(k1_relat_1(X3)=X1&(r1_tarski(k2_relat_1(X3),X1)&(v2_funct_1(X2)&k5_relat_1(X3,X2)=X2))))=>X3=k6_relat_1(X1)))), file('data/MPTP/proofs/d5fe536cc8d741b286ac33e56f08892b.E_input', t50_funct_1)).
fof(t49_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)<=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r1_tarski(k2_relat_1(X2),k1_relat_1(X1))&(r1_tarski(k2_relat_1(X3),k1_relat_1(X1))&(k1_relat_1(X2)=k1_relat_1(X3)&k5_relat_1(X2,X1)=k5_relat_1(X3,X1))))=>X2=X3))))), file('data/MPTP/proofs/d5fe536cc8d741b286ac33e56f08892b.E_input', t49_funct_1)).
fof(t25_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&r1_tarski(k2_relat_1(X1),k2_relat_1(X2)))))), file('data/MPTP/proofs/d5fe536cc8d741b286ac33e56f08892b.E_input', t25_relat_1)).
fof(t71_relat_1, axiom, ![X1]:(k1_relat_1(k6_relat_1(X1))=X1&k2_relat_1(k6_relat_1(X1))=X1), file('data/MPTP/proofs/d5fe536cc8d741b286ac33e56f08892b.E_input', t71_relat_1)).
fof(fc3_funct_1, axiom, ![X1]:(v1_relat_1(k6_relat_1(X1))&v1_funct_1(k6_relat_1(X1))), file('data/MPTP/proofs/d5fe536cc8d741b286ac33e56f08892b.E_input', fc3_funct_1)).
fof(t76_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(r1_tarski(k5_relat_1(X2,k6_relat_1(X1)),X2)&r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2))), file('data/MPTP/proofs/d5fe536cc8d741b286ac33e56f08892b.E_input', t76_relat_1)).
fof(t78_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k5_relat_1(k6_relat_1(k1_relat_1(X1)),X1)=X1), file('data/MPTP/proofs/d5fe536cc8d741b286ac33e56f08892b.E_input', t78_relat_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X2)=X1&(k1_relat_1(X3)=X1&(r1_tarski(k2_relat_1(X3),X1)&(v2_funct_1(X2)&k5_relat_1(X3,X2)=X2))))=>X3=k6_relat_1(X1))))), inference(assume_negation,[status(cth)],[t50_funct_1])).
fof(c_0_8, plain, ![X7, X8, X9]:((~v2_funct_1(X7)|(~v1_relat_1(X8)|~v1_funct_1(X8)|(~v1_relat_1(X9)|~v1_funct_1(X9)|(~r1_tarski(k2_relat_1(X8),k1_relat_1(X7))|(~r1_tarski(k2_relat_1(X9),k1_relat_1(X7))|(k1_relat_1(X8)!=k1_relat_1(X9)|k5_relat_1(X8,X7)!=k5_relat_1(X9,X7)))|X8=X9)))|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(((v1_relat_1(esk4_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(v1_funct_1(esk4_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(((v1_relat_1(esk5_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(v1_funct_1(esk5_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))&(((r1_tarski(k2_relat_1(esk4_1(X7)),k1_relat_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&((r1_tarski(k2_relat_1(esk5_1(X7)),k1_relat_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&((k1_relat_1(esk4_1(X7))=k1_relat_1(esk5_1(X7))|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7)))&(k5_relat_1(esk4_1(X7),X7)=k5_relat_1(esk5_1(X7),X7)|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))))&(esk4_1(X7)!=esk5_1(X7)|v2_funct_1(X7)|(~v1_relat_1(X7)|~v1_funct_1(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_funct_1])])])])])).
fof(c_0_9, negated_conjecture, ((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v1_relat_1(esk3_0)&v1_funct_1(esk3_0))&((k1_relat_1(esk2_0)=esk1_0&(k1_relat_1(esk3_0)=esk1_0&(r1_tarski(k2_relat_1(esk3_0),esk1_0)&(v2_funct_1(esk2_0)&k5_relat_1(esk3_0,esk2_0)=esk2_0))))&esk3_0!=k6_relat_1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])).
cnf(c_0_10, plain, (X2=X3|~v2_funct_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~r1_tarski(k2_relat_1(X2),k1_relat_1(X1))|~r1_tarski(k2_relat_1(X3),k1_relat_1(X1))|k1_relat_1(X2)!=k1_relat_1(X3)|k5_relat_1(X2,X1)!=k5_relat_1(X3,X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, negated_conjecture, (k1_relat_1(esk2_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, negated_conjecture, (v2_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_13, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_15, plain, ![X12, X13]:((r1_tarski(k1_relat_1(X12),k1_relat_1(X13))|~r1_tarski(X12,X13)|~v1_relat_1(X13)|~v1_relat_1(X12))&(r1_tarski(k2_relat_1(X12),k2_relat_1(X13))|~r1_tarski(X12,X13)|~v1_relat_1(X13)|~v1_relat_1(X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_relat_1])])])])).
fof(c_0_16, plain, ![X14]:(k1_relat_1(k6_relat_1(X14))=X14&k2_relat_1(k6_relat_1(X14))=X14), inference(variable_rename,[status(thm)],[t71_relat_1])).
fof(c_0_17, plain, ![X16]:(v1_relat_1(k6_relat_1(X16))&v1_funct_1(k6_relat_1(X16))), inference(variable_rename,[status(thm)],[fc3_funct_1])).
fof(c_0_18, plain, ![X17, X18]:((r1_tarski(k5_relat_1(X18,k6_relat_1(X17)),X18)|~v1_relat_1(X18))&(r1_tarski(k5_relat_1(k6_relat_1(X17),X18),X18)|~v1_relat_1(X18))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t76_relat_1])])])).
fof(c_0_19, plain, ![X15]:(~v1_relat_1(X15)|k5_relat_1(k6_relat_1(k1_relat_1(X15)),X15)=X15), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t78_relat_1])])).
cnf(c_0_20, negated_conjecture, (X1=X2|k5_relat_1(X1,esk2_0)!=k5_relat_1(X2,esk2_0)|k1_relat_1(X1)!=k1_relat_1(X2)|~r1_tarski(k2_relat_1(X2),esk1_0)|~r1_tarski(k2_relat_1(X1),esk1_0)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13]), c_0_14])])).
cnf(c_0_21, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_22, negated_conjecture, (k5_relat_1(esk3_0,esk2_0)=esk2_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_23, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_25, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_26, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_27, plain, (k1_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_28, plain, (v1_relat_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (r1_tarski(k5_relat_1(k6_relat_1(X1),X2),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, plain, (k5_relat_1(k6_relat_1(k1_relat_1(X1)),X1)=X1|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, negated_conjecture, (X1=esk3_0|k5_relat_1(X1,esk2_0)!=esk2_0|k1_relat_1(X1)!=esk1_0|~r1_tarski(k2_relat_1(X1),esk1_0)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_20, c_0_21]), c_0_22]), c_0_23]), c_0_24]), c_0_25])])).
cnf(c_0_32, plain, (k2_relat_1(k6_relat_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_33, plain, (v1_funct_1(k6_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_34, plain, (r1_tarski(X1,k1_relat_1(X2))|~r1_tarski(k6_relat_1(X1),X2)|~v1_relat_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_35, plain, (r1_tarski(X1,X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, negated_conjecture, (k6_relat_1(X1)=esk3_0|k5_relat_1(k6_relat_1(X1),esk2_0)!=esk2_0|X1!=esk1_0|~r1_tarski(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_27]), c_0_33]), c_0_28])])).
cnf(c_0_37, plain, (r1_tarski(X1,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_27]), c_0_28])])).
cnf(c_0_38, negated_conjecture, (esk3_0!=k6_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_39, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_30]), c_0_11]), c_0_11]), c_0_11]), c_0_37]), c_0_14])]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 40
# Proof object clause steps            : 25
# Proof object formula steps           : 15
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 18
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 71
# ...of these trivial                  : 1
# ...subsumed                          : 5
# ...remaining for further processing  : 65
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 4
# Generated clauses                    : 85
# ...of the previous two non-trivial   : 67
# Contextual simplify-reflections      : 4
# Paramodulations                      : 85
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 60
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 42
# Current number of unprocessed clauses: 24
# ...number of literals in the above   : 202
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 632
# Rec. Clause-clause subsumption calls : 179
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 32
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3531

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28680 pages
