# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d7_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>![X5]:(m1_subset_1(X5,X3)=>(X5=k7_mcart_1(X1,X2,X3,X4)<=>![X6, X7, X8]:(X4=k3_mcart_1(X6,X7,X8)=>X5=X8))))))))), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', d7_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', d3_zfmisc_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', d3_mcart_1)).
fof(dt_k7_mcart_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>m1_subset_1(k7_mcart_1(X1,X2,X3,X4),X3)), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', dt_k7_mcart_1)).
fof(t47_mcart_1, conjecture, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&?[X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))&?[X5, X6, X7]:(X4=k3_mcart_1(X5,X6,X7)&~((k5_mcart_1(X1,X2,X3,X4)=X5&(k6_mcart_1(X1,X2,X3,X4)=X6&k7_mcart_1(X1,X2,X3,X4)=X7))))))))), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', t47_mcart_1)).
fof(d6_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>![X5]:(m1_subset_1(X5,X2)=>(X5=k6_mcart_1(X1,X2,X3,X4)<=>![X6, X7, X8]:(X4=k3_mcart_1(X6,X7,X8)=>X5=X7))))))))), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', d6_mcart_1)).
fof(dt_k6_mcart_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>m1_subset_1(k6_mcart_1(X1,X2,X3,X4),X2)), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', dt_k6_mcart_1)).
fof(d5_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>![X5]:(m1_subset_1(X5,X1)=>(X5=k5_mcart_1(X1,X2,X3,X4)<=>![X6, X7, X8]:(X4=k3_mcart_1(X6,X7,X8)=>X5=X6))))))))), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', d5_mcart_1)).
fof(dt_k5_mcart_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>m1_subset_1(k5_mcart_1(X1,X2,X3,X4),X1)), file('data/MPTP/proofs/d61a0cba7ddd49bd83f1feae1e3e12b8.E_input', dt_k5_mcart_1)).
fof(c_0_9, plain, ![X56, X57, X58, X59, X60, X61, X62, X63]:((X60!=k7_mcart_1(X56,X57,X58,X59)|(X59!=k3_mcart_1(X61,X62,X63)|X60=X63)|~m1_subset_1(X60,X58)|~m1_subset_1(X59,k3_zfmisc_1(X56,X57,X58))|X58=k1_xboole_0|X57=k1_xboole_0|X56=k1_xboole_0)&((X59=k3_mcart_1(esk14_5(X56,X57,X58,X59,X60),esk15_5(X56,X57,X58,X59,X60),esk16_5(X56,X57,X58,X59,X60))|X60=k7_mcart_1(X56,X57,X58,X59)|~m1_subset_1(X60,X58)|~m1_subset_1(X59,k3_zfmisc_1(X56,X57,X58))|X58=k1_xboole_0|X57=k1_xboole_0|X56=k1_xboole_0)&(X60!=esk16_5(X56,X57,X58,X59,X60)|X60=k7_mcart_1(X56,X57,X58,X59)|~m1_subset_1(X60,X58)|~m1_subset_1(X59,k3_zfmisc_1(X56,X57,X58))|X58=k1_xboole_0|X57=k1_xboole_0|X56=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_mcart_1])])])])])])])).
fof(c_0_10, plain, ![X42, X43, X44]:k3_zfmisc_1(X42,X43,X44)=k2_zfmisc_1(k2_zfmisc_1(X42,X43),X44), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_11, plain, ![X24, X25, X26]:k3_mcart_1(X24,X25,X26)=k4_tarski(k4_tarski(X24,X25),X26), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_12, plain, ![X38, X39, X40, X41]:(~m1_subset_1(X41,k3_zfmisc_1(X38,X39,X40))|m1_subset_1(k7_mcart_1(X38,X39,X40,X41),X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_mcart_1])])).
fof(c_0_13, negated_conjecture, ~(![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&?[X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))&?[X5, X6, X7]:(X4=k3_mcart_1(X5,X6,X7)&~((k5_mcart_1(X1,X2,X3,X4)=X5&(k6_mcart_1(X1,X2,X3,X4)=X6&k7_mcart_1(X1,X2,X3,X4)=X7)))))))))), inference(assume_negation,[status(cth)],[t47_mcart_1])).
cnf(c_0_14, plain, (X1=X8|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1!=k7_mcart_1(X2,X3,X4,X5)|X5!=k3_mcart_1(X6,X7,X8)|~m1_subset_1(X1,X4)|~m1_subset_1(X5,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (m1_subset_1(k7_mcart_1(X2,X3,X4,X1),X4)|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, negated_conjecture, (esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&(m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&(esk4_0=k3_mcart_1(esk5_0,esk6_0,esk7_0)&(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk5_0|(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk6_0|k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk7_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])).
fof(c_0_19, plain, ![X45, X46, X47, X48, X49, X50, X51, X52]:((X49!=k6_mcart_1(X45,X46,X47,X48)|(X48!=k3_mcart_1(X50,X51,X52)|X49=X51)|~m1_subset_1(X49,X46)|~m1_subset_1(X48,k3_zfmisc_1(X45,X46,X47))|X47=k1_xboole_0|X46=k1_xboole_0|X45=k1_xboole_0)&((X48=k3_mcart_1(esk11_5(X45,X46,X47,X48,X49),esk12_5(X45,X46,X47,X48,X49),esk13_5(X45,X46,X47,X48,X49))|X49=k6_mcart_1(X45,X46,X47,X48)|~m1_subset_1(X49,X46)|~m1_subset_1(X48,k3_zfmisc_1(X45,X46,X47))|X47=k1_xboole_0|X46=k1_xboole_0|X45=k1_xboole_0)&(X49!=esk12_5(X45,X46,X47,X48,X49)|X49=k6_mcart_1(X45,X46,X47,X48)|~m1_subset_1(X49,X46)|~m1_subset_1(X48,k3_zfmisc_1(X45,X46,X47))|X47=k1_xboole_0|X46=k1_xboole_0|X45=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_mcart_1])])])])])])])).
fof(c_0_20, plain, ![X20, X21, X22, X23]:(~m1_subset_1(X23,k3_zfmisc_1(X20,X21,X22))|m1_subset_1(k6_mcart_1(X20,X21,X22,X23),X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k6_mcart_1])])).
cnf(c_0_21, plain, (X1=X8|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1!=k7_mcart_1(X2,X3,X4,X5)|X5!=k4_tarski(k4_tarski(X6,X7),X8)|~m1_subset_1(X1,X4)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_22, plain, (m1_subset_1(k7_mcart_1(X2,X3,X4,X1),X4)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[c_0_17, c_0_15])).
cnf(c_0_23, negated_conjecture, (esk4_0=k3_mcart_1(esk5_0,esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (X1=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1!=k6_mcart_1(X2,X3,X4,X5)|X5!=k3_mcart_1(X6,X7,X8)|~m1_subset_1(X1,X3)|~m1_subset_1(X5,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (m1_subset_1(k6_mcart_1(X2,X3,X4,X1),X3)|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X27, X28, X29, X30, X31, X32, X33, X34]:((X31!=k5_mcart_1(X27,X28,X29,X30)|(X30!=k3_mcart_1(X32,X33,X34)|X31=X32)|~m1_subset_1(X31,X27)|~m1_subset_1(X30,k3_zfmisc_1(X27,X28,X29))|X29=k1_xboole_0|X28=k1_xboole_0|X27=k1_xboole_0)&((X30=k3_mcart_1(esk8_5(X27,X28,X29,X30,X31),esk9_5(X27,X28,X29,X30,X31),esk10_5(X27,X28,X29,X30,X31))|X31=k5_mcart_1(X27,X28,X29,X30)|~m1_subset_1(X31,X27)|~m1_subset_1(X30,k3_zfmisc_1(X27,X28,X29))|X29=k1_xboole_0|X28=k1_xboole_0|X27=k1_xboole_0)&(X31!=esk8_5(X27,X28,X29,X30,X31)|X31=k5_mcart_1(X27,X28,X29,X30)|~m1_subset_1(X31,X27)|~m1_subset_1(X30,k3_zfmisc_1(X27,X28,X29))|X29=k1_xboole_0|X28=k1_xboole_0|X27=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_mcart_1])])])])])])])).
fof(c_0_27, plain, ![X16, X17, X18, X19]:(~m1_subset_1(X19,k3_zfmisc_1(X16,X17,X18))|m1_subset_1(k5_mcart_1(X16,X17,X18,X19),X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_mcart_1])])).
cnf(c_0_28, plain, (k7_mcart_1(X1,X2,X3,k4_tarski(k4_tarski(X4,X5),X6))=X6|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(X4,X5),X6),k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_21])]), c_0_22])).
cnf(c_0_29, negated_conjecture, (esk4_0=k4_tarski(k4_tarski(esk5_0,esk6_0),esk7_0)), inference(rw,[status(thm)],[c_0_23, c_0_16])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk4_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, plain, (X1=X7|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1!=k6_mcart_1(X2,X3,X4,X5)|X5!=k4_tarski(k4_tarski(X6,X7),X8)|~m1_subset_1(X1,X3)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_15]), c_0_16])).
cnf(c_0_32, plain, (m1_subset_1(k6_mcart_1(X2,X3,X4,X1),X3)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[c_0_25, c_0_15])).
cnf(c_0_33, plain, (X1=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1!=k5_mcart_1(X2,X3,X4,X5)|X5!=k3_mcart_1(X6,X7,X8)|~m1_subset_1(X1,X2)|~m1_subset_1(X5,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (m1_subset_1(k5_mcart_1(X2,X3,X4,X1),X2)|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_35, negated_conjecture, (k7_mcart_1(X1,X2,X3,esk4_0)=esk7_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk4_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk4_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_30, c_0_15])).
cnf(c_0_37, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_38, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_39, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_40, plain, (k6_mcart_1(X1,X2,X3,k4_tarski(k4_tarski(X4,X5),X6))=X5|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(X4,X5),X6),k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])]), c_0_32])).
cnf(c_0_41, plain, (X1=X6|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1!=k5_mcart_1(X2,X3,X4,X5)|X5!=k4_tarski(k4_tarski(X6,X7),X8)|~m1_subset_1(X1,X2)|~m1_subset_1(X5,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_15]), c_0_16])).
cnf(c_0_42, plain, (m1_subset_1(k5_mcart_1(X2,X3,X4,X1),X2)|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[c_0_34, c_0_15])).
cnf(c_0_43, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk5_0|k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk6_0|k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk7_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_44, negated_conjecture, (k7_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk7_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_45, negated_conjecture, (k6_mcart_1(X1,X2,X3,esk4_0)=esk6_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk4_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_40, c_0_29])).
cnf(c_0_46, plain, (k5_mcart_1(X1,X2,X3,k4_tarski(k4_tarski(X4,X5),X6))=X4|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(X4,X5),X6),k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(csr,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_41])]), c_0_42])).
cnf(c_0_47, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk5_0|k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_44])])).
cnf(c_0_48, negated_conjecture, (k6_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)=esk6_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_36]), c_0_37]), c_0_38]), c_0_39])).
cnf(c_0_49, negated_conjecture, (k5_mcart_1(X1,X2,X3,esk4_0)=esk5_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk4_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_46, c_0_29])).
cnf(c_0_50, negated_conjecture, (k5_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0)!=esk5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48])])).
cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_36]), c_0_50]), c_0_37]), c_0_38]), c_0_39]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 33
# Proof object formula steps           : 19
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 18
# Processed clauses                    : 49
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 49
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 26
# ...of the previous two non-trivial   : 27
# Contextual simplify-reflections      : 3
# Paramodulations                      : 23
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 14
# ...number of literals in the above   : 144
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 80
# Rec. Clause-clause subsumption calls : 7
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 21
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2628

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28664 pages
