# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t88_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))=X4=>![X6]:(m1_subset_1(X6,u1_struct_0(X4))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X6=X7&r1_tmap_1(X3,X2,k3_tmap_1(X1,X2,X4,X3,X5),X7))=>r1_tmap_1(X4,X2,X5,X6)))))))))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', t88_tmap_1)).
fof(t11_tmap_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))&m1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', t11_tmap_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', dt_m1_pre_topc)).
fof(free_g1_pre_topc, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3, X4]:(g1_pre_topc(X1,X2)=g1_pre_topc(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', free_g1_pre_topc)).
fof(abstractness_v1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', abstractness_v1_pre_topc)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', dt_u1_pre_topc)).
fof(t16_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', t16_tsep_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', t1_tsep_1)).
fof(t65_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(m1_pre_topc(X1,X2)<=>m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', t65_pre_topc)).
fof(t86_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>((v1_tsep_1(X3,X4)&m1_pre_topc(X3,X4))=>![X6]:(m1_subset_1(X6,u1_struct_0(X4))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>(X6=X7=>(r1_tmap_1(X4,X2,X5,X6)<=>r1_tmap_1(X3,X2,k3_tmap_1(X1,X2,X4,X3,X5),X7))))))))))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', t86_tmap_1)).
fof(t7_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_pre_topc(X3,X2)=>m1_pre_topc(X3,X1)))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', t7_tsep_1)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', cc1_pre_topc)).
fof(t2_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>m1_pre_topc(X1,X1)), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', t2_tsep_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', d1_pre_topc)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/d70106550b4c4637bcf6aa1f6178dbb8.E_input', d5_pre_topc)).
fof(c_0_15, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))=X4=>![X6]:(m1_subset_1(X6,u1_struct_0(X4))=>![X7]:(m1_subset_1(X7,u1_struct_0(X3))=>((X6=X7&r1_tmap_1(X3,X2,k3_tmap_1(X1,X2,X4,X3,X5),X7))=>r1_tmap_1(X4,X2,X5,X6))))))))))), inference(assume_negation,[status(cth)],[t88_tmap_1])).
fof(c_0_16, plain, ![X15, X16]:((v1_pre_topc(g1_pre_topc(u1_struct_0(X16),u1_pre_topc(X16)))|~m1_pre_topc(X16,X15)|~l1_pre_topc(X15))&(m1_pre_topc(g1_pre_topc(u1_struct_0(X16),u1_pre_topc(X16)),X15)|~m1_pre_topc(X16,X15)|~l1_pre_topc(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_tmap_1])])])])).
fof(c_0_17, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&m1_pre_topc(esk3_0,esk1_0))&((~v2_struct_0(esk4_0)&m1_pre_topc(esk4_0,esk1_0))&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))))&(g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk3_0))=esk4_0&(m1_subset_1(esk6_0,u1_struct_0(esk4_0))&(m1_subset_1(esk7_0,u1_struct_0(esk3_0))&((esk6_0=esk7_0&r1_tmap_1(esk3_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk5_0),esk7_0))&~r1_tmap_1(esk4_0,esk2_0,esk5_0,esk6_0)))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_15])])])])).
fof(c_0_18, plain, ![X17, X18]:(~l1_pre_topc(X17)|(~m1_pre_topc(X18,X17)|l1_pre_topc(X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_19, plain, ![X19, X20, X21, X22]:((X19=X21|g1_pre_topc(X19,X20)!=g1_pre_topc(X21,X22)|~m1_subset_1(X20,k1_zfmisc_1(k1_zfmisc_1(X19))))&(X20=X22|g1_pre_topc(X19,X20)!=g1_pre_topc(X21,X22)|~m1_subset_1(X20,k1_zfmisc_1(k1_zfmisc_1(X19))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_pre_topc])])])])).
fof(c_0_20, plain, ![X23]:(~l1_pre_topc(X23)|(~v1_pre_topc(X23)|X23=g1_pre_topc(u1_struct_0(X23),u1_pre_topc(X23)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
cnf(c_0_21, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk3_0))=esk4_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, negated_conjecture, (m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (X1=X2|g1_pre_topc(X3,X1)!=g1_pre_topc(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3)))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_28, plain, ![X24]:(~l1_pre_topc(X24)|m1_subset_1(u1_pre_topc(X24),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X24))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_29, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, negated_conjecture, (v1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23]), c_0_24])])).
cnf(c_0_31, negated_conjecture, (l1_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_24])])).
cnf(c_0_32, negated_conjecture, (u1_pre_topc(esk3_0)=X1|g1_pre_topc(X2,X1)!=esk4_0|~m1_subset_1(u1_pre_topc(esk3_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk3_0))))), inference(spm,[status(thm)],[c_0_27, c_0_23])).
cnf(c_0_33, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_22]), c_0_24])])).
cnf(c_0_35, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, negated_conjecture, (g1_pre_topc(u1_struct_0(esk4_0),u1_pre_topc(esk4_0))=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
cnf(c_0_37, negated_conjecture, (u1_pre_topc(esk3_0)=X1|g1_pre_topc(X2,X1)!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])])).
fof(c_0_38, plain, ![X33, X34, X35]:((~v1_tsep_1(X34,X33)|~m1_pre_topc(X34,X33)|v3_pre_topc(X35,X33)|X35!=u1_struct_0(X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X33)))|~m1_pre_topc(X34,X33)|(~v2_pre_topc(X33)|~l1_pre_topc(X33)))&((v1_tsep_1(X34,X33)|~v3_pre_topc(X35,X33)|X35!=u1_struct_0(X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X33)))|~m1_pre_topc(X34,X33)|(~v2_pre_topc(X33)|~l1_pre_topc(X33)))&(m1_pre_topc(X34,X33)|~v3_pre_topc(X35,X33)|X35!=u1_struct_0(X34)|~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X33)))|~m1_pre_topc(X34,X33)|(~v2_pre_topc(X33)|~l1_pre_topc(X33))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_tsep_1])])])])).
fof(c_0_39, plain, ![X36, X37]:(~l1_pre_topc(X36)|(~m1_pre_topc(X37,X36)|m1_subset_1(u1_struct_0(X37),k1_zfmisc_1(u1_struct_0(X36))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
cnf(c_0_40, negated_conjecture, (u1_struct_0(esk4_0)=X1|g1_pre_topc(X1,X2)!=esk4_0|~m1_subset_1(u1_pre_topc(esk4_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk4_0))))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, negated_conjecture, (u1_pre_topc(esk3_0)=u1_pre_topc(esk4_0)), inference(spm,[status(thm)],[c_0_37, c_0_36])).
fof(c_0_42, plain, ![X25, X26]:((~m1_pre_topc(X25,X26)|m1_pre_topc(X25,g1_pre_topc(u1_struct_0(X26),u1_pre_topc(X26)))|~l1_pre_topc(X26)|~l1_pre_topc(X25))&(~m1_pre_topc(X25,g1_pre_topc(u1_struct_0(X26),u1_pre_topc(X26)))|m1_pre_topc(X25,X26)|~l1_pre_topc(X26)|~l1_pre_topc(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_pre_topc])])])])).
fof(c_0_43, plain, ![X38, X39, X40, X41, X42, X43, X44]:((~r1_tmap_1(X41,X39,X42,X43)|r1_tmap_1(X40,X39,k3_tmap_1(X38,X39,X41,X40,X42),X44)|X43!=X44|~m1_subset_1(X44,u1_struct_0(X40))|~m1_subset_1(X43,u1_struct_0(X41))|(~v1_tsep_1(X40,X41)|~m1_pre_topc(X40,X41))|(~v1_funct_1(X42)|(~v1_funct_2(X42,u1_struct_0(X41),u1_struct_0(X39))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X39))))))|(v2_struct_0(X41)|~m1_pre_topc(X41,X38))|(v2_struct_0(X40)|~m1_pre_topc(X40,X38))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))&(~r1_tmap_1(X40,X39,k3_tmap_1(X38,X39,X41,X40,X42),X44)|r1_tmap_1(X41,X39,X42,X43)|X43!=X44|~m1_subset_1(X44,u1_struct_0(X40))|~m1_subset_1(X43,u1_struct_0(X41))|(~v1_tsep_1(X40,X41)|~m1_pre_topc(X40,X41))|(~v1_funct_1(X42)|(~v1_funct_2(X42,u1_struct_0(X41),u1_struct_0(X39))|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X41),u1_struct_0(X39))))))|(v2_struct_0(X41)|~m1_pre_topc(X41,X38))|(v2_struct_0(X40)|~m1_pre_topc(X40,X38))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39)))|(v2_struct_0(X38)|(~v2_pre_topc(X38)|~l1_pre_topc(X38))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t86_tmap_1])])])])])).
fof(c_0_44, plain, ![X27, X28, X29]:(~v2_pre_topc(X27)|~l1_pre_topc(X27)|(~m1_pre_topc(X28,X27)|(~m1_pre_topc(X29,X28)|m1_pre_topc(X29,X27)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tsep_1])])])).
cnf(c_0_45, plain, (v1_tsep_1(X1,X2)|~v3_pre_topc(X3,X2)|X3!=u1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, negated_conjecture, (u1_struct_0(esk4_0)=X1|g1_pre_topc(X1,X2)!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_33]), c_0_31])])).
cnf(c_0_48, negated_conjecture, (g1_pre_topc(u1_struct_0(esk3_0),u1_pre_topc(esk4_0))=esk4_0), inference(rw,[status(thm)],[c_0_23, c_0_41])).
cnf(c_0_49, plain, (v3_pre_topc(X3,X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(X1,X2)|X3!=u1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_51, plain, ![X31, X32]:(~v2_pre_topc(X31)|~l1_pre_topc(X31)|(~m1_pre_topc(X32,X31)|v2_pre_topc(X32))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
cnf(c_0_52, plain, (r1_tmap_1(X4,X2,X5,X7)|v2_struct_0(X4)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~r1_tmap_1(X1,X2,k3_tmap_1(X3,X2,X4,X1,X5),X6)|X7!=X6|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X4))|~v1_tsep_1(X1,X4)|~m1_pre_topc(X1,X4)|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_pre_topc(X4,X3)|~m1_pre_topc(X1,X3)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (m1_pre_topc(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_pre_topc(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_54, plain, (v1_tsep_1(X1,X2)|~v3_pre_topc(u1_struct_0(X1),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_45]), c_0_46])).
cnf(c_0_55, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk4_0)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_56, plain, (v3_pre_topc(X3,X2)|X3!=u1_struct_0(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)|~v1_tsep_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(cn,[status(thm)],[c_0_49])).
cnf(c_0_57, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_58, plain, (m1_pre_topc(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[c_0_50, c_0_25])).
fof(c_0_59, plain, ![X30]:(~l1_pre_topc(X30)|m1_pre_topc(X30,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tsep_1])])).
fof(c_0_60, plain, ![X45, X46, X47, X48]:(((r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|~v2_pre_topc(X45)|~l1_pre_topc(X45))&((~m1_subset_1(X46,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X45))))|(~r1_tarski(X46,u1_pre_topc(X45))|r2_hidden(k5_setfam_1(u1_struct_0(X45),X46),u1_pre_topc(X45)))|~v2_pre_topc(X45)|~l1_pre_topc(X45))&(~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X45)))|(~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X45)))|(~r2_hidden(X47,u1_pre_topc(X45))|~r2_hidden(X48,u1_pre_topc(X45))|r2_hidden(k9_subset_1(u1_struct_0(X45),X47,X48),u1_pre_topc(X45))))|~v2_pre_topc(X45)|~l1_pre_topc(X45))))&(((m1_subset_1(esk9_1(X45),k1_zfmisc_1(u1_struct_0(X45)))|m1_subset_1(esk8_1(X45),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X45))))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&((m1_subset_1(esk10_1(X45),k1_zfmisc_1(u1_struct_0(X45)))|m1_subset_1(esk8_1(X45),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X45))))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&(((r2_hidden(esk9_1(X45),u1_pre_topc(X45))|m1_subset_1(esk8_1(X45),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X45))))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&(r2_hidden(esk10_1(X45),u1_pre_topc(X45))|m1_subset_1(esk8_1(X45),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X45))))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45)))&(~r2_hidden(k9_subset_1(u1_struct_0(X45),esk9_1(X45),esk10_1(X45)),u1_pre_topc(X45))|m1_subset_1(esk8_1(X45),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X45))))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45)))))&(((m1_subset_1(esk9_1(X45),k1_zfmisc_1(u1_struct_0(X45)))|r1_tarski(esk8_1(X45),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&((m1_subset_1(esk10_1(X45),k1_zfmisc_1(u1_struct_0(X45)))|r1_tarski(esk8_1(X45),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&(((r2_hidden(esk9_1(X45),u1_pre_topc(X45))|r1_tarski(esk8_1(X45),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&(r2_hidden(esk10_1(X45),u1_pre_topc(X45))|r1_tarski(esk8_1(X45),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45)))&(~r2_hidden(k9_subset_1(u1_struct_0(X45),esk9_1(X45),esk10_1(X45)),u1_pre_topc(X45))|r1_tarski(esk8_1(X45),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45)))))&((m1_subset_1(esk9_1(X45),k1_zfmisc_1(u1_struct_0(X45)))|~r2_hidden(k5_setfam_1(u1_struct_0(X45),esk8_1(X45)),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&((m1_subset_1(esk10_1(X45),k1_zfmisc_1(u1_struct_0(X45)))|~r2_hidden(k5_setfam_1(u1_struct_0(X45),esk8_1(X45)),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&(((r2_hidden(esk9_1(X45),u1_pre_topc(X45))|~r2_hidden(k5_setfam_1(u1_struct_0(X45),esk8_1(X45)),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45))&(r2_hidden(esk10_1(X45),u1_pre_topc(X45))|~r2_hidden(k5_setfam_1(u1_struct_0(X45),esk8_1(X45)),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45)))&(~r2_hidden(k9_subset_1(u1_struct_0(X45),esk9_1(X45),esk10_1(X45)),u1_pre_topc(X45))|~r2_hidden(k5_setfam_1(u1_struct_0(X45),esk8_1(X45)),u1_pre_topc(X45))|~r2_hidden(u1_struct_0(X45),u1_pre_topc(X45))|v2_pre_topc(X45)|~l1_pre_topc(X45)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
cnf(c_0_61, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_63, plain, (r1_tmap_1(X1,X2,X3,X4)|v2_struct_0(X1)|v2_struct_0(X5)|v2_struct_0(X2)|v2_struct_0(X6)|~v1_tsep_1(X6,X1)|~r1_tmap_1(X6,X2,k3_tmap_1(X5,X2,X1,X6,X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X6))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~m1_pre_topc(X1,X5)|~m1_pre_topc(X6,X1)|~l1_pre_topc(X5)|~l1_pre_topc(X2)|~v2_pre_topc(X5)|~v2_pre_topc(X2)), inference(er,[status(thm)],[inference(csr,[status(thm)],[c_0_52, c_0_53])])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_65, negated_conjecture, (~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_66, negated_conjecture, (v1_tsep_1(esk3_0,X1)|~v3_pre_topc(u1_struct_0(esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_67, plain, (v3_pre_topc(u1_struct_0(X1),X2)|~v1_tsep_1(X1,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_56]), c_0_46])).
cnf(c_0_68, negated_conjecture, (m1_pre_topc(esk4_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_57, c_0_23])).
cnf(c_0_69, negated_conjecture, (m1_pre_topc(X1,esk4_0)|~m1_pre_topc(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_23]), c_0_34])])).
cnf(c_0_70, plain, (m1_pre_topc(X1,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_71, plain, ![X52, X53]:((~v3_pre_topc(X53,X52)|r2_hidden(X53,u1_pre_topc(X52))|~m1_subset_1(X53,k1_zfmisc_1(u1_struct_0(X52)))|~l1_pre_topc(X52))&(~r2_hidden(X53,u1_pre_topc(X52))|v3_pre_topc(X53,X52)|~m1_subset_1(X53,k1_zfmisc_1(u1_struct_0(X52)))|~l1_pre_topc(X52))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
cnf(c_0_72, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_73, negated_conjecture, (v2_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_22]), c_0_24]), c_0_62])])).
cnf(c_0_74, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~m1_pre_topc(X3,X2)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_57]), c_0_25])).
cnf(c_0_75, negated_conjecture, (r1_tmap_1(esk4_0,X1,X2,esk6_0)|v2_struct_0(X3)|v2_struct_0(X1)|v2_struct_0(X4)|~v1_tsep_1(X4,esk4_0)|~r1_tmap_1(X4,X1,k3_tmap_1(X3,X1,esk4_0,X4,X2),esk6_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~m1_subset_1(esk6_0,u1_struct_0(X4))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(X1))|~v1_funct_1(X2)|~m1_pre_topc(esk4_0,X3)|~m1_pre_topc(X4,esk4_0)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X3)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
cnf(c_0_76, negated_conjecture, (v1_tsep_1(esk3_0,X1)|~v1_tsep_1(esk4_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_77, negated_conjecture, (m1_pre_topc(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_34])])).
cnf(c_0_78, negated_conjecture, (v2_pre_topc(esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_26]), c_0_24]), c_0_62])])).
cnf(c_0_79, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_80, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_81, negated_conjecture, (r2_hidden(u1_struct_0(esk3_0),u1_pre_topc(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_41]), c_0_34]), c_0_73])])).
cnf(c_0_82, negated_conjecture, (m1_pre_topc(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_70]), c_0_34])])).
cnf(c_0_83, plain, (m1_pre_topc(X1,X2)|~m1_pre_topc(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)),X2)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_53, c_0_58])).
cnf(c_0_84, negated_conjecture, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),esk1_0)|~m1_pre_topc(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_22]), c_0_24]), c_0_62])])).
cnf(c_0_85, negated_conjecture, (r1_tmap_1(esk4_0,X1,X2,esk6_0)|v2_struct_0(X1)|v2_struct_0(X3)|~v1_tsep_1(esk4_0,esk4_0)|~r1_tmap_1(esk3_0,X1,k3_tmap_1(X3,X1,esk4_0,esk3_0,X2),esk6_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(X1))|~v1_funct_1(X2)|~m1_pre_topc(esk4_0,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X3)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_55]), c_0_64]), c_0_77]), c_0_31]), c_0_78])]), c_0_79])).
cnf(c_0_86, plain, (v1_tsep_1(X1,X2)|~r2_hidden(u1_struct_0(X1),u1_pre_topc(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_80]), c_0_46])).
cnf(c_0_87, negated_conjecture, (r2_hidden(u1_struct_0(esk4_0),u1_pre_topc(esk4_0))), inference(rw,[status(thm)],[c_0_81, c_0_55])).
cnf(c_0_88, negated_conjecture, (m1_pre_topc(esk4_0,esk4_0)), inference(spm,[status(thm)],[c_0_69, c_0_82])).
cnf(c_0_89, negated_conjecture, (m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X2,esk3_0)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_24]), c_0_62])])).
cnf(c_0_90, negated_conjecture, (m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_53, c_0_77])).
cnf(c_0_91, negated_conjecture, (r1_tmap_1(esk4_0,X1,X2,esk6_0)|v2_struct_0(X3)|v2_struct_0(X1)|~r1_tmap_1(esk3_0,X1,k3_tmap_1(X3,X1,esk4_0,esk3_0,X2),esk6_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(X1))|~v1_funct_1(X2)|~m1_pre_topc(esk4_0,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X3)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_88]), c_0_31]), c_0_78])])).
cnf(c_0_92, negated_conjecture, (m1_pre_topc(X1,esk1_0)|~m1_pre_topc(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_34]), c_0_82]), c_0_73])])).
cnf(c_0_93, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_94, negated_conjecture, (r1_tmap_1(esk3_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk5_0),esk7_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_95, negated_conjecture, (esk6_0=esk7_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_96, negated_conjecture, (r1_tmap_1(esk4_0,X1,X2,esk6_0)|v2_struct_0(X1)|~r1_tmap_1(esk3_0,X1,k3_tmap_1(esk1_0,X1,esk4_0,esk3_0,X2),esk6_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_24]), c_0_62]), c_0_82])]), c_0_93])).
cnf(c_0_97, negated_conjecture, (v1_funct_2(esk5_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_98, negated_conjecture, (r1_tmap_1(esk3_0,esk2_0,k3_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk5_0),esk6_0)), inference(rw,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_99, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_100, negated_conjecture, (v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_101, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_102, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_103, negated_conjecture, (~r1_tmap_1(esk4_0,esk2_0,esk5_0,esk6_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_104, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_105, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98]), c_0_99]), c_0_100]), c_0_101]), c_0_102])]), c_0_103]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 75
# Proof object formula steps           : 31
# Proof object conjectures             : 53
# Proof object clause conjectures      : 50
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 76
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 57
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 56
# Processed clauses                    : 823
# ...of these trivial                  : 15
# ...subsumed                          : 481
# ...remaining for further processing  : 327
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 20
# Backward-rewritten                   : 8
# Generated clauses                    : 1995
# ...of the previous two non-trivial   : 1390
# Contextual simplify-reflections      : 64
# Paramodulations                      : 1979
# Factorizations                       : 0
# Equation resolutions                 : 16
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 241
#    Positive orientable unit clauses  : 28
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 207
# Current number of unprocessed clauses: 655
# ...number of literals in the above   : 3001
# Current number of archived formulas  : 0
# Current number of archived clauses   : 84
# Clause-clause subsumption calls (NU) : 21687
# Rec. Clause-clause subsumption calls : 7948
# Non-unit clause-clause subsumptions  : 564
# Unit Clause-clause subsumption calls : 274
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 50842

# -------------------------------------------------
# User time                : 0.031 s
# System time              : 0.007 s
# Total time               : 0.038 s
# Maximum resident set size: 28632 pages
