# No SInE strategy applied
# Auto-Mode selected heuristic G_E___200_B02_F1_SE_CS_SP_PI_S2k
# and selection function SelectCQArNTEqFirst.
#
# Preprocessing time       : 0.010 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t68_tex_2, conjecture, ![X1]:((v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1)))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2)))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t68_tex_2)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t5_subset)).
fof(dt_k6_partfun1, axiom, ![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', dt_k6_partfun1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t3_subset)).
fof(t17_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v3_pre_topc(X2,X1)))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t17_tdlat_3)).
fof(cc1_tdlat_3, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_tdlat_3(X1)=>v2_pre_topc(X1))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', cc1_tdlat_3)).
fof(dt_k8_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k8_relset_1(X1,X2,X3,X4),k1_zfmisc_1(X1))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', dt_k8_relset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', d1_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t113_zfmisc_1)).
fof(rc1_funct_2, axiom, ![X1, X2]:?[X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&(v1_relat_1(X3)&(v4_relat_1(X3,X1)&(v5_relat_1(X3,X2)&(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', rc1_funct_2)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', d4_subset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', d10_xboole_0)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', cc4_relset_1)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t8_boole)).
fof(t55_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>((k2_struct_0(X2)=k1_xboole_0=>k2_struct_0(X1)=k1_xboole_0)=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X4,X2)=>v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4),X1)))))))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t55_tops_2)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', l13_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', fc1_xboole_0)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k1_relset_1(X1,X2,X3)=k1_relat_1(X3)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', redefinition_k1_relset_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t2_xboole_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', d3_struct_0)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', fc6_relat_1)).
fof(t38_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(k7_relset_1(X1,X2,X3,X1)=k2_relset_1(X1,X2,X3)&k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t38_relset_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', dt_l1_pre_topc)).
fof(t56_tops_2, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>(v5_pre_topc(X3,X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))=>r1_tarski(k2_pre_topc(X1,k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,X4)),k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X2,X4)))))))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t56_tops_2)).
fof(d4_partfun1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>(v1_partfun1(X2,X1)<=>k1_relat_1(X2)=X1)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', d4_partfun1)).
fof(t18_tdlat_3, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>(v1_tdlat_3(X1)<=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>v4_pre_topc(X2,X1)))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t18_tdlat_3)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t4_subset_1)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/d70f6e1b427845109f8b7c3a8bea6ccf.E_input', t52_pre_topc)).
fof(c_0_29, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&(v1_tdlat_3(X1)&l1_pre_topc(X1)))=>![X2]:((v2_pre_topc(X2)&l1_pre_topc(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))))=>v5_pre_topc(X3,X1,X2))))), inference(assume_negation,[status(cth)],[t68_tex_2])).
fof(c_0_30, plain, ![X16, X17, X18]:(~r2_hidden(X16,X17)|(~m1_subset_1(X17,k1_zfmisc_1(X18))|~v1_xboole_0(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_31, plain, ![X19]:(v1_partfun1(k6_partfun1(X19),X19)&m1_subset_1(k6_partfun1(X19),k1_zfmisc_1(k2_zfmisc_1(X19,X19)))), inference(variable_rename,[status(thm)],[dt_k6_partfun1])).
fof(c_0_32, plain, ![X20, X21]:((~m1_subset_1(X20,k1_zfmisc_1(X21))|r1_tarski(X20,X21))&(~r1_tarski(X20,X21)|m1_subset_1(X20,k1_zfmisc_1(X21)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_33, negated_conjecture, ((v2_pre_topc(esk1_0)&(v1_tdlat_3(esk1_0)&l1_pre_topc(esk1_0)))&((v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))))&~v5_pre_topc(esk3_0,esk1_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_29])])])).
fof(c_0_34, plain, ![X8, X9]:((~v1_tdlat_3(X8)|(~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(X8)))|v3_pre_topc(X9,X8))|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))&((m1_subset_1(esk4_1(X8),k1_zfmisc_1(u1_struct_0(X8)))|v1_tdlat_3(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))&(~v3_pre_topc(esk4_1(X8),X8)|v1_tdlat_3(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_tdlat_3])])])])])).
fof(c_0_35, plain, ![X11]:(~l1_pre_topc(X11)|(~v1_tdlat_3(X11)|v2_pre_topc(X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_tdlat_3])])).
fof(c_0_36, plain, ![X12, X13, X14, X15]:(~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X12,X13)))|m1_subset_1(k8_relset_1(X12,X13,X14,X15),k1_zfmisc_1(X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_relset_1])])).
cnf(c_0_37, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, ![X22, X23, X24]:((~v1_xboole_0(X22)|~r2_hidden(X23,X22))&(r2_hidden(esk5_1(X24),X24)|v1_xboole_0(X24))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_40, plain, ![X26, X27]:((k2_zfmisc_1(X26,X27)!=k1_xboole_0|(X26=k1_xboole_0|X27=k1_xboole_0))&((X26!=k1_xboole_0|k2_zfmisc_1(X26,X27)=k1_xboole_0)&(X27!=k1_xboole_0|k2_zfmisc_1(X26,X27)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_41, plain, ![X28, X29]:(m1_subset_1(esk6_2(X28,X29),k1_zfmisc_1(k2_zfmisc_1(X28,X29)))&(v1_relat_1(esk6_2(X28,X29))&(v4_relat_1(esk6_2(X28,X29),X28)&(v5_relat_1(esk6_2(X28,X29),X29)&(v1_funct_1(esk6_2(X28,X29))&v1_funct_2(esk6_2(X28,X29),X28,X29)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_funct_2])])).
fof(c_0_42, plain, ![X44]:m1_subset_1(k2_subset_1(X44),k1_zfmisc_1(X44)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_43, plain, ![X45]:k2_subset_1(X45)=X45, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_44, plain, ![X31, X32]:(((r1_tarski(X31,X32)|X31!=X32)&(r1_tarski(X32,X31)|X31!=X32))&(~r1_tarski(X31,X32)|~r1_tarski(X32,X31)|X31=X32)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_45, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_47, plain, (v3_pre_topc(X2,X1)|~v1_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48, plain, (v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_tdlat_3(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_49, plain, (m1_subset_1(k8_relset_1(X2,X3,X1,X4),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_50, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X1))|~r2_hidden(X2,k6_partfun1(X1))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_51, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_52, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_53, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_54, plain, ![X51, X52, X53]:(~v1_xboole_0(X51)|(~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X52,X51)))|v1_xboole_0(X53))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_55, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_57, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_58, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
fof(c_0_59, plain, ![X38, X39]:(~v1_xboole_0(X38)|(X38=X39|~v1_xboole_0(X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
fof(c_0_60, plain, ![X33, X34, X35, X36]:(((~v5_pre_topc(X35,X33,X34)|(~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X34)))|(~v3_pre_topc(X36,X34)|v3_pre_topc(k8_relset_1(u1_struct_0(X33),u1_struct_0(X34),X35,X36),X33)))|k2_struct_0(X34)=k1_xboole_0|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(X33),u1_struct_0(X34))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X34))))))|~l1_pre_topc(X34)|~l1_pre_topc(X33))&((m1_subset_1(esk7_3(X33,X34,X35),k1_zfmisc_1(u1_struct_0(X34)))|v5_pre_topc(X35,X33,X34)|k2_struct_0(X34)=k1_xboole_0|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(X33),u1_struct_0(X34))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X34))))))|~l1_pre_topc(X34)|~l1_pre_topc(X33))&((v3_pre_topc(esk7_3(X33,X34,X35),X34)|v5_pre_topc(X35,X33,X34)|k2_struct_0(X34)=k1_xboole_0|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(X33),u1_struct_0(X34))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X34))))))|~l1_pre_topc(X34)|~l1_pre_topc(X33))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X33),u1_struct_0(X34),X35,esk7_3(X33,X34,X35)),X33)|v5_pre_topc(X35,X33,X34)|k2_struct_0(X34)=k1_xboole_0|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(X33),u1_struct_0(X34))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X34))))))|~l1_pre_topc(X34)|~l1_pre_topc(X33)))))&((~v5_pre_topc(X35,X33,X34)|(~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X34)))|(~v3_pre_topc(X36,X34)|v3_pre_topc(k8_relset_1(u1_struct_0(X33),u1_struct_0(X34),X35,X36),X33)))|k2_struct_0(X33)!=k1_xboole_0|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(X33),u1_struct_0(X34))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X34))))))|~l1_pre_topc(X34)|~l1_pre_topc(X33))&((m1_subset_1(esk7_3(X33,X34,X35),k1_zfmisc_1(u1_struct_0(X34)))|v5_pre_topc(X35,X33,X34)|k2_struct_0(X33)!=k1_xboole_0|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(X33),u1_struct_0(X34))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X34))))))|~l1_pre_topc(X34)|~l1_pre_topc(X33))&((v3_pre_topc(esk7_3(X33,X34,X35),X34)|v5_pre_topc(X35,X33,X34)|k2_struct_0(X33)!=k1_xboole_0|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(X33),u1_struct_0(X34))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X34))))))|~l1_pre_topc(X34)|~l1_pre_topc(X33))&(~v3_pre_topc(k8_relset_1(u1_struct_0(X33),u1_struct_0(X34),X35,esk7_3(X33,X34,X35)),X33)|v5_pre_topc(X35,X33,X34)|k2_struct_0(X33)!=k1_xboole_0|(~v1_funct_1(X35)|(~v1_funct_2(X35,u1_struct_0(X33),u1_struct_0(X34))|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X33),u1_struct_0(X34))))))|~l1_pre_topc(X34)|~l1_pre_topc(X33)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_tops_2])])])])])).
cnf(c_0_61, plain, (v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_62, negated_conjecture, (m1_subset_1(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_49, c_0_46])).
cnf(c_0_63, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_64, negated_conjecture, (v1_tdlat_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_65, plain, ![X41]:(~v1_xboole_0(X41)|X41=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_66, plain, (v1_xboole_0(k6_partfun1(X1))|~v1_xboole_0(k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_67, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_68, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_69, plain, (~v1_xboole_0(k2_zfmisc_1(X1,X2))|~r2_hidden(X3,esk6_2(X1,X2))), inference(spm,[status(thm)],[c_0_37, c_0_53])).
fof(c_0_70, plain, ![X48, X49, X50]:(~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(X48,X49)))|k1_relset_1(X48,X49,X50)=k1_relat_1(X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_71, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_72, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_73, negated_conjecture, (k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))=esk3_0|~r1_tarski(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)),esk3_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_74, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_75, plain, ![X42]:r1_tarski(k1_xboole_0,X42), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_76, plain, ![X40]:(~l1_struct_0(X40)|k2_struct_0(X40)=u1_struct_0(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_77, plain, (v5_pre_topc(X3,X1,X2)|k2_struct_0(X2)=k1_xboole_0|~v3_pre_topc(k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk7_3(X1,X2,X3)),X1)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_78, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_79, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_80, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_81, negated_conjecture, (~v5_pre_topc(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_82, negated_conjecture, (v3_pre_topc(k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,X1),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_64])])).
cnf(c_0_83, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_84, plain, (v1_xboole_0(k6_partfun1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])])).
fof(c_0_85, plain, ![X46, X47]:v1_relat_1(k2_zfmisc_1(X46,X47)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
cnf(c_0_86, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_87, plain, (v1_xboole_0(esk6_2(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_69, c_0_51])).
fof(c_0_88, plain, ![X56, X57, X58]:((k7_relset_1(X56,X57,X58,X56)=k2_relset_1(X56,X57,X58)|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(X56,X57))))&(k8_relset_1(X56,X57,X58,X57)=k1_relset_1(X56,X57,X58)|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(X56,X57))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_relset_1])])])).
cnf(c_0_89, plain, (k1_relset_1(X2,X3,X1)=k1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_90, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_91, negated_conjecture, (X1=esk3_0|~r1_tarski(X1,esk3_0)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_92, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_93, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_94, negated_conjecture, (k2_struct_0(esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_46]), c_0_79]), c_0_80]), c_0_63])]), c_0_81]), c_0_82])])).
fof(c_0_95, plain, ![X43]:(~l1_pre_topc(X43)|l1_struct_0(X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_96, plain, ![X59, X60, X61, X62]:((~v5_pre_topc(X61,X59,X60)|(~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X60)))|r1_tarski(k2_pre_topc(X59,k8_relset_1(u1_struct_0(X59),u1_struct_0(X60),X61,X62)),k8_relset_1(u1_struct_0(X59),u1_struct_0(X60),X61,k2_pre_topc(X60,X62))))|(~v1_funct_1(X61)|(~v1_funct_2(X61,u1_struct_0(X59),u1_struct_0(X60))|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X59),u1_struct_0(X60))))))|(~v2_pre_topc(X60)|~l1_pre_topc(X60))|(~v2_pre_topc(X59)|~l1_pre_topc(X59)))&((m1_subset_1(esk8_3(X59,X60,X61),k1_zfmisc_1(u1_struct_0(X60)))|v5_pre_topc(X61,X59,X60)|(~v1_funct_1(X61)|(~v1_funct_2(X61,u1_struct_0(X59),u1_struct_0(X60))|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X59),u1_struct_0(X60))))))|(~v2_pre_topc(X60)|~l1_pre_topc(X60))|(~v2_pre_topc(X59)|~l1_pre_topc(X59)))&(~r1_tarski(k2_pre_topc(X59,k8_relset_1(u1_struct_0(X59),u1_struct_0(X60),X61,esk8_3(X59,X60,X61))),k8_relset_1(u1_struct_0(X59),u1_struct_0(X60),X61,k2_pre_topc(X60,esk8_3(X59,X60,X61))))|v5_pre_topc(X61,X59,X60)|(~v1_funct_1(X61)|(~v1_funct_2(X61,u1_struct_0(X59),u1_struct_0(X60))|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X59),u1_struct_0(X60))))))|(~v2_pre_topc(X60)|~l1_pre_topc(X60))|(~v2_pre_topc(X59)|~l1_pre_topc(X59))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_tops_2])])])])])).
fof(c_0_97, plain, ![X54, X55]:((~v1_partfun1(X55,X54)|k1_relat_1(X55)=X54|(~v1_relat_1(X55)|~v4_relat_1(X55,X54)))&(k1_relat_1(X55)!=X54|v1_partfun1(X55,X54)|(~v1_relat_1(X55)|~v4_relat_1(X55,X54)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_partfun1])])])).
cnf(c_0_98, plain, (v1_partfun1(k6_partfun1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_99, plain, (k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_100, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_101, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_86])).
cnf(c_0_102, plain, (v1_xboole_0(esk6_2(k1_xboole_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_67]), c_0_68])])).
fof(c_0_103, plain, ![X64, X65]:((~v1_tdlat_3(X64)|(~m1_subset_1(X65,k1_zfmisc_1(u1_struct_0(X64)))|v4_pre_topc(X65,X64))|(~v2_pre_topc(X64)|~l1_pre_topc(X64)))&((m1_subset_1(esk9_1(X64),k1_zfmisc_1(u1_struct_0(X64)))|v1_tdlat_3(X64)|(~v2_pre_topc(X64)|~l1_pre_topc(X64)))&(~v4_pre_topc(esk9_1(X64),X64)|v1_tdlat_3(X64)|(~v2_pre_topc(X64)|~l1_pre_topc(X64))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_tdlat_3])])])])])).
cnf(c_0_104, plain, (k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_105, negated_conjecture, (k1_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0)=k1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_89, c_0_46])).
cnf(c_0_106, plain, (v1_xboole_0(esk6_2(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_87, c_0_90])).
cnf(c_0_107, negated_conjecture, (k1_xboole_0=esk3_0|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_68])])).
cnf(c_0_108, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0|~l1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_109, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_110, plain, (m1_subset_1(esk8_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|v5_pre_topc(X3,X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_111, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_112, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_113, plain, (k1_relat_1(X1)=X2|~v1_partfun1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_114, plain, (v1_partfun1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_115, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_116, plain, (v4_relat_1(esk6_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_117, plain, (esk6_2(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_83, c_0_102])).
cnf(c_0_118, plain, (v4_pre_topc(X2,X1)|~v1_tdlat_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_119, negated_conjecture, (k8_relset_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0),esk3_0,u1_struct_0(esk2_0))=k1_relat_1(esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_46]), c_0_105])).
cnf(c_0_120, plain, (v1_funct_2(esk6_2(X1,X2),X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_121, plain, (esk6_2(X1,X2)=k1_xboole_0|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_83, c_0_106])).
cnf(c_0_122, negated_conjecture, (k1_xboole_0=esk3_0|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_107, c_0_90])).
cnf(c_0_123, negated_conjecture, (u1_struct_0(esk2_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_80])])).
fof(c_0_124, plain, ![X67]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X67)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_125, negated_conjecture, (m1_subset_1(esk8_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_78]), c_0_46]), c_0_79]), c_0_80]), c_0_63]), c_0_111]), c_0_112])]), c_0_81])).
cnf(c_0_126, plain, (v1_xboole_0(esk6_2(X1,k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_101]), c_0_68])])).
cnf(c_0_127, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0|~v4_relat_1(k1_xboole_0,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_115])])).
cnf(c_0_128, plain, (v4_relat_1(k1_xboole_0,k1_xboole_0)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
fof(c_0_129, plain, ![X68, X69]:((~v4_pre_topc(X69,X68)|k2_pre_topc(X68,X69)=X69|~m1_subset_1(X69,k1_zfmisc_1(u1_struct_0(X68)))|~l1_pre_topc(X68))&(~v2_pre_topc(X68)|k2_pre_topc(X68,X69)!=X69|v4_pre_topc(X69,X68)|~m1_subset_1(X69,k1_zfmisc_1(u1_struct_0(X68)))|~l1_pre_topc(X68))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
cnf(c_0_130, plain, (v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_tdlat_3(X2)), inference(csr,[status(thm)],[c_0_118, c_0_48])).
cnf(c_0_131, negated_conjecture, (m1_subset_1(k1_relat_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_62, c_0_119])).
cnf(c_0_132, plain, (v1_funct_2(k1_xboole_0,X1,X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_133, negated_conjecture, (k1_xboole_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123]), c_0_68])])).
cnf(c_0_134, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_135, negated_conjecture, (r1_tarski(esk8_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_45, c_0_125])).
cnf(c_0_136, plain, (k1_relset_1(X1,X2,esk6_2(X1,X2))=k1_relat_1(esk6_2(X1,X2))), inference(spm,[status(thm)],[c_0_89, c_0_53])).
cnf(c_0_137, plain, (esk6_2(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_83, c_0_126])).
cnf(c_0_138, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_128])])).
cnf(c_0_139, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_140, negated_conjecture, (v4_pre_topc(k1_relat_1(esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_63]), c_0_64])])).
cnf(c_0_141, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_37, c_0_46])).
cnf(c_0_142, plain, (v5_pre_topc(X3,X1,X2)|~r1_tarski(k2_pre_topc(X1,k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,esk8_3(X1,X2,X3))),k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),X3,k2_pre_topc(X2,esk8_3(X1,X2,X3))))|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_143, plain, (v1_funct_2(esk3_0,X1,X2)|~v1_xboole_0(X2)), inference(rw,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_144, plain, (m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_134, c_0_133])).
cnf(c_0_145, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk2_0)|~r1_tarski(u1_struct_0(esk2_0),esk8_3(esk1_0,esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_57, c_0_135])).
cnf(c_0_146, plain, (k8_relset_1(X1,X2,esk6_2(X1,X2),X2)=k1_relat_1(esk6_2(X1,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_53]), c_0_136])).
cnf(c_0_147, plain, (esk6_2(X1,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_137, c_0_133]), c_0_133])).
cnf(c_0_148, plain, (k1_relat_1(esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_138, c_0_133]), c_0_133])).
cnf(c_0_149, negated_conjecture, (k2_pre_topc(esk1_0,k1_relat_1(esk3_0))=k1_relat_1(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_131]), c_0_63])]), c_0_140])])).
cnf(c_0_150, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_141, c_0_51])).
cnf(c_0_151, plain, (v5_pre_topc(esk3_0,X1,X2)|~r1_tarski(k2_pre_topc(X1,k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),esk3_0,esk8_3(X1,X2,esk3_0))),k8_relset_1(u1_struct_0(X1),u1_struct_0(X2),esk3_0,k2_pre_topc(X2,esk8_3(X1,X2,esk3_0))))|~v1_xboole_0(u1_struct_0(X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143]), c_0_144]), c_0_79])])).
cnf(c_0_152, negated_conjecture, (esk8_3(esk1_0,esk2_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_145, c_0_123]), c_0_123]), c_0_92])]), c_0_133])).
cnf(c_0_153, plain, (k8_relset_1(X1,esk3_0,esk3_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_148])).
cnf(c_0_154, negated_conjecture, (k2_pre_topc(esk1_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_149, c_0_148]), c_0_148])).
cnf(c_0_155, plain, (r1_tarski(esk3_0,X1)), inference(rw,[status(thm)],[c_0_92, c_0_133])).
cnf(c_0_156, negated_conjecture, (v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_150, c_0_123]), c_0_101]), c_0_68])])).
cnf(c_0_157, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_123]), c_0_133]), c_0_153]), c_0_154]), c_0_123]), c_0_133]), c_0_155]), c_0_123]), c_0_133]), c_0_156]), c_0_80]), c_0_63]), c_0_111]), c_0_112])]), c_0_81]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 158
# Proof object clause steps            : 100
# Proof object formula steps           : 58
# Proof object conjectures             : 37
# Proof object clause conjectures      : 34
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 42
# Proof object initial formulas used   : 29
# Proof object generating inferences   : 43
# Proof object simplifying inferences  : 88
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 65
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 64
# Processed clauses                    : 3105
# ...of these trivial                  : 61
# ...subsumed                          : 2240
# ...remaining for further processing  : 804
# Other redundant clauses eliminated   : 32
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 92
# Backward-rewritten                   : 164
# Generated clauses                    : 12095
# ...of the previous two non-trivial   : 11142
# Contextual simplify-reflections      : 21
# Paramodulations                      : 12063
# Factorizations                       : 0
# Equation resolutions                 : 32
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 543
#    Positive orientable unit clauses  : 67
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 474
# Current number of unprocessed clauses: 7506
# ...number of literals in the above   : 41047
# Current number of archived formulas  : 0
# Current number of archived clauses   : 257
# Clause-clause subsumption calls (NU) : 200272
# Rec. Clause-clause subsumption calls : 90462
# Non-unit clause-clause subsumptions  : 2189
# Unit Clause-clause subsumption calls : 5365
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 207
# BW rewrite match successes           : 29
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 174596

# -------------------------------------------------
# User time                : 0.156 s
# System time              : 0.003 s
# Total time               : 0.159 s
# Maximum resident set size: 28756 pages
