# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', d10_xboole_0)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', d1_zfmisc_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', t4_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', t3_subset)).
fof(t48_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r1_tarski(X2,X3)=>r1_tarski(k1_tops_1(X1,X2),k1_tops_1(X1,X3)))))), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', t48_tops_1)).
fof(t9_connsp_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((m1_connsp_2(X4,X1,X3)&r1_tarski(X4,X2)))))))))), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', t9_connsp_2)).
fof(t44_tops_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r1_tarski(k1_tops_1(X1,X2),X2))), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', t44_tops_1)).
fof(t5_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((v3_pre_topc(X2,X1)&r2_hidden(X3,X2))=>m1_connsp_2(X2,X1,X3))))), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', t5_connsp_2)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', d3_tarski)).
fof(d1_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(m1_connsp_2(X3,X1,X2)<=>r2_hidden(X2,k1_tops_1(X1,X3)))))), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', d1_connsp_2)).
fof(fc9_tops_1, axiom, ![X1, X2]:(((v2_pre_topc(X1)&l1_pre_topc(X1))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>v3_pre_topc(k1_tops_1(X1,X2),X1)), file('data/MPTP/proofs/d7298a9b01064dabad0164039535c540.E_input', fc9_tops_1)).
fof(c_0_11, plain, ![X5, X6]:(((r1_tarski(X5,X6)|X5!=X6)&(r1_tarski(X6,X5)|X5!=X6))&(~r1_tarski(X5,X6)|~r1_tarski(X6,X5)|X5=X6)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_12, plain, ![X7, X8, X9, X10, X11, X12]:(((~r2_hidden(X9,X8)|r1_tarski(X9,X7)|X8!=k1_zfmisc_1(X7))&(~r1_tarski(X10,X7)|r2_hidden(X10,X8)|X8!=k1_zfmisc_1(X7)))&((~r2_hidden(esk1_2(X11,X12),X12)|~r1_tarski(esk1_2(X11,X12),X11)|X12=k1_zfmisc_1(X11))&(r2_hidden(esk1_2(X11,X12),X12)|r1_tarski(esk1_2(X11,X12),X11)|X12=k1_zfmisc_1(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_13, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X23, X24, X25]:(~r2_hidden(X23,X24)|~m1_subset_1(X24,k1_zfmisc_1(X25))|m1_subset_1(X23,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_15, plain, ![X32, X33]:((~m1_subset_1(X32,k1_zfmisc_1(X33))|r1_tarski(X32,X33))&(~r1_tarski(X32,X33)|m1_subset_1(X32,k1_zfmisc_1(X33)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_16, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_17, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, plain, (r2_hidden(X1,k1_zfmisc_1(X1))), inference(er,[status(thm)],[c_0_20])).
fof(c_0_23, plain, ![X34, X35, X36]:(~l1_pre_topc(X34)|(~m1_subset_1(X35,k1_zfmisc_1(u1_struct_0(X34)))|(~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X34)))|(~r1_tarski(X35,X36)|r1_tarski(k1_tops_1(X34,X35),k1_tops_1(X34,X36)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t48_tops_1])])])).
cnf(c_0_24, plain, (m1_subset_1(X1,X2)|~r1_tarski(k1_zfmisc_1(X1),X2)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_25, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r2_hidden(X3,X2)&![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))=>~((m1_connsp_2(X4,X1,X3)&r1_tarski(X4,X2))))))))))), inference(assume_negation,[status(cth)],[t9_connsp_2])).
fof(c_0_26, plain, ![X37, X38]:(~l1_pre_topc(X37)|(~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37)))|r1_tarski(k1_tops_1(X37,X38),X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_tops_1])])])).
cnf(c_0_27, plain, (r1_tarski(k1_tops_1(X1,X2),k1_tops_1(X1,X3))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_24, c_0_17])).
cnf(c_0_29, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_30, negated_conjecture, ![X17, X18]:((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&(((m1_subset_1(esk4_0,u1_struct_0(esk2_0))|~v3_pre_topc(esk3_0,esk2_0))&((r2_hidden(esk4_0,esk3_0)|~v3_pre_topc(esk3_0,esk2_0))&(~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(esk2_0)))|(~m1_connsp_2(X17,esk2_0,esk4_0)|~r1_tarski(X17,esk3_0))|~v3_pre_topc(esk3_0,esk2_0))))&((m1_subset_1(esk5_1(X18),k1_zfmisc_1(u1_struct_0(esk2_0)))|~r2_hidden(X18,esk3_0)|~m1_subset_1(X18,u1_struct_0(esk2_0))|v3_pre_topc(esk3_0,esk2_0))&((m1_connsp_2(esk5_1(X18),esk2_0,X18)|~r2_hidden(X18,esk3_0)|~m1_subset_1(X18,u1_struct_0(esk2_0))|v3_pre_topc(esk3_0,esk2_0))&(r1_tarski(esk5_1(X18),esk3_0)|~r2_hidden(X18,esk3_0)|~m1_subset_1(X18,u1_struct_0(esk2_0))|v3_pre_topc(esk3_0,esk2_0))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_25])])])])])])).
fof(c_0_31, plain, ![X20, X21, X22]:(v2_struct_0(X20)|(~v2_pre_topc(X20)|~l1_pre_topc(X20))|(~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|(~m1_subset_1(X22,u1_struct_0(X20))|(~v3_pre_topc(X21,X20)|~r2_hidden(X22,X21)|m1_connsp_2(X21,X20,X22))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_connsp_2])])])])).
fof(c_0_32, plain, ![X26, X27, X28, X29, X30]:((~r1_tarski(X26,X27)|(~r2_hidden(X28,X26)|r2_hidden(X28,X27)))&((r2_hidden(esk6_2(X29,X30),X29)|r1_tarski(X29,X30))&(~r2_hidden(esk6_2(X29,X30),X30)|r1_tarski(X29,X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_33, plain, (r1_tarski(k1_tops_1(X1,X2),X2)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (r1_tarski(k1_tops_1(X1,X2),k1_tops_1(X1,u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_35, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (v2_struct_0(X1)|m1_connsp_2(X2,X1,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_pre_topc(X2,X1)|~r2_hidden(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (r1_tarski(k1_tops_1(X1,u1_struct_0(X1)),u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_40, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,esk3_0),k1_tops_1(esk2_0,u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_36])])).
cnf(c_0_41, negated_conjecture, (~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_connsp_2(X1,esk2_0,esk4_0)|~r1_tarski(X1,esk3_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (m1_connsp_2(X1,X2,X3)|v2_struct_0(X2)|~v3_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~r2_hidden(X3,X1)), inference(csr,[status(thm)],[c_0_37, c_0_18])).
cnf(c_0_43, negated_conjecture, (v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_45, plain, (r2_hidden(X1,u1_struct_0(X2))|~l1_pre_topc(X2)|~r2_hidden(X1,k1_tops_1(X2,u1_struct_0(X2)))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,u1_struct_0(esk2_0)))|~r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_38, c_0_40])).
cnf(c_0_47, negated_conjecture, (~m1_connsp_2(esk3_0,esk2_0,esk4_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_35]), c_0_17])])).
cnf(c_0_48, negated_conjecture, (m1_connsp_2(esk3_0,esk2_0,X1)|~v3_pre_topc(esk3_0,esk2_0)|~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_35]), c_0_36]), c_0_43])]), c_0_44])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_50, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_51, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46]), c_0_36])])).
cnf(c_0_52, negated_conjecture, (r1_tarski(esk5_1(X1),esk3_0)|v3_pre_topc(esk3_0,esk2_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_53, negated_conjecture, (~v3_pre_topc(esk3_0,esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_49])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_18, c_0_35])).
cnf(c_0_55, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk2_0))|~r2_hidden(esk6_2(X1,u1_struct_0(esk2_0)),k1_tops_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_56, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_57, negated_conjecture, (r2_hidden(esk5_1(X1),X2)|X2!=k1_zfmisc_1(esk3_0)|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_52]), c_0_53]), c_0_54])).
cnf(c_0_58, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_59, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_60, negated_conjecture, (m1_subset_1(esk5_1(X1),X2)|X3!=k1_zfmisc_1(esk3_0)|~r2_hidden(X1,esk3_0)|~r1_tarski(X3,X2)), inference(spm,[status(thm)],[c_0_21, c_0_57])).
cnf(c_0_61, plain, (r1_tarski(X1,X2)|k1_zfmisc_1(X1)!=k1_zfmisc_1(X2)), inference(spm,[status(thm)],[c_0_58, c_0_22])).
cnf(c_0_62, negated_conjecture, (r2_hidden(k1_tops_1(esk2_0,esk3_0),X1)|X1!=k1_zfmisc_1(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_16, c_0_59])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk5_1(X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|v3_pre_topc(esk3_0,esk2_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_64, negated_conjecture, (m1_subset_1(esk5_1(X1),X2)|k1_zfmisc_1(X3)!=k1_zfmisc_1(X2)|X3!=k1_zfmisc_1(esk3_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_65, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_18, c_0_28])).
cnf(c_0_66, negated_conjecture, (r2_hidden(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_67, negated_conjecture, (r1_tarski(esk5_1(X1),u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_63]), c_0_53]), c_0_54])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(esk5_1(X1),X2)|k1_zfmisc_1(k1_zfmisc_1(esk3_0))!=k1_zfmisc_1(X2)|~r2_hidden(X1,esk3_0)), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_69, negated_conjecture, (m1_subset_1(k1_tops_1(esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_70, negated_conjecture, (r2_hidden(esk5_1(X1),X2)|X2!=k1_zfmisc_1(u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_16, c_0_67])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(esk5_1(X1),k1_zfmisc_1(esk3_0))|~r2_hidden(X1,esk3_0)), inference(er,[status(thm)],[c_0_68])).
fof(c_0_72, plain, ![X39, X40, X41]:((~m1_connsp_2(X41,X39,X40)|r2_hidden(X40,k1_tops_1(X39,X41))|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X39)))|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))))&(~r2_hidden(X40,k1_tops_1(X39,X41))|m1_connsp_2(X41,X39,X40)|~m1_subset_1(X41,k1_zfmisc_1(u1_struct_0(X39)))|~m1_subset_1(X40,u1_struct_0(X39))|(v2_struct_0(X39)|(~v2_pre_topc(X39)|~l1_pre_topc(X39))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_connsp_2])])])])])).
cnf(c_0_73, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,k1_tops_1(esk2_0,esk3_0)),k1_tops_1(esk2_0,u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_69]), c_0_36])])).
cnf(c_0_74, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,X1),k1_tops_1(esk2_0,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))|~r1_tarski(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_35]), c_0_36])])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk5_1(X1),X2)|X2!=k1_zfmisc_1(u1_struct_0(esk2_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_65, c_0_70])).
cnf(c_0_76, negated_conjecture, (r1_tarski(esk5_1(X1),esk3_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_29, c_0_71])).
cnf(c_0_77, plain, (r2_hidden(X3,k1_tops_1(X2,X1))|v2_struct_0(X2)|~m1_connsp_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_78, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,u1_struct_0(esk2_0)))|~r2_hidden(X1,k1_tops_1(esk2_0,k1_tops_1(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_38, c_0_73])).
cnf(c_0_79, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,esk5_1(X1)),k1_tops_1(esk2_0,esk3_0))|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
cnf(c_0_80, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,esk5_1(X2)))|~m1_connsp_2(esk5_1(X2),esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk2_0))|~r2_hidden(X2,esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_63]), c_0_36]), c_0_43])]), c_0_44]), c_0_53]), c_0_54])).
cnf(c_0_81, negated_conjecture, (m1_connsp_2(esk5_1(X1),esk2_0,X1)|v3_pre_topc(esk3_0,esk2_0)|~r2_hidden(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_82, plain, ![X42, X43]:(~v2_pre_topc(X42)|~l1_pre_topc(X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X42)))|v3_pre_topc(k1_tops_1(X42,X43),X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc9_tops_1])])).
cnf(c_0_83, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk2_0))|~r2_hidden(X1,k1_tops_1(esk2_0,k1_tops_1(esk2_0,esk3_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_78]), c_0_36])])).
cnf(c_0_84, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))|~r2_hidden(X1,k1_tops_1(esk2_0,esk5_1(X2)))|~r2_hidden(X2,esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_79])).
cnf(c_0_85, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,esk5_1(X1)))|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_53]), c_0_54])).
cnf(c_0_86, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_87, plain, (v3_pre_topc(k1_tops_1(X1,X2),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_88, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk2_0))|~r2_hidden(esk6_2(X1,u1_struct_0(esk2_0)),k1_tops_1(esk2_0,k1_tops_1(esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_50, c_0_83])).
cnf(c_0_89, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,esk3_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_35]), c_0_36])])).
cnf(c_0_90, negated_conjecture, (r2_hidden(X1,k1_tops_1(esk2_0,esk3_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_91, plain, (X1=X2|r2_hidden(esk6_2(X2,X1),X2)|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_86, c_0_56])).
cnf(c_0_92, plain, (v3_pre_topc(k1_tops_1(X1,X2),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~r1_tarski(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_87, c_0_19])).
cnf(c_0_93, negated_conjecture, (r1_tarski(k1_tops_1(esk2_0,k1_tops_1(esk2_0,esk3_0)),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_88, c_0_56])).
cnf(c_0_94, negated_conjecture, (k1_tops_1(esk2_0,esk3_0)=esk3_0|~r1_tarski(esk3_0,k1_tops_1(esk2_0,esk3_0))), inference(spm,[status(thm)],[c_0_86, c_0_89])).
cnf(c_0_95, negated_conjecture, (r1_tarski(X1,k1_tops_1(esk2_0,esk3_0))|~r2_hidden(esk6_2(X1,k1_tops_1(esk2_0,esk3_0)),esk3_0)), inference(spm,[status(thm)],[c_0_50, c_0_90])).
cnf(c_0_96, negated_conjecture, (k1_tops_1(esk2_0,esk3_0)=esk3_0|r2_hidden(esk6_2(esk3_0,k1_tops_1(esk2_0,esk3_0)),esk3_0)), inference(spm,[status(thm)],[c_0_91, c_0_89])).
cnf(c_0_97, negated_conjecture, (v3_pre_topc(k1_tops_1(esk2_0,k1_tops_1(esk2_0,k1_tops_1(esk2_0,esk3_0))),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_36]), c_0_43])])).
cnf(c_0_98, negated_conjecture, (k1_tops_1(esk2_0,esk3_0)=esk3_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])).
cnf(c_0_99, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_98]), c_0_98]), c_0_98]), c_0_53]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 100
# Proof object clause steps            : 77
# Proof object formula steps           : 23
# Proof object conjectures             : 51
# Proof object clause conjectures      : 48
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 50
# Proof object simplifying inferences  : 43
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 3284
# ...of these trivial                  : 12
# ...subsumed                          : 1380
# ...remaining for further processing  : 1892
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 74
# Backward-rewritten                   : 598
# Generated clauses                    : 40816
# ...of the previous two non-trivial   : 40104
# Contextual simplify-reflections      : 43
# Paramodulations                      : 40569
# Factorizations                       : 26
# Equation resolutions                 : 221
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1190
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 1179
# Current number of unprocessed clauses: 36769
# ...number of literals in the above   : 194684
# Current number of archived formulas  : 0
# Current number of archived clauses   : 700
# Clause-clause subsumption calls (NU) : 261724
# Rec. Clause-clause subsumption calls : 100232
# Non-unit clause-clause subsumptions  : 1398
# Unit Clause-clause subsumption calls : 81
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 766332

# -------------------------------------------------
# User time                : 0.606 s
# System time              : 0.007 s
# Total time               : 0.613 s
# Maximum resident set size: 32224 pages
