# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t25_zfmisc_1, conjecture, ![X1, X2, X3]:~((r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))&(~(X1=X2)&~(X1=X3)))), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t25_zfmisc_1)).
fof(t98_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1)), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t98_xboole_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t100_xboole_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t75_enumset1)).
fof(t68_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8)), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t68_enumset1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t28_xboole_1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t92_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t5_boole)).
fof(t119_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X4,X2,X1), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t119_enumset1)).
fof(t103_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X2,X4,X3), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t103_enumset1)).
fof(t54_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_tarski(X5,X6)), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t54_enumset1)).
fof(l75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8]:k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8)), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', l75_enumset1)).
fof(t107_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t107_enumset1)).
fof(t125_enumset1, axiom, ![X1, X2, X3, X4]:k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t125_enumset1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', commutativity_k3_xboole_0)).
fof(t131_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_enumset1(X6,X7,X8,X9)), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t131_enumset1)).
fof(t132_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7, X8, X9]:k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9)), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t132_enumset1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t4_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', d1_xboole_0)).
fof(fc3_xboole_0, axiom, ![X1, X2]:~(v1_xboole_0(k2_tarski(X1,X2))), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', fc3_xboole_0)).
fof(l97_xboole_1, axiom, ![X1, X2]:r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2)), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', l97_xboole_1)).
fof(t136_enumset1, axiom, ![X1, X2, X3]:~((~(X1=X2)&(~(X1=X3)&~(k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3))))), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', t136_enumset1)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/d862b2ca15da4bd7bacd18097d43c936.E_input', commutativity_k5_xboole_0)).
fof(c_0_29, negated_conjecture, ~(![X1, X2, X3]:~((r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))&(~(X1=X2)&~(X1=X3))))), inference(assume_negation,[status(cth)],[t25_zfmisc_1])).
fof(c_0_30, plain, ![X49, X50]:k2_xboole_0(X49,X50)=k5_xboole_0(X49,k4_xboole_0(X50,X49)), inference(variable_rename,[status(thm)],[t98_xboole_1])).
fof(c_0_31, plain, ![X117, X118]:k4_xboole_0(X117,X118)=k5_xboole_0(X117,k3_xboole_0(X117,X118)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_32, negated_conjecture, (r1_tarski(k1_tarski(esk1_0),k2_tarski(esk2_0,esk3_0))&(esk1_0!=esk2_0&esk1_0!=esk3_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_29])])])])).
fof(c_0_33, plain, ![X14]:k2_tarski(X14,X14)=k1_tarski(X14), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_34, plain, ![X124, X125]:k1_enumset1(X124,X124,X125)=k2_tarski(X124,X125), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_35, plain, ![X84, X85, X86]:k2_enumset1(X84,X84,X85,X86)=k1_enumset1(X84,X85,X86), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_36, plain, ![X93, X94, X95, X96]:k3_enumset1(X93,X93,X94,X95,X96)=k2_enumset1(X93,X94,X95,X96), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_37, plain, ![X119, X120, X121, X122, X123]:k4_enumset1(X119,X119,X120,X121,X122,X123)=k3_enumset1(X119,X120,X121,X122,X123), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_38, plain, ![X87, X88, X89, X90, X91, X92]:k5_enumset1(X87,X87,X88,X89,X90,X91,X92)=k4_enumset1(X87,X88,X89,X90,X91,X92), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_39, plain, ![X42, X43, X44, X45, X46, X47, X48]:k6_enumset1(X42,X42,X43,X44,X45,X46,X47,X48)=k5_enumset1(X42,X43,X44,X45,X46,X47,X48), inference(variable_rename,[status(thm)],[t75_enumset1])).
fof(c_0_40, plain, ![X68, X69, X70, X71, X72, X73, X74, X75]:k6_enumset1(X68,X69,X70,X71,X72,X73,X74,X75)=k2_xboole_0(k5_enumset1(X68,X69,X70,X71,X72,X73,X74),k1_tarski(X75)), inference(variable_rename,[status(thm)],[t68_enumset1])).
cnf(c_0_41, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_43, plain, ![X105, X106]:(~r1_tarski(X105,X106)|k3_xboole_0(X105,X106)=X105), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_44, negated_conjecture, (r1_tarski(k1_tarski(esk1_0),k2_tarski(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_45, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_49, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_50, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_51, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_52, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k5_enumset1(X1,X2,X3,X4,X5,X6,X7),k1_tarski(X8))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_53, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X2,X1)))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_54, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, negated_conjecture, (r1_tarski(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_45]), c_0_46]), c_0_46]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51])).
fof(c_0_56, plain, ![X13]:k5_xboole_0(X13,X13)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
fof(c_0_57, plain, ![X26]:k5_xboole_0(X26,k1_xboole_0)=X26, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_58, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k5_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k3_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_45]), c_0_46]), c_0_53]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_59, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_60, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_61, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_62, plain, ![X107, X108, X109, X110]:k2_enumset1(X107,X108,X109,X110)=k2_enumset1(X109,X110,X108,X107), inference(variable_rename,[status(thm)],[t119_enumset1])).
fof(c_0_63, plain, ![X27, X28, X29, X30]:k2_enumset1(X27,X28,X29,X30)=k2_enumset1(X27,X28,X30,X29), inference(variable_rename,[status(thm)],[t103_enumset1])).
fof(c_0_64, plain, ![X99, X100, X101, X102, X103, X104]:k4_enumset1(X99,X100,X101,X102,X103,X104)=k2_xboole_0(k2_enumset1(X99,X100,X101,X102),k2_tarski(X103,X104)), inference(variable_rename,[status(thm)],[t54_enumset1])).
fof(c_0_65, plain, ![X76, X77, X78, X79, X80, X81, X82, X83]:k6_enumset1(X76,X77,X78,X79,X80,X81,X82,X83)=k2_xboole_0(k2_enumset1(X76,X77,X78,X79),k2_enumset1(X80,X81,X82,X83)), inference(variable_rename,[status(thm)],[l75_enumset1])).
fof(c_0_66, plain, ![X55, X56, X57, X58]:k2_enumset1(X55,X56,X57,X58)=k2_enumset1(X55,X58,X57,X56), inference(variable_rename,[status(thm)],[t107_enumset1])).
cnf(c_0_67, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0,esk1_0)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60]), c_0_61])).
cnf(c_0_68, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X3,X4,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_69, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X2,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_70, plain, ![X51, X52, X53, X54]:k2_enumset1(X51,X52,X53,X54)=k2_enumset1(X54,X53,X52,X51), inference(variable_rename,[status(thm)],[t125_enumset1])).
cnf(c_0_71, plain, (k4_enumset1(X1,X2,X3,X4,X5,X6)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_tarski(X5,X6))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k2_xboole_0(k2_enumset1(X1,X2,X3,X4),k2_enumset1(X5,X6,X7,X8))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_73, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X1,X4,X3,X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0,esk1_0,X1)=k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_67]), c_0_58])).
cnf(c_0_75, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X3,X3,X3,X3,X3,X4,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51])).
cnf(c_0_76, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X2,X4,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51])).
cnf(c_0_77, plain, (k2_enumset1(X1,X2,X3,X4)=k2_enumset1(X4,X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_78, plain, (k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k5_xboole_0(k6_enumset1(X5,X5,X5,X5,X5,X5,X5,X6),k3_xboole_0(k6_enumset1(X5,X5,X5,X5,X5,X5,X5,X6),k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_46]), c_0_53]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_79, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4),k5_xboole_0(k6_enumset1(X5,X5,X5,X5,X5,X6,X7,X8),k3_xboole_0(k6_enumset1(X5,X5,X5,X5,X5,X6,X7,X8),k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_53]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_80, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X1,X1,X1,X1,X1,X4,X3,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51])).
cnf(c_0_81, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_75]), c_0_76])).
cnf(c_0_82, plain, (k6_enumset1(X1,X1,X1,X1,X1,X2,X3,X4)=k6_enumset1(X4,X4,X4,X4,X4,X3,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_51]), c_0_51])).
fof(c_0_83, plain, ![X31, X32]:k3_xboole_0(X31,X32)=k3_xboole_0(X32,X31), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_84, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X5,X6)=k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)), inference(rw,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_85, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_76])).
cnf(c_0_86, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[c_0_67, c_0_82])).
cnf(c_0_87, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_88, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,X1,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_89, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_86]), c_0_80])).
cnf(c_0_90, plain, (k5_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k5_xboole_0(k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8),k3_xboole_0(k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7),k6_enumset1(X8,X8,X8,X8,X8,X8,X8,X8))))=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(spm,[status(thm)],[c_0_58, c_0_87])).
cnf(c_0_91, negated_conjecture, (k6_enumset1(esk2_0,X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_84, c_0_88])).
fof(c_0_92, plain, ![X33, X34, X35, X36, X37, X38, X39, X40, X41]:k7_enumset1(X33,X34,X35,X36,X37,X38,X39,X40,X41)=k2_xboole_0(k3_enumset1(X33,X34,X35,X36,X37),k2_enumset1(X38,X39,X40,X41)), inference(variable_rename,[status(thm)],[t131_enumset1])).
cnf(c_0_93, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_89]), c_0_60]), c_0_61])).
cnf(c_0_94, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_90]), c_0_81])).
fof(c_0_95, plain, ![X59, X60, X61, X62, X63, X64, X65, X66, X67]:k7_enumset1(X59,X60,X61,X62,X63,X64,X65,X66,X67)=k2_xboole_0(k4_enumset1(X59,X60,X61,X62,X63,X64),k1_enumset1(X65,X66,X67)), inference(variable_rename,[status(thm)],[t132_enumset1])).
cnf(c_0_96, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k3_enumset1(X1,X2,X3,X4,X5),k2_enumset1(X6,X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_97, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_98, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k2_xboole_0(k4_enumset1(X1,X2,X3,X4,X5,X6),k1_enumset1(X7,X8,X9))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_99, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5),k5_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X7,X8,X9),k3_xboole_0(k6_enumset1(X6,X6,X6,X6,X6,X7,X8,X9),k6_enumset1(X1,X1,X1,X1,X2,X3,X4,X5))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_96, c_0_53]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_100, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_97]), c_0_90]), c_0_94])).
cnf(c_0_101, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X7,X8,X9)=k5_xboole_0(k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6),k5_xboole_0(k6_enumset1(X7,X7,X7,X7,X7,X7,X8,X9),k3_xboole_0(k6_enumset1(X7,X7,X7,X7,X7,X7,X8,X9),k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_53]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_102, plain, (k7_enumset1(X1,X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[c_0_79, c_0_99])).
cnf(c_0_103, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0,X1,X2)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_81]), c_0_90])).
cnf(c_0_104, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_81])).
cnf(c_0_105, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_84, c_0_100])).
cnf(c_0_106, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X6,X6,X6)=k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)), inference(spm,[status(thm)],[c_0_58, c_0_99])).
cnf(c_0_107, plain, (k7_enumset1(X1,X2,X3,X4,X5,X6,X6,X7,X8)=k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_99]), c_0_102])).
cnf(c_0_108, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1,X2,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X2,esk2_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_80, c_0_103])).
cnf(c_0_109, negated_conjecture, (k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,X1,esk3_0,X2)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_104]), c_0_90])).
cnf(c_0_110, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1,esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_105])).
cnf(c_0_111, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X5,X5)=k6_enumset1(X4,X4,X4,X4,X3,X2,X1,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_82]), c_0_99]), c_0_102])).
cnf(c_0_112, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X6,X6)=k6_enumset1(X1,X1,X1,X2,X3,X4,X5,X6)), inference(rw,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_113, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_104, c_0_108])).
cnf(c_0_114, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_80, c_0_100])).
cnf(c_0_115, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,X2)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X2,esk2_0,esk3_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_109]), c_0_108])).
cnf(c_0_116, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X4,X3,X2,X1,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_82]), c_0_99]), c_0_102]), c_0_102])).
cnf(c_0_117, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(X1,esk2_0,esk1_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_118, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,X1)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_100, c_0_112])).
cnf(c_0_119, negated_conjecture, (k6_enumset1(esk1_0,X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,X1)), inference(spm,[status(thm)],[c_0_84, c_0_105])).
cnf(c_0_120, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_113]), c_0_114])).
cnf(c_0_121, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1,esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_115])).
cnf(c_0_122, negated_conjecture, (k6_enumset1(X1,esk2_0,esk1_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_123, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk3_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_105, c_0_118])).
cnf(c_0_124, negated_conjecture, (k6_enumset1(esk1_0,X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_119, c_0_115])).
cnf(c_0_125, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X6,X7)=k6_enumset1(X4,X4,X3,X1,X2,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_75]), c_0_99]), c_0_102]), c_0_102])).
cnf(c_0_126, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk2_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_115, c_0_120])).
cnf(c_0_127, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X6,X7)=k6_enumset1(X4,X4,X3,X2,X1,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_82]), c_0_99]), c_0_102]), c_0_102])).
cnf(c_0_128, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk1_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_111]), c_0_122])).
cnf(c_0_129, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,esk3_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_123, c_0_116])).
cnf(c_0_130, negated_conjecture, (k6_enumset1(esk1_0,X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1,esk3_0)), inference(spm,[status(thm)],[c_0_84, c_0_110])).
cnf(c_0_131, negated_conjecture, (k6_enumset1(esk1_0,X1,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_132, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X6,X7)=k6_enumset1(X1,X1,X2,X4,X3,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_76]), c_0_99]), c_0_102]), c_0_102])).
cnf(c_0_133, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_113, c_0_126])).
cnf(c_0_134, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_127, c_0_128])).
cnf(c_0_135, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0,esk3_0,esk3_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_116, c_0_129])).
cnf(c_0_136, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X6,X6)=k6_enumset1(X1,X2,X3,X4,X4,X4,X5,X6)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_84]), c_0_99]), c_0_107])).
fof(c_0_137, plain, ![X111, X112, X114, X115, X116]:((r1_xboole_0(X111,X112)|r2_hidden(esk5_2(X111,X112),k3_xboole_0(X111,X112)))&(~r2_hidden(X116,k3_xboole_0(X114,X115))|~r1_xboole_0(X114,X115))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_138, plain, ![X20, X21, X22]:((~v1_xboole_0(X20)|~r2_hidden(X21,X20))&(r2_hidden(esk4_1(X22),X22)|v1_xboole_0(X22))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_139, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1,esk3_0)), inference(rw,[status(thm)],[c_0_130, c_0_131])).
cnf(c_0_140, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X5,X5)=k6_enumset1(X1,X1,X1,X1,X4,X3,X2,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_80]), c_0_99]), c_0_102])).
cnf(c_0_141, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X3,X4,X2,X1,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_75]), c_0_99]), c_0_102]), c_0_102])).
cnf(c_0_142, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0,esk1_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,X1,esk1_0,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_131]), c_0_132])).
cnf(c_0_143, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X5,X5)=k6_enumset1(X4,X4,X4,X4,X3,X1,X2,X5)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_75]), c_0_99]), c_0_102])).
cnf(c_0_144, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_133, c_0_134])).
cnf(c_0_145, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_146, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_147, plain, (r2_hidden(esk4_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_138])).
fof(c_0_148, plain, ![X15, X16]:~v1_xboole_0(k2_tarski(X15,X16)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
fof(c_0_149, plain, ![X97, X98]:r1_xboole_0(k3_xboole_0(X97,X98),k5_xboole_0(X97,X98)), inference(variable_rename,[status(thm)],[l97_xboole_1])).
cnf(c_0_150, negated_conjecture, (k6_enumset1(esk1_0,X1,esk2_0,esk1_0,esk3_0,esk3_0,esk3_0,esk3_0)=k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_151, plain, (k6_enumset1(X1,X2,X3,X4,X5,X6,X7,X8)=k6_enumset1(X1,X2,X4,X3,X5,X6,X7,X8)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_76]), c_0_99]), c_0_102]), c_0_102])).
cnf(c_0_152, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0,esk1_0)=k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_153, negated_conjecture, (k6_enumset1(X1,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_122]), c_0_145])).
cnf(c_0_154, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_155, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_156, plain, (~r1_xboole_0(X1,X2)|~r2_hidden(X3,k3_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_146, c_0_87])).
cnf(c_0_157, plain, (r1_xboole_0(X1,X2)|r2_hidden(esk5_2(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_158, plain, (r1_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_149])).
fof(c_0_159, plain, ![X17, X18, X19]:(X17=X18|(X17=X19|k4_xboole_0(k1_enumset1(X17,X18,X19),k1_tarski(X17))=k2_tarski(X18,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_enumset1])])])).
cnf(c_0_160, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk2_0,esk3_0,esk3_0,esk3_0,esk3_0)=k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0,esk1_0)), inference(rw,[status(thm)],[c_0_150, c_0_151])).
cnf(c_0_161, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_134]), c_0_152])).
cnf(c_0_162, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_127, c_0_153])).
cnf(c_0_163, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_154, c_0_87])).
cnf(c_0_164, plain, (~v1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_155, c_0_46]), c_0_47]), c_0_48]), c_0_49]), c_0_50]), c_0_51])).
cnf(c_0_165, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_166, negated_conjecture, (r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_158, c_0_59])).
cnf(c_0_167, plain, (X1=X2|X1=X3|k4_xboole_0(k1_enumset1(X1,X2,X3),k1_tarski(X1))=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_159])).
cnf(c_0_168, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk2_0,esk3_0,esk3_0,esk3_0,esk3_0)=k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_160, c_0_152])).
cnf(c_0_169, plain, (k6_enumset1(X1,X2,X3,X4,X5,X5,X6,X7)=k6_enumset1(X1,X1,X4,X3,X2,X5,X6,X7)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_80]), c_0_99]), c_0_102]), c_0_102])).
cnf(c_0_170, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,X1,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_161, c_0_162])).
cnf(c_0_171, negated_conjecture, (~r1_xboole_0(k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_59]), c_0_164])).
cnf(c_0_172, negated_conjecture, (r1_xboole_0(k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk2_0,esk3_0)),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_173, plain, (X1=X3|X1=X2|k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1)))=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_167, c_0_45]), c_0_46]), c_0_46]), c_0_42]), c_0_47]), c_0_47]), c_0_47]), c_0_47]), c_0_48]), c_0_48]), c_0_48]), c_0_48]), c_0_49]), c_0_49]), c_0_49]), c_0_49]), c_0_50]), c_0_50]), c_0_50]), c_0_50]), c_0_51]), c_0_51]), c_0_51]), c_0_51])).
cnf(c_0_174, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,X1,esk3_0)), inference(spm,[status(thm)],[c_0_76, c_0_100])).
cnf(c_0_175, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk2_0,esk2_0,esk2_0,esk3_0,esk3_0)=k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_168, c_0_136])).
cnf(c_0_176, negated_conjecture, (k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)=k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_169, c_0_170])).
fof(c_0_177, plain, ![X24, X25]:k5_xboole_0(X24,X25)=k5_xboole_0(X25,X24), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
cnf(c_0_178, negated_conjecture, (~r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk3_0,esk2_0,esk2_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[c_0_171, c_0_86])).
cnf(c_0_179, negated_conjecture, (r1_xboole_0(k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0)),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_172, c_0_86]), c_0_80])).
cnf(c_0_180, plain, (k5_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3),k3_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X1),k6_enumset1(X1,X1,X1,X1,X1,X1,X2,X3)))=k6_enumset1(X2,X2,X2,X2,X2,X2,X2,X3)|X1=X2|X1=X3), inference(spm,[status(thm)],[c_0_173, c_0_87])).
cnf(c_0_181, negated_conjecture, (k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,X1,esk2_0,esk3_0)=k6_enumset1(esk1_0,X1,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_140]), c_0_151]), c_0_136]), c_0_175]), c_0_176])).
cnf(c_0_182, negated_conjecture, (k3_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)), inference(rw,[status(thm)],[c_0_89, c_0_97])).
cnf(c_0_183, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_177])).
cnf(c_0_184, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_185, negated_conjecture, (esk1_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_186, negated_conjecture, (~r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0,esk3_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[c_0_178, c_0_80])).
cnf(c_0_187, negated_conjecture, (r1_xboole_0(k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[c_0_179, c_0_97])).
cnf(c_0_188, negated_conjecture, (k5_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0))=k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_182]), c_0_183]), c_0_85]), c_0_144]), c_0_184]), c_0_185])).
cnf(c_0_189, negated_conjecture, (~r1_xboole_0(k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk3_0),k6_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[c_0_186, c_0_97])).
cnf(c_0_190, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_187, c_0_188]), c_0_189]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 191
# Proof object clause steps            : 132
# Proof object formula steps           : 59
# Proof object conjectures             : 72
# Proof object clause conjectures      : 69
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 29
# Proof object generating inferences   : 64
# Proof object simplifying inferences  : 247
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 33
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 24
# Processed clauses                    : 3546
# ...of these trivial                  : 261
# ...subsumed                          : 2579
# ...remaining for further processing  : 706
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 291
# Generated clauses                    : 43371
# ...of the previous two non-trivial   : 38036
# Contextual simplify-reflections      : 0
# Paramodulations                      : 43371
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 391
#    Positive orientable unit clauses  : 214
#    Positive unorientable unit clauses: 66
#    Negative unit clauses             : 75
#    Non-unit-clauses                  : 36
# Current number of unprocessed clauses: 32684
# ...number of literals in the above   : 54526
# Current number of archived formulas  : 0
# Current number of archived clauses   : 324
# Clause-clause subsumption calls (NU) : 381
# Rec. Clause-clause subsumption calls : 87
# Non-unit clause-clause subsumptions  : 56
# Unit Clause-clause subsumption calls : 7733
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 35841
# BW rewrite match successes           : 4369
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 702099

# -------------------------------------------------
# User time                : 0.312 s
# System time              : 0.017 s
# Total time               : 0.328 s
# Maximum resident set size: 26608 pages
