# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k5_waybel_0(X1,X3))))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', t48_waybel_0)).
fof(fraenkel_a_2_0_waybel_0, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l1_orders_2(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_0_waybel_0(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(X5,X3))))))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', fraenkel_a_2_0_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', t24_orders_2)).
fof(d15_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', d15_waybel_0)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', l3_subset_1)).
fof(dt_k5_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', dt_k5_waybel_0)).
fof(t17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', t17_waybel_0)).
fof(d19_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v12_waybel_0(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_hidden(X3,X2)&r1_orders_2(X1,X4,X3))=>r2_hidden(X4,X2))))))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', d19_waybel_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', d10_xboole_0)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', t3_subset)).
fof(d9_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X2)=>r1_orders_2(X1,X4,X3)))))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', d9_lattice3)).
fof(d21_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_hidden(X3,X2)&r2_lattice3(X1,X2,X3)))))), file('data/MPTP/proofs/d9655e732e4c469f90940ed12caf2a1c.E_input', d21_waybel_0)).
fof(c_0_12, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:((~(v1_xboole_0(X2))&(v1_waybel_0(X2,X1)&(v12_waybel_0(X2,X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))))=>(v14_waybel_0(X2,X1)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&X2=k5_waybel_0(X1,X3)))))), inference(assume_negation,[status(cth)],[t48_waybel_0])).
fof(c_0_13, plain, ![X23, X24, X25, X28, X29]:(((m1_subset_1(esk6_3(X23,X24,X25),u1_struct_0(X24))|~r2_hidden(X23,a_2_0_waybel_0(X24,X25))|(v2_struct_0(X24)|~l1_orders_2(X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))))&((X23=esk6_3(X23,X24,X25)|~r2_hidden(X23,a_2_0_waybel_0(X24,X25))|(v2_struct_0(X24)|~l1_orders_2(X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))))&((m1_subset_1(esk7_3(X23,X24,X25),u1_struct_0(X24))|~r2_hidden(X23,a_2_0_waybel_0(X24,X25))|(v2_struct_0(X24)|~l1_orders_2(X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))))&((r1_orders_2(X24,esk6_3(X23,X24,X25),esk7_3(X23,X24,X25))|~r2_hidden(X23,a_2_0_waybel_0(X24,X25))|(v2_struct_0(X24)|~l1_orders_2(X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))))&(r2_hidden(esk7_3(X23,X24,X25),X25)|~r2_hidden(X23,a_2_0_waybel_0(X24,X25))|(v2_struct_0(X24)|~l1_orders_2(X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))))))))&(~m1_subset_1(X28,u1_struct_0(X24))|(X23!=X28|(~m1_subset_1(X29,u1_struct_0(X24))|(~r1_orders_2(X24,X28,X29)|~r2_hidden(X29,X25))))|r2_hidden(X23,a_2_0_waybel_0(X24,X25))|(v2_struct_0(X24)|~l1_orders_2(X24)|~m1_subset_1(X25,k1_zfmisc_1(u1_struct_0(X24)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_0_waybel_0])])])])])])).
fof(c_0_14, plain, ![X21, X22]:(v2_struct_0(X21)|(~v3_orders_2(X21)|~l1_orders_2(X21))|(~m1_subset_1(X22,u1_struct_0(X21))|r1_orders_2(X21,X22,X22))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_15, negated_conjecture, ![X8]:((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&l1_orders_2(esk1_0))))&((~v1_xboole_0(esk2_0)&(v1_waybel_0(esk2_0,esk1_0)&(v12_waybel_0(esk2_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))))&((~v14_waybel_0(esk2_0,esk1_0)|(~m1_subset_1(X8,u1_struct_0(esk1_0))|esk2_0!=k5_waybel_0(esk1_0,X8)))&((m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v14_waybel_0(esk2_0,esk1_0))&(esk2_0=k5_waybel_0(esk1_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_12])])])])])])).
cnf(c_0_16, plain, (r2_hidden(X3,a_2_0_waybel_0(X2,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~m1_subset_1(X4,u1_struct_0(X2))|~r1_orders_2(X2,X1,X4)|~r2_hidden(X4,X5)|~l1_orders_2(X2)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, plain, ![X41, X42, X43, X44, X46, X48]:((((m1_subset_1(esk10_4(X41,X42,X43,X44),u1_struct_0(X41))|~r2_hidden(X44,X43)|~m1_subset_1(X44,u1_struct_0(X41))|X43!=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41))&((r1_orders_2(X41,X44,esk10_4(X41,X42,X43,X44))|~r2_hidden(X44,X43)|~m1_subset_1(X44,u1_struct_0(X41))|X43!=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41))&(r2_hidden(esk10_4(X41,X42,X43,X44),X42)|~r2_hidden(X44,X43)|~m1_subset_1(X44,u1_struct_0(X41))|X43!=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41))))&(~m1_subset_1(X46,u1_struct_0(X41))|(~r1_orders_2(X41,X44,X46)|~r2_hidden(X46,X42))|r2_hidden(X44,X43)|~m1_subset_1(X44,u1_struct_0(X41))|X43!=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41)))&((m1_subset_1(esk11_3(X41,X42,X43),u1_struct_0(X41))|X43=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41))&((~r2_hidden(esk11_3(X41,X42,X43),X43)|(~m1_subset_1(X48,u1_struct_0(X41))|(~r1_orders_2(X41,esk11_3(X41,X42,X43),X48)|~r2_hidden(X48,X42)))|X43=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41))&((m1_subset_1(esk12_3(X41,X42,X43),u1_struct_0(X41))|r2_hidden(esk11_3(X41,X42,X43),X43)|X43=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41))&((r1_orders_2(X41,esk11_3(X41,X42,X43),esk12_3(X41,X42,X43))|r2_hidden(esk11_3(X41,X42,X43),X43)|X43=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41))&(r2_hidden(esk12_3(X41,X42,X43),X42)|r2_hidden(esk11_3(X41,X42,X43),X43)|X43=k3_waybel_0(X41,X42)|~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X41)))|~m1_subset_1(X42,k1_zfmisc_1(u1_struct_0(X41)))|~l1_orders_2(X41))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])).
fof(c_0_18, plain, ![X52, X53, X54]:(~m1_subset_1(X53,k1_zfmisc_1(X52))|(~r2_hidden(X54,X53)|r2_hidden(X54,X52))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_19, plain, ![X35, X36]:(v2_struct_0(X35)|~l1_orders_2(X35)|~m1_subset_1(X36,u1_struct_0(X35))|m1_subset_1(k5_waybel_0(X35,X36),k1_zfmisc_1(u1_struct_0(X35)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
fof(c_0_20, plain, ![X18, X19, X20]:((~r2_hidden(X20,k5_waybel_0(X18,X19))|r1_orders_2(X18,X20,X19)|~m1_subset_1(X20,u1_struct_0(X18))|~m1_subset_1(X19,u1_struct_0(X18))|(v2_struct_0(X18)|~l1_orders_2(X18)))&(~r1_orders_2(X18,X20,X19)|r2_hidden(X20,k5_waybel_0(X18,X19))|~m1_subset_1(X20,u1_struct_0(X18))|~m1_subset_1(X19,u1_struct_0(X18))|(v2_struct_0(X18)|~l1_orders_2(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_waybel_0])])])])])).
cnf(c_0_21, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_25, plain, ![X10, X11, X12, X13]:((~v12_waybel_0(X11,X10)|(~m1_subset_1(X12,u1_struct_0(X10))|(~m1_subset_1(X13,u1_struct_0(X10))|(~r2_hidden(X12,X11)|~r1_orders_2(X10,X13,X12)|r2_hidden(X13,X11))))|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_orders_2(X10))&((m1_subset_1(esk4_2(X10,X11),u1_struct_0(X10))|v12_waybel_0(X11,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_orders_2(X10))&((m1_subset_1(esk5_2(X10,X11),u1_struct_0(X10))|v12_waybel_0(X11,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_orders_2(X10))&(((r2_hidden(esk4_2(X10,X11),X11)|v12_waybel_0(X11,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_orders_2(X10))&(r1_orders_2(X10,esk5_2(X10,X11),esk4_2(X10,X11))|v12_waybel_0(X11,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_orders_2(X10)))&(~r2_hidden(esk5_2(X10,X11),X11)|v12_waybel_0(X11,X10)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|~l1_orders_2(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d19_waybel_0])])])])])).
cnf(c_0_26, plain, (r2_hidden(X1,a_2_0_waybel_0(X2,X3))|v2_struct_0(X2)|~r1_orders_2(X2,X1,X4)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_27, plain, (r1_orders_2(X1,esk11_3(X1,X2,X3),esk12_3(X1,X2,X3))|r2_hidden(esk11_3(X1,X2,X3),X3)|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (m1_subset_1(esk11_3(X1,X2,X3),u1_struct_0(X1))|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_29, plain, (m1_subset_1(esk12_3(X1,X2,X3),u1_struct_0(X1))|r2_hidden(esk11_3(X1,X2,X3),X3)|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, plain, (r1_orders_2(X1,esk6_3(X2,X1,X3),esk7_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~l1_orders_2(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_31, plain, (X1=esk6_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_32, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_33, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_34, plain, (v2_struct_0(X1)|m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_35, plain, (r2_hidden(X2,k5_waybel_0(X1,X3))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])]), c_0_24])).
fof(c_0_37, plain, ![X16, X17]:(((r1_tarski(X16,X17)|X16!=X17)&(r1_tarski(X17,X16)|X16!=X17))&(~r1_tarski(X16,X17)|~r1_tarski(X17,X16)|X16=X17)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_38, plain, (r2_hidden(X4,X1)|~v12_waybel_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r2_hidden(X3,X1)|~r1_orders_2(X2,X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_39, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_40, plain, (X1=k3_waybel_0(X2,X3)|r2_hidden(esk11_3(X2,X3,X1),a_2_0_waybel_0(X2,X4))|r2_hidden(esk11_3(X2,X3,X1),X1)|v2_struct_0(X2)|~r2_hidden(esk12_3(X2,X3,X1),X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29])).
cnf(c_0_41, plain, (r2_hidden(esk12_3(X1,X2,X3),X2)|r2_hidden(esk11_3(X1,X2,X3),X3)|X3=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_42, plain, (r1_orders_2(X1,X2,esk7_3(X2,X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,a_2_0_waybel_0(X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_43, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_32, c_0_31])).
cnf(c_0_44, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_45, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_23])]), c_0_24])).
fof(c_0_46, plain, ![X50, X51]:((~m1_subset_1(X50,k1_zfmisc_1(X51))|r1_tarski(X50,X51))&(~r1_tarski(X50,X51)|m1_subset_1(X50,k1_zfmisc_1(X51)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_47, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (r2_hidden(esk6_3(X1,X2,X3),X4)|v2_struct_0(X2)|~r2_hidden(esk7_3(X1,X2,X3),X4)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v12_waybel_0(X4,X2)|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_30]), c_0_39]), c_0_32])).
cnf(c_0_49, plain, (r2_hidden(esk7_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_50, plain, (X1=k3_waybel_0(X2,X3)|r2_hidden(esk11_3(X2,X3,X1),a_2_0_waybel_0(X2,X3))|r2_hidden(esk11_3(X2,X3,X1),X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_52, plain, (r2_hidden(X1,a_2_0_waybel_0(X2,X3))|v2_struct_0(X2)|~r2_hidden(esk7_3(X1,X2,X4),X3)|~r2_hidden(X1,a_2_0_waybel_0(X2,X4))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_42]), c_0_43]), c_0_39])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_23])]), c_0_24])).
cnf(c_0_54, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (r2_hidden(esk6_3(X1,X2,X3),X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v12_waybel_0(X3,X2)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_57, negated_conjecture, (X1=k3_waybel_0(esk1_0,esk2_0)|r2_hidden(esk11_3(esk1_0,esk2_0,X1),a_2_0_waybel_0(esk1_0,esk2_0))|r2_hidden(esk11_3(esk1_0,esk2_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_23])]), c_0_24])).
cnf(c_0_58, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(X2,u1_struct_0(esk1_0)))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_59, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_60, plain, (r2_hidden(esk6_3(X1,X2,X3),a_2_0_waybel_0(X2,X4))|v2_struct_0(X2)|~r2_hidden(esk7_3(X1,X2,X3),X4)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_30]), c_0_32]), c_0_39])).
cnf(c_0_61, plain, (r1_orders_2(X2,X1,X3)|v2_struct_0(X2)|~r2_hidden(X1,k5_waybel_0(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_62, plain, (X3=k3_waybel_0(X1,X2)|~r2_hidden(esk11_3(X1,X2,X3),X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r1_orders_2(X1,esk11_3(X1,X2,X3),X4)|~r2_hidden(X4,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_63, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~r2_hidden(X1,a_2_0_waybel_0(X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~v12_waybel_0(X2,X3)|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_56, c_0_31])).
cnf(c_0_64, negated_conjecture, (k3_waybel_0(esk1_0,esk2_0)=esk2_0|r2_hidden(esk11_3(esk1_0,esk2_0,esk2_0),a_2_0_waybel_0(esk1_0,esk2_0))|r2_hidden(esk11_3(esk1_0,esk2_0,esk2_0),esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_51])).
cnf(c_0_65, negated_conjecture, (v12_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_66, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_39]), c_0_59]), c_0_23])]), c_0_24])).
cnf(c_0_67, plain, (r2_hidden(esk6_3(X1,X2,X3),a_2_0_waybel_0(X2,X3))|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_60, c_0_49])).
fof(c_0_68, plain, ![X30, X31, X32, X33]:((~r2_lattice3(X30,X31,X32)|(~m1_subset_1(X33,u1_struct_0(X30))|(~r2_hidden(X33,X31)|r1_orders_2(X30,X33,X32)))|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&((m1_subset_1(esk8_3(X30,X31,X32),u1_struct_0(X30))|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&((r2_hidden(esk8_3(X30,X31,X32),X31)|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))&(~r1_orders_2(X30,esk8_3(X30,X31,X32),X32)|r2_lattice3(X30,X31,X32)|~m1_subset_1(X32,u1_struct_0(X30))|~l1_orders_2(X30))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_lattice3])])])])])).
cnf(c_0_69, negated_conjecture, (esk2_0=k5_waybel_0(esk1_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))|v14_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_71, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k3_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_72, plain, (r1_orders_2(X1,esk7_3(X2,X3,k5_waybel_0(X1,X4)),X4)|v2_struct_0(X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_0_waybel_0(X3,k5_waybel_0(X1,X4)))|~m1_subset_1(esk7_3(X2,X3,k5_waybel_0(X1,X4)),u1_struct_0(X1))|~m1_subset_1(k5_waybel_0(X1,X4),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~l1_orders_2(X3)), inference(spm,[status(thm)],[c_0_61, c_0_49])).
cnf(c_0_73, negated_conjecture, (X1=k3_waybel_0(esk1_0,X2)|~r2_hidden(esk11_3(esk1_0,X2,X1),X1)|~r2_hidden(esk11_3(esk1_0,X2,X1),X2)|~m1_subset_1(esk11_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_36]), c_0_23])])).
cnf(c_0_74, negated_conjecture, (k3_waybel_0(esk1_0,esk2_0)=esk2_0|r2_hidden(esk11_3(esk1_0,esk2_0,esk2_0),esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_51]), c_0_65]), c_0_23])]), c_0_24])).
cnf(c_0_75, negated_conjecture, (r2_hidden(esk6_3(X1,esk1_0,X2),a_2_0_waybel_0(esk1_0,u1_struct_0(esk1_0)))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_23])]), c_0_24])).
cnf(c_0_76, plain, (r2_lattice3(X1,X2,X3)|~r1_orders_2(X1,esk8_3(X1,X2,X3),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_77, negated_conjecture, (r1_orders_2(esk1_0,X1,esk3_0)|v14_waybel_0(esk2_0,esk1_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_69]), c_0_23])]), c_0_24]), c_0_70])).
cnf(c_0_78, negated_conjecture, (r2_hidden(X1,esk2_0)|k3_waybel_0(esk1_0,X2)!=esk2_0|~r1_orders_2(esk1_0,X1,X3)|~r2_hidden(X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_51]), c_0_23])])).
cnf(c_0_79, plain, (r2_hidden(esk7_3(X1,X2,k5_waybel_0(X3,X4)),X5)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,k5_waybel_0(X3,X4)))|~r2_hidden(X4,X5)|~m1_subset_1(esk7_3(X1,X2,k5_waybel_0(X3,X4)),u1_struct_0(X3))|~m1_subset_1(k5_waybel_0(X3,X4),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X3))|~v12_waybel_0(X5,X3)|~l1_orders_2(X3)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_38, c_0_72])).
cnf(c_0_80, negated_conjecture, (k3_waybel_0(esk1_0,esk2_0)=esk2_0|~m1_subset_1(esk11_3(esk1_0,esk2_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_51])]), c_0_74])).
cnf(c_0_81, negated_conjecture, (m1_subset_1(esk6_3(X1,esk1_0,X2),u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_75]), c_0_59]), c_0_23])]), c_0_24])).
cnf(c_0_82, negated_conjecture, (r2_lattice3(esk1_0,X1,esk3_0)|v14_waybel_0(esk2_0,esk1_0)|~r2_hidden(esk8_3(esk1_0,X1,esk3_0),esk2_0)|~m1_subset_1(esk8_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_23])]), c_0_70])).
cnf(c_0_83, plain, (r2_hidden(esk8_3(X1,X2,X3),X2)|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_84, plain, ![X37, X38, X40]:(((m1_subset_1(esk9_2(X37,X38),u1_struct_0(X37))|~v14_waybel_0(X38,X37)|(v1_xboole_0(X38)|(~v1_waybel_0(X38,X37)|(~v12_waybel_0(X38,X37)|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37))))))|(v2_struct_0(X37)|(~v3_orders_2(X37)|(~v4_orders_2(X37)|~l1_orders_2(X37)))))&((r2_hidden(esk9_2(X37,X38),X38)|~v14_waybel_0(X38,X37)|(v1_xboole_0(X38)|(~v1_waybel_0(X38,X37)|(~v12_waybel_0(X38,X37)|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37))))))|(v2_struct_0(X37)|(~v3_orders_2(X37)|(~v4_orders_2(X37)|~l1_orders_2(X37)))))&(r2_lattice3(X37,X38,esk9_2(X37,X38))|~v14_waybel_0(X38,X37)|(v1_xboole_0(X38)|(~v1_waybel_0(X38,X37)|(~v12_waybel_0(X38,X37)|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37))))))|(v2_struct_0(X37)|(~v3_orders_2(X37)|(~v4_orders_2(X37)|~l1_orders_2(X37)))))))&(~m1_subset_1(X40,u1_struct_0(X37))|(~r2_hidden(X40,X38)|~r2_lattice3(X37,X38,X40))|v14_waybel_0(X38,X37)|(v1_xboole_0(X38)|(~v1_waybel_0(X38,X37)|(~v12_waybel_0(X38,X37)|~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37))))))|(v2_struct_0(X37)|(~v3_orders_2(X37)|(~v4_orders_2(X37)|~l1_orders_2(X37)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_waybel_0])])])])])])).
cnf(c_0_85, negated_conjecture, (r2_hidden(X1,esk2_0)|k3_waybel_0(esk1_0,X2)!=esk2_0|~r2_hidden(esk7_3(X1,esk1_0,X3),X2)|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X3))|~m1_subset_1(esk7_3(X1,esk1_0,X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_42]), c_0_23])]), c_0_24])).
cnf(c_0_86, negated_conjecture, (r2_hidden(esk7_3(X1,X2,k5_waybel_0(esk1_0,X3)),esk2_0)|v2_struct_0(X2)|~r2_hidden(X1,a_2_0_waybel_0(X2,k5_waybel_0(esk1_0,X3)))|~r2_hidden(X3,esk2_0)|~m1_subset_1(esk7_3(X1,X2,k5_waybel_0(esk1_0,X3)),u1_struct_0(esk1_0))|~m1_subset_1(k5_waybel_0(esk1_0,X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~l1_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_65]), c_0_51]), c_0_23])]), c_0_24])).
cnf(c_0_87, negated_conjecture, (k3_waybel_0(esk1_0,esk2_0)=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_28]), c_0_51]), c_0_23])])).
cnf(c_0_88, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_31]), c_0_23])]), c_0_24])).
cnf(c_0_89, plain, (r1_orders_2(X1,esk5_2(X1,X2),esk4_2(X1,X2))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_90, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_91, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_92, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)|~m1_subset_1(esk8_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_23])]), c_0_70])).
cnf(c_0_93, plain, (m1_subset_1(esk8_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_94, plain, (r1_orders_2(X1,X4,X3)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_95, plain, (r2_lattice3(X1,X2,esk9_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~v14_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_96, plain, (m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~v14_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_97, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,k5_waybel_0(esk1_0,X2)))|~r2_hidden(X2,esk2_0)|~m1_subset_1(esk7_3(X1,esk1_0,k5_waybel_0(esk1_0,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k5_waybel_0(esk1_0,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_87]), c_0_51]), c_0_23])]), c_0_24]), c_0_88])).
cnf(c_0_98, plain, (r2_hidden(esk5_2(X1,X2),a_2_0_waybel_0(X1,X3))|v12_waybel_0(X2,X1)|v2_struct_0(X1)|~r2_hidden(esk4_2(X1,X2),X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_89]), c_0_90]), c_0_91])).
cnf(c_0_99, plain, (X1=k3_waybel_0(X2,X3)|r2_hidden(esk11_3(X2,X3,X1),X1)|r2_hidden(esk11_3(X2,X3,X1),X4)|~r2_hidden(esk12_3(X2,X3,X1),X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v12_waybel_0(X4,X2)|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_27]), c_0_29]), c_0_28])).
cnf(c_0_100, plain, (v14_waybel_0(X3,X2)|v1_xboole_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,X3)|~r2_lattice3(X2,X3,X1)|~v1_waybel_0(X3,X2)|~v12_waybel_0(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)|~v4_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_101, negated_conjecture, (r2_lattice3(esk1_0,esk2_0,esk3_0)|v14_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_23])]), c_0_70])).
cnf(c_0_102, negated_conjecture, (v1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_103, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_104, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_105, negated_conjecture, (r2_hidden(esk3_0,esk2_0)|v14_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_69]), c_0_70])).
cnf(c_0_106, plain, (r1_orders_2(X1,X2,esk9_2(X1,X3))|v1_xboole_0(X3)|v2_struct_0(X1)|~r2_hidden(X2,X3)|~v14_waybel_0(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~v12_waybel_0(X3,X1)|~v1_waybel_0(X3,X1)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96])).
cnf(c_0_107, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,a_2_0_waybel_0(esk1_0,k5_waybel_0(esk1_0,X2)))|~r2_hidden(X2,esk2_0)|~m1_subset_1(k5_waybel_0(esk1_0,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_39]), c_0_23])]), c_0_24])).
cnf(c_0_108, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),a_2_0_waybel_0(esk1_0,u1_struct_0(esk1_0)))|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_98]), c_0_23])]), c_0_24])).
cnf(c_0_109, plain, (r2_hidden(esk4_2(X1,X2),X2)|v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_110, plain, (X1=k3_waybel_0(X2,X3)|r2_hidden(esk11_3(X2,X3,X1),X3)|r2_hidden(esk11_3(X2,X3,X1),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v12_waybel_0(X3,X2)|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_99, c_0_41])).
cnf(c_0_111, negated_conjecture, (~v14_waybel_0(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|esk2_0!=k5_waybel_0(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_112, negated_conjecture, (v14_waybel_0(esk2_0,esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_51]), c_0_65]), c_0_102]), c_0_23]), c_0_103]), c_0_22])]), c_0_104]), c_0_24]), c_0_70]), c_0_105])).
cnf(c_0_113, plain, (v12_waybel_0(X2,X1)|~r2_hidden(esk5_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_114, plain, (r2_hidden(X1,k5_waybel_0(X2,esk9_2(X2,X3)))|v1_xboole_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,X3)|~v14_waybel_0(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~v12_waybel_0(X3,X2)|~v1_waybel_0(X3,X2)|~l1_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_106]), c_0_96])).
cnf(c_0_115, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk2_0)|v12_waybel_0(X1,esk1_0)|~r2_hidden(esk4_2(esk1_0,X1),k5_waybel_0(esk1_0,X2))|~r2_hidden(X2,esk2_0)|~m1_subset_1(k5_waybel_0(esk1_0,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_98]), c_0_23])]), c_0_24])).
cnf(c_0_116, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),a_2_0_waybel_0(esk1_0,u1_struct_0(esk1_0)))|v12_waybel_0(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_23])])).
cnf(c_0_117, plain, (r1_orders_2(X1,X2,esk10_4(X1,X3,X4,X2))|~r2_hidden(X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k3_waybel_0(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_118, plain, (m1_subset_1(esk10_4(X1,X2,X3,X4),u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_119, negated_conjecture, (X1=esk2_0|r2_hidden(esk11_3(esk1_0,esk2_0,X1),esk2_0)|r2_hidden(esk11_3(esk1_0,esk2_0,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_51]), c_0_87]), c_0_65]), c_0_23])])).
cnf(c_0_120, negated_conjecture, (k5_waybel_0(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_121, plain, (r2_hidden(X1,X2)|v1_xboole_0(X3)|v2_struct_0(X4)|~r2_hidden(esk9_2(X4,X3),X2)|~r2_hidden(X1,X3)|~v14_waybel_0(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X1,u1_struct_0(X4))|~v12_waybel_0(X2,X4)|~v12_waybel_0(X3,X4)|~v1_waybel_0(X3,X4)|~l1_orders_2(X4)|~v4_orders_2(X4)|~v3_orders_2(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_106]), c_0_96])).
cnf(c_0_122, plain, (v12_waybel_0(k5_waybel_0(X1,esk9_2(X1,X2)),X3)|v1_xboole_0(X2)|v2_struct_0(X1)|~r2_hidden(esk5_2(X3,k5_waybel_0(X1,esk9_2(X1,X2))),X2)|~v14_waybel_0(X2,X1)|~m1_subset_1(esk5_2(X3,k5_waybel_0(X1,esk9_2(X1,X2))),u1_struct_0(X1))|~m1_subset_1(k5_waybel_0(X1,esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v12_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~l1_orders_2(X3)|~l1_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_123, negated_conjecture, (r2_hidden(esk5_2(esk1_0,k5_waybel_0(esk1_0,X1)),esk2_0)|v12_waybel_0(k5_waybel_0(esk1_0,X1),esk1_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(k5_waybel_0(esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_109]), c_0_23])])).
cnf(c_0_124, negated_conjecture, (m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|v12_waybel_0(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_116]), c_0_59]), c_0_23])]), c_0_24])).
cnf(c_0_125, plain, (r2_hidden(X1,a_2_0_waybel_0(X2,X3))|v2_struct_0(X2)|X4!=k3_waybel_0(X2,X5)|~r2_hidden(esk10_4(X2,X5,X4,X1),X3)|~r2_hidden(X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_117]), c_0_118])).
cnf(c_0_126, plain, (r2_hidden(esk10_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_127, negated_conjecture, (r2_hidden(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),k5_waybel_0(esk1_0,X1))|r2_hidden(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_34]), c_0_23])]), c_0_24]), c_0_120])).
cnf(c_0_128, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk1_0,esk9_2(X2,X3)))|v1_xboole_0(X3)|v2_struct_0(X2)|~r2_hidden(X1,X3)|~v14_waybel_0(X3,X2)|~m1_subset_1(k5_waybel_0(esk1_0,esk9_2(X2,X3)),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk9_2(X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~v12_waybel_0(k5_waybel_0(esk1_0,esk9_2(X2,X3)),X2)|~v12_waybel_0(X3,X2)|~v1_waybel_0(X3,X2)|~l1_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_121, c_0_45])).
cnf(c_0_129, negated_conjecture, (v12_waybel_0(k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0)),esk1_0)|~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_112]), c_0_51]), c_0_65]), c_0_102]), c_0_23]), c_0_103]), c_0_22])]), c_0_104]), c_0_24]), c_0_124])).
cnf(c_0_130, plain, (X1=k3_waybel_0(X2,X3)|v2_struct_0(X2)|~r2_hidden(esk7_3(esk11_3(X2,X3,X1),X2,X4),X3)|~r2_hidden(esk11_3(X2,X3,X1),a_2_0_waybel_0(X2,X4))|~r2_hidden(esk11_3(X2,X3,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_42]), c_0_39])).
cnf(c_0_131, plain, (r2_hidden(X1,a_2_0_waybel_0(X2,X3))|v2_struct_0(X2)|X4!=k3_waybel_0(X2,X3)|~r2_hidden(X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_125, c_0_126])).
cnf(c_0_132, negated_conjecture, (r1_orders_2(esk1_0,esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),X1)|r2_hidden(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),esk2_0)|~m1_subset_1(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_127]), c_0_23])]), c_0_24])).
cnf(c_0_133, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0)))|~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_112]), c_0_51]), c_0_65]), c_0_102]), c_0_23]), c_0_103]), c_0_22])]), c_0_104]), c_0_24])).
cnf(c_0_134, plain, (X1=k3_waybel_0(X2,X3)|v2_struct_0(X2)|~r2_hidden(esk11_3(X2,X3,X1),a_2_0_waybel_0(X2,X3))|~r2_hidden(esk11_3(X2,X3,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_130, c_0_49])).
cnf(c_0_135, negated_conjecture, (r2_hidden(X1,a_2_0_waybel_0(esk1_0,X2))|k3_waybel_0(esk1_0,X2)!=esk2_0|~r2_hidden(X1,esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_51]), c_0_23])]), c_0_24])).
cnf(c_0_136, negated_conjecture, (r2_hidden(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),esk2_0)|k3_waybel_0(esk1_0,X2)!=esk2_0|~r2_hidden(X1,X2)|~m1_subset_1(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_78, c_0_132])).
cnf(c_0_137, negated_conjecture, (r1_orders_2(esk1_0,X1,esk9_2(esk1_0,esk2_0))|~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_133]), c_0_23])]), c_0_24])).
cnf(c_0_138, negated_conjecture, (X1=k3_waybel_0(esk1_0,X2)|k3_waybel_0(esk1_0,X2)!=esk2_0|~r2_hidden(esk11_3(esk1_0,X2,X1),esk2_0)|~r2_hidden(esk11_3(esk1_0,X2,X1),X1)|~m1_subset_1(esk11_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_23])]), c_0_24])).
cnf(c_0_139, negated_conjecture, (r2_hidden(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),esk2_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_87]), c_0_51])])).
cnf(c_0_140, negated_conjecture, (r1_orders_2(esk1_0,X1,esk9_2(esk1_0,esk2_0))|~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_34]), c_0_23])]), c_0_24])).
cnf(c_0_141, negated_conjecture, (~r2_hidden(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),k5_waybel_0(esk1_0,X1))|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k5_waybel_0(esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_87]), c_0_87]), c_0_51])]), c_0_120])).
cnf(c_0_142, negated_conjecture, (r2_hidden(X1,k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0)))|~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_140]), c_0_23])]), c_0_24])).
cnf(c_0_143, negated_conjecture, (~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(esk11_3(esk1_0,esk2_0,k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0))),u1_struct_0(esk1_0))|~m1_subset_1(k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142]), c_0_139])).
cnf(c_0_144, negated_conjecture, (~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(k5_waybel_0(esk1_0,esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_28]), c_0_87]), c_0_51]), c_0_23])]), c_0_120])).
cnf(c_0_145, negated_conjecture, (~r2_hidden(esk9_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_34]), c_0_23])]), c_0_24])).
cnf(c_0_146, plain, (r2_hidden(esk9_2(X1,X2),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v14_waybel_0(X2,X1)|~v1_waybel_0(X2,X1)|~v12_waybel_0(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_147, negated_conjecture, (~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_112]), c_0_51]), c_0_65]), c_0_102]), c_0_23]), c_0_103]), c_0_22])]), c_0_104]), c_0_24])).
cnf(c_0_148, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_96]), c_0_112]), c_0_51]), c_0_65]), c_0_102]), c_0_23]), c_0_103]), c_0_22])]), c_0_104]), c_0_24]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 149
# Proof object clause steps            : 124
# Proof object formula steps           : 25
# Proof object conjectures             : 64
# Proof object clause conjectures      : 61
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 47
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 74
# Proof object simplifying inferences  : 200
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 50
# Processed clauses                    : 15299
# ...of these trivial                  : 460
# ...subsumed                          : 12035
# ...remaining for further processing  : 2804
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 634
# Backward-rewritten                   : 141
# Generated clauses                    : 55709
# ...of the previous two non-trivial   : 54453
# Contextual simplify-reflections      : 409
# Paramodulations                      : 55687
# Factorizations                       : 0
# Equation resolutions                 : 22
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1977
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 1957
# Current number of unprocessed clauses: 37388
# ...number of literals in the above   : 491172
# Current number of archived formulas  : 0
# Current number of archived clauses   : 824
# Clause-clause subsumption calls (NU) : 1673113
# Rec. Clause-clause subsumption calls : 57529
# Non-unit clause-clause subsumptions  : 8836
# Unit Clause-clause subsumption calls : 1452
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3085339

# -------------------------------------------------
# User time                : 2.652 s
# System time              : 0.033 s
# Total time               : 2.685 s
# Maximum resident set size: 54908 pages
