# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t76_mcart_1, conjecture, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(X2))=>![X6]:(m1_subset_1(X6,k1_zfmisc_1(X3))=>![X7]:(m1_subset_1(X7,k3_zfmisc_1(X1,X2,X3))=>(r2_hidden(X7,k3_zfmisc_1(X4,X5,X6))=>(r2_hidden(k5_mcart_1(X1,X2,X3,X7),X4)&(r2_hidden(k6_mcart_1(X1,X2,X3,X7),X5)&r2_hidden(k7_mcart_1(X1,X2,X3,X7),X6)))))))), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', t76_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', d3_zfmisc_1)).
fof(t73_mcart_1, axiom, ![X1, X2, X3, X4, X5, X6]:(r2_hidden(k3_mcart_1(X1,X2,X3),k3_zfmisc_1(X4,X5,X6))<=>(r2_hidden(X1,X4)&(r2_hidden(X2,X5)&r2_hidden(X3,X6)))), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', t73_mcart_1)).
fof(t10_mcart_1, axiom, ![X1, X2, X3]:(r2_hidden(X1,k2_zfmisc_1(X2,X3))=>(r2_hidden(k1_mcart_1(X1),X2)&r2_hidden(k2_mcart_1(X1),X3))), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', t10_mcart_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', d1_xboole_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', t2_subset)).
fof(t50_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&~(![X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))=>(k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))&(k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))&k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4))))))))), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', t50_mcart_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', cc1_subset_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/da00e21de3c24aa4a435b03e5a7627e3.E_input', t2_xboole_1)).
fof(c_0_10, negated_conjecture, ~(![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(X2))=>![X6]:(m1_subset_1(X6,k1_zfmisc_1(X3))=>![X7]:(m1_subset_1(X7,k3_zfmisc_1(X1,X2,X3))=>(r2_hidden(X7,k3_zfmisc_1(X4,X5,X6))=>(r2_hidden(k5_mcart_1(X1,X2,X3,X7),X4)&(r2_hidden(k6_mcart_1(X1,X2,X3,X7),X5)&r2_hidden(k7_mcart_1(X1,X2,X3,X7),X6))))))))), inference(assume_negation,[status(cth)],[t76_mcart_1])).
fof(c_0_11, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk5_0,k1_zfmisc_1(esk2_0))&(m1_subset_1(esk6_0,k1_zfmisc_1(esk3_0))&(m1_subset_1(esk7_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&(r2_hidden(esk7_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))&(~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk4_0)|(~r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk5_0)|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
fof(c_0_12, plain, ![X15, X16, X17]:k3_zfmisc_1(X15,X16,X17)=k2_zfmisc_1(k2_zfmisc_1(X15,X16),X17), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
fof(c_0_13, plain, ![X18, X19, X20, X21, X22, X23]:(((r2_hidden(X18,X21)|~r2_hidden(k3_mcart_1(X18,X19,X20),k3_zfmisc_1(X21,X22,X23)))&((r2_hidden(X19,X22)|~r2_hidden(k3_mcart_1(X18,X19,X20),k3_zfmisc_1(X21,X22,X23)))&(r2_hidden(X20,X23)|~r2_hidden(k3_mcart_1(X18,X19,X20),k3_zfmisc_1(X21,X22,X23)))))&(~r2_hidden(X18,X21)|(~r2_hidden(X19,X22)|~r2_hidden(X20,X23))|r2_hidden(k3_mcart_1(X18,X19,X20),k3_zfmisc_1(X21,X22,X23)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_mcart_1])])])).
fof(c_0_14, plain, ![X32, X33, X34]:((r2_hidden(k1_mcart_1(X32),X33)|~r2_hidden(X32,k2_zfmisc_1(X33,X34)))&(r2_hidden(k2_mcart_1(X32),X34)|~r2_hidden(X32,k2_zfmisc_1(X33,X34)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_mcart_1])])])).
cnf(c_0_15, negated_conjecture, (r2_hidden(esk7_0,k3_zfmisc_1(esk4_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_17, plain, ![X35, X36, X37]:((~v1_xboole_0(X35)|~r2_hidden(X36,X35))&(r2_hidden(esk8_1(X37),X37)|v1_xboole_0(X37))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_18, plain, (r2_hidden(k3_mcart_1(X1,X3,X5),k3_zfmisc_1(X2,X4,X6))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|~r2_hidden(X5,X6)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, ![X30, X31]:(~m1_subset_1(X30,X31)|(v1_xboole_0(X31)|r2_hidden(X30,X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(esk7_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21, plain, (r2_hidden(k2_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (r2_hidden(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk4_0,esk5_0),esk6_0))), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, plain, (r2_hidden(k1_mcart_1(X1),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_24, plain, ![X24, X25, X26, X27]:((k5_mcart_1(X24,X25,X26,X27)=k1_mcart_1(k1_mcart_1(X27))|~m1_subset_1(X27,k3_zfmisc_1(X24,X25,X26))|X26=k1_xboole_0|X25=k1_xboole_0|X24=k1_xboole_0)&((k6_mcart_1(X24,X25,X26,X27)=k2_mcart_1(k1_mcart_1(X27))|~m1_subset_1(X27,k3_zfmisc_1(X24,X25,X26))|X26=k1_xboole_0|X25=k1_xboole_0|X24=k1_xboole_0)&(k7_mcart_1(X24,X25,X26,X27)=k2_mcart_1(X27)|~m1_subset_1(X27,k3_zfmisc_1(X24,X25,X26))|X26=k1_xboole_0|X25=k1_xboole_0|X24=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t50_mcart_1])])])])])).
cnf(c_0_25, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (r2_hidden(k3_mcart_1(X1,X3,X5),k2_zfmisc_1(k2_zfmisc_1(X2,X4),X6))|~r2_hidden(X5,X6)|~r2_hidden(X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_18, c_0_16])).
cnf(c_0_27, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_20, c_0_16])).
fof(c_0_29, plain, ![X28, X29]:(~v1_xboole_0(X28)|(~m1_subset_1(X29,k1_zfmisc_1(X28))|v1_xboole_0(X29))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k2_mcart_1(esk7_0),esk6_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_31, negated_conjecture, (r2_hidden(k1_mcart_1(esk7_0),k2_zfmisc_1(esk4_0,esk5_0))), inference(spm,[status(thm)],[c_0_23, c_0_22])).
cnf(c_0_32, plain, (k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (~v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))|~r2_hidden(X4,X3)|~r2_hidden(X5,X2)|~r2_hidden(X6,X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34, negated_conjecture, (v1_xboole_0(k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|r2_hidden(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_36, negated_conjecture, (m1_subset_1(esk6_0,k1_zfmisc_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_37, negated_conjecture, (~v1_xboole_0(esk6_0)), inference(spm,[status(thm)],[c_0_25, c_0_30])).
cnf(c_0_38, negated_conjecture, (r2_hidden(k2_mcart_1(k1_mcart_1(esk7_0)),esk5_0)), inference(spm,[status(thm)],[c_0_21, c_0_31])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_40, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k6_mcart_1(X1,X2,X3,X4)=k2_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_32, c_0_16])).
cnf(c_0_41, plain, (k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|~r2_hidden(X1,esk3_0)|~r2_hidden(X2,esk2_0)|~r2_hidden(X3,esk1_0)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_43, plain, (r2_hidden(esk8_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_44, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37])).
cnf(c_0_45, negated_conjecture, (~v1_xboole_0(esk5_0)), inference(spm,[status(thm)],[c_0_25, c_0_38])).
cnf(c_0_46, negated_conjecture, (v1_xboole_0(esk5_0)|~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_39])).
cnf(c_0_47, negated_conjecture, (r2_hidden(k1_mcart_1(k1_mcart_1(esk7_0)),esk4_0)), inference(spm,[status(thm)],[c_0_23, c_0_31])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_49, negated_conjecture, (~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk4_0)|~r2_hidden(k6_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk5_0)|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_50, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|r2_hidden(k6_mcart_1(X1,X2,X3,esk7_0),esk5_0)|~m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_38, c_0_40])).
cnf(c_0_51, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=k1_mcart_1(k1_mcart_1(X4))|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_41, c_0_16])).
fof(c_0_52, plain, ![X40, X41]:(~r2_hidden(X40,X41)|~r1_tarski(X41,X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_53, plain, ![X39]:r1_tarski(k1_xboole_0,X39), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_54, plain, (r2_hidden(k2_mcart_1(k3_mcart_1(X1,X2,X3)),X4)|~r2_hidden(X3,X4)|~r2_hidden(X2,X5)|~r2_hidden(X1,X6)), inference(spm,[status(thm)],[c_0_21, c_0_26])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_56, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_57, negated_conjecture, (~v1_xboole_0(esk4_0)), inference(spm,[status(thm)],[c_0_25, c_0_47])).
cnf(c_0_58, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_35, c_0_48])).
cnf(c_0_59, negated_conjecture, (esk3_0=k1_xboole_0|esk2_0=k1_xboole_0|esk1_0=k1_xboole_0|~r2_hidden(k5_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk4_0)|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_28])])).
cnf(c_0_60, negated_conjecture, (X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|r2_hidden(k5_mcart_1(X1,X2,X3,esk7_0),esk4_0)|~m1_subset_1(esk7_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_47, c_0_51])).
cnf(c_0_61, plain, (k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_62, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_64, negated_conjecture, (r2_hidden(k2_mcart_1(k3_mcart_1(X1,esk7_0,X2)),X3)|~r2_hidden(X2,X3)|~r2_hidden(X1,X4)), inference(spm,[status(thm)],[c_0_54, c_0_22])).
cnf(c_0_65, negated_conjecture, (r2_hidden(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_43]), c_0_56])).
cnf(c_0_66, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_67, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0|esk3_0=k1_xboole_0|~r2_hidden(k7_mcart_1(esk1_0,esk2_0,esk3_0,esk7_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_60]), c_0_28])])).
cnf(c_0_68, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k7_mcart_1(X1,X2,X3,X4)=k2_mcart_1(X4)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_61, c_0_16])).
cnf(c_0_69, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_70, negated_conjecture, (r2_hidden(k2_mcart_1(k3_mcart_1(k2_mcart_1(esk7_0),esk7_0,X1)),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_64, c_0_30])).
cnf(c_0_71, negated_conjecture, (r2_hidden(esk7_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_43]), c_0_66])).
cnf(c_0_72, negated_conjecture, (esk3_0=k1_xboole_0|esk2_0=k1_xboole_0|esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_30]), c_0_28])])).
cnf(c_0_73, plain, (v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_69, c_0_43])).
cnf(c_0_74, negated_conjecture, (r2_hidden(k1_mcart_1(k2_mcart_1(k3_mcart_1(k2_mcart_1(esk7_0),esk7_0,X1))),X2)|~r2_hidden(X1,k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_23, c_0_70])).
cnf(c_0_75, negated_conjecture, (r2_hidden(k1_mcart_1(esk7_0),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_23, c_0_71])).
cnf(c_0_76, negated_conjecture, (esk1_0=k1_xboole_0|esk2_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_72]), c_0_73])])).
cnf(c_0_77, negated_conjecture, (r2_hidden(k1_mcart_1(k2_mcart_1(k3_mcart_1(k2_mcart_1(esk7_0),esk7_0,k1_mcart_1(esk7_0)))),esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_78, negated_conjecture, (esk1_0=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_76]), c_0_73])])).
cnf(c_0_79, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_78]), c_0_69]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 80
# Proof object clause steps            : 59
# Proof object formula steps           : 21
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 23
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 21
# Processed clauses                    : 194
# ...of these trivial                  : 0
# ...subsumed                          : 63
# ...remaining for further processing  : 131
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 19
# Generated clauses                    : 612
# ...of the previous two non-trivial   : 602
# Contextual simplify-reflections      : 0
# Paramodulations                      : 612
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 86
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 446
# ...number of literals in the above   : 1652
# Current number of archived formulas  : 0
# Current number of archived clauses   : 46
# Clause-clause subsumption calls (NU) : 438
# Rec. Clause-clause subsumption calls : 258
# Non-unit clause-clause subsumptions  : 27
# Unit Clause-clause subsumption calls : 69
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 17
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10739

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.006 s
# Total time               : 0.016 s
# Maximum resident set size: 28800 pages
