# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_xboole_0, axiom, ![X1, X2, X3]:(r2_hidden(X1,k5_xboole_0(X2,X3))<=>~((r2_hidden(X1,X2)<=>r2_hidden(X1,X3)))), file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', t1_xboole_0)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', t5_boole)).
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', d2_zfmisc_1)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', t134_zfmisc_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', l54_zfmisc_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', t113_zfmisc_1)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', t6_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', d3_tarski)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/dac391d42eb6421694b986a544f18da0.E_input', d8_xboole_0)).
fof(c_0_9, plain, ![X29, X30, X31]:(((~r2_hidden(X29,X30)|~r2_hidden(X29,X31)|~r2_hidden(X29,k5_xboole_0(X30,X31)))&(r2_hidden(X29,X30)|r2_hidden(X29,X31)|~r2_hidden(X29,k5_xboole_0(X30,X31))))&((~r2_hidden(X29,X30)|r2_hidden(X29,X31)|r2_hidden(X29,k5_xboole_0(X30,X31)))&(~r2_hidden(X29,X31)|r2_hidden(X29,X30)|r2_hidden(X29,k5_xboole_0(X30,X31))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_0])])])).
fof(c_0_10, plain, ![X15]:k5_xboole_0(X15,k1_xboole_0)=X15, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_11, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, plain, (r2_hidden(X1,X3)|r2_hidden(X1,k5_xboole_0(X3,X2))|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_15, plain, ![X32, X33, X34, X35, X38, X39, X40, X41, X42, X43, X45, X46]:((((r2_hidden(esk7_4(X32,X33,X34,X35),X32)|~r2_hidden(X35,X34)|X34!=k2_zfmisc_1(X32,X33))&((r2_hidden(esk8_4(X32,X33,X34,X35),X33)|~r2_hidden(X35,X34)|X34!=k2_zfmisc_1(X32,X33))&(X35=k4_tarski(esk7_4(X32,X33,X34,X35),esk8_4(X32,X33,X34,X35))|~r2_hidden(X35,X34)|X34!=k2_zfmisc_1(X32,X33))))&(~r2_hidden(X39,X32)|(~r2_hidden(X40,X33)|X38!=k4_tarski(X39,X40))|r2_hidden(X38,X34)|X34!=k2_zfmisc_1(X32,X33)))&((~r2_hidden(esk9_3(X41,X42,X43),X43)|(~r2_hidden(X45,X41)|(~r2_hidden(X46,X42)|esk9_3(X41,X42,X43)!=k4_tarski(X45,X46)))|X43=k2_zfmisc_1(X41,X42))&((r2_hidden(esk10_3(X41,X42,X43),X41)|r2_hidden(esk9_3(X41,X42,X43),X43)|X43=k2_zfmisc_1(X41,X42))&((r2_hidden(esk11_3(X41,X42,X43),X42)|r2_hidden(esk9_3(X41,X42,X43),X43)|X43=k2_zfmisc_1(X41,X42))&(esk9_3(X41,X42,X43)=k4_tarski(esk10_3(X41,X42,X43),esk11_3(X41,X42,X43))|r2_hidden(esk9_3(X41,X42,X43),X43)|X43=k2_zfmisc_1(X41,X42)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
cnf(c_0_16, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_12]), c_0_14])).
cnf(c_0_17, plain, (r2_hidden(esk11_3(X1,X2,X3),X2)|r2_hidden(esk9_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (r2_hidden(esk8_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|X3!=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (X1=k2_zfmisc_1(X2,k1_xboole_0)|r2_hidden(esk9_3(X2,k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20, plain, (r2_hidden(esk9_3(X1,k1_xboole_0,X2),X2)|~r2_hidden(X3,X2)), inference(er,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_16])])).
cnf(c_0_21, plain, (X1=k2_zfmisc_1(X2,k1_xboole_0)|r2_hidden(esk9_3(X3,k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_20, c_0_19])).
cnf(c_0_22, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_16, c_0_21])).
fof(c_0_23, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
cnf(c_0_24, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_22]), c_0_16])).
fof(c_0_25, plain, ![X11, X12, X13, X14]:(((r2_hidden(X11,X13)|~r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(X13,X14)))&(r2_hidden(X12,X14)|~r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(X13,X14))))&(~r2_hidden(X11,X13)|~r2_hidden(X12,X14)|r2_hidden(k4_tarski(X11,X12),k2_zfmisc_1(X13,X14)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
fof(c_0_26, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk1_0!=esk3_0|esk2_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_23])])])).
cnf(c_0_27, plain, (X1=k2_zfmisc_1(X2,X3)|r2_hidden(esk11_3(X2,X3,X1),X3)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_17])).
cnf(c_0_28, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_30, plain, ![X22, X23]:((k2_zfmisc_1(X22,X23)!=k1_xboole_0|(X22=k1_xboole_0|X23=k1_xboole_0))&((X22!=k1_xboole_0|k2_zfmisc_1(X22,X23)=k1_xboole_0)&(X23!=k1_xboole_0|k2_zfmisc_1(X22,X23)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_31, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|r2_hidden(esk11_3(X1,X2,k1_xboole_0),X2)), inference(er,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (r2_hidden(esk10_3(X1,X2,X3),X1)|r2_hidden(esk9_3(X1,X2,X3),X3)|X3=k2_zfmisc_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_34, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_36, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|r2_hidden(esk11_3(esk3_0,esk4_0,k1_xboole_0),esk4_0)), inference(spm,[status(thm)],[c_0_29, c_0_31])).
cnf(c_0_37, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_39, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_32, c_0_29])).
cnf(c_0_40, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (X1=k2_zfmisc_1(k1_xboole_0,X2)|r2_hidden(esk9_3(k1_xboole_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_16, c_0_33])).
cnf(c_0_42, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_32, c_0_34])).
cnf(c_0_43, negated_conjecture, (r2_hidden(esk11_3(esk3_0,esk4_0,k1_xboole_0),esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_37]), c_0_38])).
fof(c_0_44, plain, ![X24, X25]:((r2_hidden(esk6_2(X24,X25),X25)|~r2_xboole_0(X24,X25))&(~r2_hidden(esk6_2(X24,X25),X24)|~r2_xboole_0(X24,X25))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
fof(c_0_45, plain, ![X16, X17, X18, X19, X20]:((~r1_tarski(X16,X17)|(~r2_hidden(X18,X16)|r2_hidden(X18,X17)))&((r2_hidden(esk5_2(X19,X20),X19)|r1_tarski(X19,X20))&(~r2_hidden(esk5_2(X19,X20),X20)|r1_tarski(X19,X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_46, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X2,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_39, c_0_28])).
cnf(c_0_47, plain, (X1=k1_xboole_0|r2_hidden(esk9_3(k1_xboole_0,X2,X1),X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_48, plain, (X1=k2_zfmisc_1(X2,X3)|r2_hidden(esk10_3(X2,X3,X1),X2)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_33])).
cnf(c_0_49, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50, plain, (r2_hidden(esk6_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_51, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_37])).
cnf(c_0_53, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|r2_hidden(esk10_3(X1,X2,k1_xboole_0),X1)), inference(er,[status(thm)],[c_0_48])).
cnf(c_0_54, plain, (~r2_hidden(esk6_2(X1,X2),X1)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, negated_conjecture, (r2_hidden(esk6_2(X1,esk3_0),esk1_0)|~r2_xboole_0(X1,esk3_0)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
fof(c_0_56, plain, ![X27, X28]:(((r1_tarski(X27,X28)|~r2_xboole_0(X27,X28))&(X27!=X28|~r2_xboole_0(X27,X28)))&(~r1_tarski(X27,X28)|X27=X28|r2_xboole_0(X27,X28))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
cnf(c_0_57, negated_conjecture, (r1_tarski(X1,esk3_0)|~r2_hidden(esk5_2(X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_58, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_59, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k1_xboole_0|r2_hidden(esk10_3(esk3_0,esk4_0,k1_xboole_0),esk3_0)), inference(spm,[status(thm)],[c_0_29, c_0_53])).
cnf(c_0_60, negated_conjecture, (~r2_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_61, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_62, negated_conjecture, (r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_63, negated_conjecture, (r2_hidden(esk10_3(esk3_0,esk4_0,k1_xboole_0),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_59]), c_0_37]), c_0_38])).
cnf(c_0_64, plain, (k2_zfmisc_1(X1,X2)!=k1_xboole_0|~r2_hidden(X3,X2)|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_24, c_0_28])).
cnf(c_0_65, negated_conjecture, (esk3_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])])).
cnf(c_0_66, negated_conjecture, (r2_hidden(esk10_3(esk3_0,esk4_0,k1_xboole_0),esk1_0)), inference(spm,[status(thm)],[c_0_49, c_0_63])).
cnf(c_0_67, plain, (r2_hidden(esk11_3(X1,X2,k1_xboole_0),X2)|~r2_hidden(X3,X2)|~r2_hidden(X4,X1)), inference(spm,[status(thm)],[c_0_64, c_0_31])).
cnf(c_0_68, negated_conjecture, (r2_hidden(esk11_3(esk1_0,esk4_0,k1_xboole_0),esk4_0)), inference(rw,[status(thm)],[c_0_43, c_0_65])).
cnf(c_0_69, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_70, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(esk1_0,esk2_0))|~r2_hidden(X2,esk4_0)|~r2_hidden(X1,esk1_0)), inference(rw,[status(thm)],[c_0_34, c_0_65])).
cnf(c_0_71, negated_conjecture, (r2_hidden(esk10_3(esk1_0,esk4_0,k1_xboole_0),esk1_0)), inference(rw,[status(thm)],[c_0_66, c_0_65])).
cnf(c_0_72, negated_conjecture, (r2_hidden(esk11_3(X1,esk4_0,k1_xboole_0),esk4_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_73, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(k4_tarski(X2,X1),k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_69, c_0_29])).
cnf(c_0_74, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_75, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk4_0)|~r2_hidden(X2,esk1_0)), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_76, negated_conjecture, (r2_hidden(esk11_3(X1,esk1_0,k1_xboole_0),esk1_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_67, c_0_71])).
cnf(c_0_77, negated_conjecture, (r2_hidden(esk11_3(esk4_0,esk4_0,k1_xboole_0),esk4_0)), inference(spm,[status(thm)],[c_0_72, c_0_68])).
cnf(c_0_78, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)|~r2_hidden(X2,esk1_0)), inference(spm,[status(thm)],[c_0_73, c_0_28])).
cnf(c_0_79, plain, (X1=k1_xboole_0|r2_hidden(esk9_3(X2,k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_74, c_0_19])).
cnf(c_0_80, negated_conjecture, (r2_hidden(esk6_2(X1,esk4_0),esk2_0)|~r2_xboole_0(X1,esk4_0)|~r2_hidden(X2,esk1_0)), inference(spm,[status(thm)],[c_0_75, c_0_50])).
cnf(c_0_81, negated_conjecture, (r2_hidden(esk11_3(esk4_0,esk1_0,k1_xboole_0),esk1_0)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_82, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_38])).
cnf(c_0_83, negated_conjecture, (r2_hidden(esk6_2(X1,esk4_0),esk2_0)|~r2_xboole_0(X1,esk4_0)), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_84, negated_conjecture, (r1_tarski(X1,esk4_0)|~r2_hidden(esk5_2(X1,esk4_0),esk2_0)), inference(spm,[status(thm)],[c_0_51, c_0_82])).
cnf(c_0_85, negated_conjecture, (esk1_0!=esk3_0|esk2_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_86, negated_conjecture, (~r2_xboole_0(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_54, c_0_83])).
cnf(c_0_87, negated_conjecture, (r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_84, c_0_58])).
cnf(c_0_88, negated_conjecture, (esk4_0!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_65])])).
cnf(c_0_89, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_61]), c_0_87])]), c_0_88]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 90
# Proof object clause steps            : 71
# Proof object formula steps           : 19
# Proof object conjectures             : 41
# Proof object clause conjectures      : 38
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 20
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 31
# Processed clauses                    : 4195
# ...of these trivial                  : 144
# ...subsumed                          : 3282
# ...remaining for further processing  : 769
# Other redundant clauses eliminated   : 254
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 66
# Backward-rewritten                   : 64
# Generated clauses                    : 72380
# ...of the previous two non-trivial   : 66242
# Contextual simplify-reflections      : 27
# Paramodulations                      : 71902
# Factorizations                       : 66
# Equation resolutions                 : 409
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 606
#    Positive orientable unit clauses  : 50
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 545
# Current number of unprocessed clauses: 61554
# ...number of literals in the above   : 344422
# Current number of archived formulas  : 0
# Current number of archived clauses   : 161
# Clause-clause subsumption calls (NU) : 97297
# Rec. Clause-clause subsumption calls : 42959
# Non-unit clause-clause subsumptions  : 2396
# Unit Clause-clause subsumption calls : 4625
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 41
# BW rewrite match successes           : 30
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1288929

# -------------------------------------------------
# User time                : 0.961 s
# System time              : 0.033 s
# Total time               : 0.994 s
# Maximum resident set size: 51300 pages
