# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.006 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_zfmisc_1, axiom, ![X1, X2, X3]:(X3=k2_zfmisc_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5, X6]:(r2_hidden(X5,X1)&(r2_hidden(X6,X2)&X4=k4_tarski(X5,X6))))), file('data/MPTP/proofs/dadb4d32ec5240f79df0e0a04f60324f.E_input', d2_zfmisc_1)).
fof(d3_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r1_tarski(X1,X2)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X1)=>r2_hidden(k4_tarski(X3,X4),X2)))), file('data/MPTP/proofs/dadb4d32ec5240f79df0e0a04f60324f.E_input', d3_relat_1)).
fof(dt_k1_wellord2, axiom, ![X1]:v1_relat_1(k1_wellord2(X1)), file('data/MPTP/proofs/dadb4d32ec5240f79df0e0a04f60324f.E_input', dt_k1_wellord2)).
fof(d1_wellord2, axiom, ![X1, X2]:(v1_relat_1(X2)=>(X2=k1_wellord2(X1)<=>(k3_relat_1(X2)=X1&![X3, X4]:((r2_hidden(X3,X1)&r2_hidden(X4,X1))=>(r2_hidden(k4_tarski(X3,X4),X2)<=>r1_tarski(X3,X4)))))), file('data/MPTP/proofs/dadb4d32ec5240f79df0e0a04f60324f.E_input', d1_wellord2)).
fof(t33_wellord2, conjecture, ![X1]:r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,X1)), file('data/MPTP/proofs/dadb4d32ec5240f79df0e0a04f60324f.E_input', t33_wellord2)).
fof(t30_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k3_relat_1(X3))&r2_hidden(X2,k3_relat_1(X3))))), file('data/MPTP/proofs/dadb4d32ec5240f79df0e0a04f60324f.E_input', t30_relat_1)).
fof(c_0_6, plain, ![X15, X16, X17, X18, X21, X22, X23, X24, X25, X26, X28, X29]:((((r2_hidden(esk4_4(X15,X16,X17,X18),X15)|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))&((r2_hidden(esk5_4(X15,X16,X17,X18),X16)|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))&(X18=k4_tarski(esk4_4(X15,X16,X17,X18),esk5_4(X15,X16,X17,X18))|~r2_hidden(X18,X17)|X17!=k2_zfmisc_1(X15,X16))))&(~r2_hidden(X22,X15)|(~r2_hidden(X23,X16)|X21!=k4_tarski(X22,X23))|r2_hidden(X21,X17)|X17!=k2_zfmisc_1(X15,X16)))&((~r2_hidden(esk6_3(X24,X25,X26),X26)|(~r2_hidden(X28,X24)|(~r2_hidden(X29,X25)|esk6_3(X24,X25,X26)!=k4_tarski(X28,X29)))|X26=k2_zfmisc_1(X24,X25))&((r2_hidden(esk7_3(X24,X25,X26),X24)|r2_hidden(esk6_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25))&((r2_hidden(esk8_3(X24,X25,X26),X25)|r2_hidden(esk6_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25))&(esk6_3(X24,X25,X26)=k4_tarski(esk7_3(X24,X25,X26),esk8_3(X24,X25,X26))|r2_hidden(esk6_3(X24,X25,X26),X26)|X26=k2_zfmisc_1(X24,X25)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_zfmisc_1])])])])])])).
fof(c_0_7, plain, ![X8, X9, X10, X11, X12]:((~r1_tarski(X8,X9)|(~r2_hidden(k4_tarski(X10,X11),X8)|r2_hidden(k4_tarski(X10,X11),X9))|~v1_relat_1(X8))&((r2_hidden(k4_tarski(esk2_2(X8,X12),esk3_2(X8,X12)),X8)|r1_tarski(X8,X12)|~v1_relat_1(X8))&(~r2_hidden(k4_tarski(esk2_2(X8,X12),esk3_2(X8,X12)),X12)|r1_tarski(X8,X12)|~v1_relat_1(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_relat_1])])])])])])).
fof(c_0_8, plain, ![X41]:v1_relat_1(k1_wellord2(X41)), inference(variable_rename,[status(thm)],[dt_k1_wellord2])).
fof(c_0_9, plain, ![X35, X36, X37, X38]:(((k3_relat_1(X36)=X35|X36!=k1_wellord2(X35)|~v1_relat_1(X36))&((~r2_hidden(k4_tarski(X37,X38),X36)|r1_tarski(X37,X38)|(~r2_hidden(X37,X35)|~r2_hidden(X38,X35))|X36!=k1_wellord2(X35)|~v1_relat_1(X36))&(~r1_tarski(X37,X38)|r2_hidden(k4_tarski(X37,X38),X36)|(~r2_hidden(X37,X35)|~r2_hidden(X38,X35))|X36!=k1_wellord2(X35)|~v1_relat_1(X36))))&(((r2_hidden(esk9_2(X35,X36),X35)|k3_relat_1(X36)!=X35|X36=k1_wellord2(X35)|~v1_relat_1(X36))&(r2_hidden(esk10_2(X35,X36),X35)|k3_relat_1(X36)!=X35|X36=k1_wellord2(X35)|~v1_relat_1(X36)))&((~r2_hidden(k4_tarski(esk9_2(X35,X36),esk10_2(X35,X36)),X36)|~r1_tarski(esk9_2(X35,X36),esk10_2(X35,X36))|k3_relat_1(X36)!=X35|X36=k1_wellord2(X35)|~v1_relat_1(X36))&(r2_hidden(k4_tarski(esk9_2(X35,X36),esk10_2(X35,X36)),X36)|r1_tarski(esk9_2(X35,X36),esk10_2(X35,X36))|k3_relat_1(X36)!=X35|X36=k1_wellord2(X35)|~v1_relat_1(X36))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_wellord2])])])])])).
cnf(c_0_10, plain, (r2_hidden(X5,X6)|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)|X5!=k4_tarski(X1,X3)|X6!=k2_zfmisc_1(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_11, negated_conjecture, ~(![X1]:r1_tarski(k1_wellord2(X1),k2_zfmisc_1(X1,X1))), inference(assume_negation,[status(cth)],[t33_wellord2])).
fof(c_0_12, plain, ![X32, X33, X34]:((r2_hidden(X32,k3_relat_1(X34))|~r2_hidden(k4_tarski(X32,X33),X34)|~v1_relat_1(X34))&(r2_hidden(X33,k3_relat_1(X34))|~r2_hidden(k4_tarski(X32,X33),X34)|~v1_relat_1(X34))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t30_relat_1])])])).
cnf(c_0_13, plain, (r2_hidden(k4_tarski(esk2_2(X1,X2),esk3_2(X1,X2)),X1)|r1_tarski(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, plain, (v1_relat_1(k1_wellord2(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (k3_relat_1(X1)=X2|X1!=k1_wellord2(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X3))|X1!=k4_tarski(X4,X5)|~r2_hidden(X5,X3)|~r2_hidden(X4,X2)), inference(er,[status(thm)],[c_0_10])).
fof(c_0_17, negated_conjecture, ~r1_tarski(k1_wellord2(esk1_0),k2_zfmisc_1(esk1_0,esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])).
cnf(c_0_18, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (r2_hidden(k4_tarski(esk2_2(k1_wellord2(X1),X2),esk3_2(k1_wellord2(X1),X2)),k1_wellord2(X1))|r1_tarski(k1_wellord2(X1),X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, plain, (k3_relat_1(k1_wellord2(X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_14])])).
cnf(c_0_21, plain, (r2_hidden(X1,k3_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_22, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(esk2_2(X1,X2),esk3_2(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (~r1_tarski(k1_wellord2(esk1_0),k2_zfmisc_1(esk1_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r2_hidden(esk3_2(k1_wellord2(X1),X2),X1)|r1_tarski(k1_wellord2(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_14])])).
cnf(c_0_26, plain, (r2_hidden(esk2_2(k1_wellord2(X1),X2),X1)|r1_tarski(k1_wellord2(X1),X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_19]), c_0_20]), c_0_14])])).
cnf(c_0_27, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(esk3_2(X1,k2_zfmisc_1(X2,X3)),X3)|~r2_hidden(esk2_2(X1,k2_zfmisc_1(X2,X3)),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk3_2(k1_wellord2(esk1_0),k2_zfmisc_1(esk1_0,esk1_0)),esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, negated_conjecture, (r2_hidden(esk2_2(k1_wellord2(esk1_0),k2_zfmisc_1(esk1_0,esk1_0)),esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_26])).
cnf(c_0_30, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_14])]), c_0_24]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 31
# Proof object clause steps            : 18
# Proof object formula steps           : 13
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 22
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 22
# Processed clauses                    : 260
# ...of these trivial                  : 0
# ...subsumed                          : 59
# ...remaining for further processing  : 201
# Other redundant clauses eliminated   : 173
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 4353
# ...of the previous two non-trivial   : 4152
# Contextual simplify-reflections      : 0
# Paramodulations                      : 4148
# Factorizations                       : 14
# Equation resolutions                 : 191
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 179
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 173
# Current number of unprocessed clauses: 3910
# ...number of literals in the above   : 27350
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 6020
# Rec. Clause-clause subsumption calls : 2986
# Non-unit clause-clause subsumptions  : 59
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 99804

# -------------------------------------------------
# User time                : 0.069 s
# System time              : 0.003 s
# Total time               : 0.072 s
# Maximum resident set size: 28716 pages
