# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t48_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k5_relat_1(X3,X1),k5_relat_1(X3,X2)))))), file('data/MPTP/proofs/db56f357aaf84528b9f853913c501e76.E_input', t48_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/db56f357aaf84528b9f853913c501e76.E_input', d3_tarski)).
fof(d8_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(X3=k5_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>?[X6]:(r2_hidden(k4_tarski(X4,X6),X1)&r2_hidden(k4_tarski(X6,X5),X2))))))), file('data/MPTP/proofs/db56f357aaf84528b9f853913c501e76.E_input', d8_relat_1)).
fof(dt_k5_relat_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_relat_1(X2))=>v1_relat_1(k5_relat_1(X1,X2))), file('data/MPTP/proofs/db56f357aaf84528b9f853913c501e76.E_input', dt_k5_relat_1)).
fof(d3_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r1_tarski(X1,X2)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X1)=>r2_hidden(k4_tarski(X3,X4),X2)))), file('data/MPTP/proofs/db56f357aaf84528b9f853913c501e76.E_input', d3_relat_1)).
fof(c_0_5, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>![X3]:(v1_relat_1(X3)=>(r1_tarski(X1,X2)=>r1_tarski(k5_relat_1(X3,X1),k5_relat_1(X3,X2))))))), inference(assume_negation,[status(cth)],[t48_relat_1])).
fof(c_0_6, plain, ![X10, X11, X12, X13, X14]:((~r1_tarski(X10,X11)|(~r2_hidden(X12,X10)|r2_hidden(X12,X11)))&((r2_hidden(esk4_2(X13,X14),X13)|r1_tarski(X13,X14))&(~r2_hidden(esk4_2(X13,X14),X14)|r1_tarski(X13,X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_7, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(v1_relat_1(esk3_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k5_relat_1(esk3_0,esk1_0),k5_relat_1(esk3_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
fof(c_0_8, plain, ![X16, X17, X18, X19, X20, X22, X23, X24, X27]:((((r2_hidden(k4_tarski(X19,esk5_5(X16,X17,X18,X19,X20)),X16)|~r2_hidden(k4_tarski(X19,X20),X18)|X18!=k5_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17)|~v1_relat_1(X16))&(r2_hidden(k4_tarski(esk5_5(X16,X17,X18,X19,X20),X20),X17)|~r2_hidden(k4_tarski(X19,X20),X18)|X18!=k5_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17)|~v1_relat_1(X16)))&(~r2_hidden(k4_tarski(X22,X24),X16)|~r2_hidden(k4_tarski(X24,X23),X17)|r2_hidden(k4_tarski(X22,X23),X18)|X18!=k5_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17)|~v1_relat_1(X16)))&((~r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk7_3(X16,X17,X18)),X18)|(~r2_hidden(k4_tarski(esk6_3(X16,X17,X18),X27),X16)|~r2_hidden(k4_tarski(X27,esk7_3(X16,X17,X18)),X17))|X18=k5_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17)|~v1_relat_1(X16))&((r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk8_3(X16,X17,X18)),X16)|r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk7_3(X16,X17,X18)),X18)|X18=k5_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17)|~v1_relat_1(X16))&(r2_hidden(k4_tarski(esk8_3(X16,X17,X18),esk7_3(X16,X17,X18)),X17)|r2_hidden(k4_tarski(esk6_3(X16,X17,X18),esk7_3(X16,X17,X18)),X18)|X18=k5_relat_1(X16,X17)|~v1_relat_1(X18)|~v1_relat_1(X17)|~v1_relat_1(X16))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_relat_1])])])])])])).
fof(c_0_9, plain, ![X29, X30]:(~v1_relat_1(X29)|~v1_relat_1(X30)|v1_relat_1(k5_relat_1(X29,X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_relat_1])])).
cnf(c_0_10, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(esk5_5(X1,X2,X3,X4,X5),X5),X2)|~r2_hidden(k4_tarski(X4,X5),X3)|X3!=k5_relat_1(X1,X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (v1_relat_1(k5_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,X4),X6)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X2,X4),X5)|X6!=k5_relat_1(X3,X5)|~v1_relat_1(X6)|~v1_relat_1(X5)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_16, plain, (r2_hidden(k4_tarski(esk5_5(X1,X2,k5_relat_1(X1,X2),X3,X4),X4),X2)|~r2_hidden(k4_tarski(X3,X4),k5_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_12]), c_0_13])).
cnf(c_0_17, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_18, plain, (r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,X4))|~r2_hidden(k4_tarski(X5,X2),X4)|~r2_hidden(k4_tarski(X1,X5),X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_14]), c_0_13])).
cnf(c_0_19, negated_conjecture, (r2_hidden(k4_tarski(esk5_5(X1,esk1_0,k5_relat_1(X1,esk1_0),X2,X3),X3),esk2_0)|~r2_hidden(k4_tarski(X2,X3),k5_relat_1(X1,esk1_0))|~v1_relat_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_20, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X1,esk5_5(X2,X3,X4,X1,X5)),X2)|~r2_hidden(k4_tarski(X1,X5),X4)|X4!=k5_relat_1(X2,X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_22, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,esk2_0))|~r2_hidden(k4_tarski(X1,esk5_5(X4,esk1_0,k5_relat_1(X4,esk1_0),X5,X2)),X3)|~r2_hidden(k4_tarski(X5,X2),k5_relat_1(X4,esk1_0))|~v1_relat_1(X3)|~v1_relat_1(X4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20])])).
cnf(c_0_23, plain, (r2_hidden(k4_tarski(X1,esk5_5(X2,X3,k5_relat_1(X2,X3),X1,X4)),X2)|~r2_hidden(k4_tarski(X1,X4),k5_relat_1(X2,X3))|~v1_relat_1(X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_21]), c_0_13])).
fof(c_0_24, plain, ![X31, X32, X33, X34, X35]:((~r1_tarski(X31,X32)|(~r2_hidden(k4_tarski(X33,X34),X31)|r2_hidden(k4_tarski(X33,X34),X32))|~v1_relat_1(X31))&((r2_hidden(k4_tarski(esk9_2(X31,X35),esk10_2(X31,X35)),X31)|r1_tarski(X31,X35)|~v1_relat_1(X31))&(~r2_hidden(k4_tarski(esk9_2(X31,X35),esk10_2(X31,X35)),X35)|r1_tarski(X31,X35)|~v1_relat_1(X31)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_relat_1])])])])])])).
cnf(c_0_25, negated_conjecture, (r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,esk2_0))|~r2_hidden(k4_tarski(X1,X2),k5_relat_1(X3,esk1_0))|~v1_relat_1(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_17])])).
cnf(c_0_26, plain, (r2_hidden(k4_tarski(esk9_2(X1,X2),esk10_2(X1,X2)),X1)|r1_tarski(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_27, plain, (r1_tarski(X1,X2)|~r2_hidden(k4_tarski(esk9_2(X1,X2),esk10_2(X1,X2)),X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_28, negated_conjecture, (r2_hidden(k4_tarski(esk9_2(k5_relat_1(X1,esk1_0),X2),esk10_2(k5_relat_1(X1,esk1_0),X2)),k5_relat_1(X1,esk2_0))|r1_tarski(k5_relat_1(X1,esk1_0),X2)|~v1_relat_1(k5_relat_1(X1,esk1_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29, negated_conjecture, (~r1_tarski(k5_relat_1(esk3_0,esk1_0),k5_relat_1(esk3_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_30, negated_conjecture, (r1_tarski(k5_relat_1(X1,esk1_0),k5_relat_1(X1,esk2_0))|~v1_relat_1(k5_relat_1(X1,esk1_0))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_31, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_32, negated_conjecture, (~v1_relat_1(k5_relat_1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_31])])).
cnf(c_0_33, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_13]), c_0_17]), c_0_31])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 34
# Proof object clause steps            : 23
# Proof object formula steps           : 11
# Proof object conjectures             : 16
# Proof object clause conjectures      : 13
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 18
# Processed clauses                    : 124
# ...of these trivial                  : 1
# ...subsumed                          : 17
# ...remaining for further processing  : 106
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 248
# ...of the previous two non-trivial   : 243
# Contextual simplify-reflections      : 6
# Paramodulations                      : 245
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 86
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 79
# Current number of unprocessed clauses: 154
# ...number of literals in the above   : 1020
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 3422
# Rec. Clause-clause subsumption calls : 472
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 5
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 10126

# -------------------------------------------------
# User time                : 0.016 s
# System time              : 0.000 s
# Total time               : 0.016 s
# Maximum resident set size: 28808 pages
