# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C47_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t89_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6]:(k17_mcart_1(k4_tarski(k4_tarski(X1,X2),X3))=X1&(k18_mcart_1(k4_tarski(k4_tarski(X1,X2),X3))=X2&(k19_mcart_1(k4_tarski(X6,k4_tarski(X4,X5)))=X4&k20_mcart_1(k4_tarski(X6,k4_tarski(X4,X5)))=X5))), file('data/MPTP/proofs/dbd73676af644e498c3d8bf17fdd8a02.E_input', t89_mcart_1)).
fof(d14_mcart_1, axiom, ![X1]:k17_mcart_1(X1)=k1_mcart_1(k1_mcart_1(X1)), file('data/MPTP/proofs/dbd73676af644e498c3d8bf17fdd8a02.E_input', d14_mcart_1)).
fof(d15_mcart_1, axiom, ![X1]:k18_mcart_1(X1)=k2_mcart_1(k1_mcart_1(X1)), file('data/MPTP/proofs/dbd73676af644e498c3d8bf17fdd8a02.E_input', d15_mcart_1)).
fof(d16_mcart_1, axiom, ![X1]:k19_mcart_1(X1)=k1_mcart_1(k2_mcart_1(X1)), file('data/MPTP/proofs/dbd73676af644e498c3d8bf17fdd8a02.E_input', d16_mcart_1)).
fof(d17_mcart_1, axiom, ![X1]:k20_mcart_1(X1)=k2_mcart_1(k2_mcart_1(X1)), file('data/MPTP/proofs/dbd73676af644e498c3d8bf17fdd8a02.E_input', d17_mcart_1)).
fof(t7_mcart_1, axiom, ![X1, X2]:(k1_mcart_1(k4_tarski(X1,X2))=X1&k2_mcart_1(k4_tarski(X1,X2))=X2), file('data/MPTP/proofs/dbd73676af644e498c3d8bf17fdd8a02.E_input', t7_mcart_1)).
fof(c_0_6, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6]:(k17_mcart_1(k4_tarski(k4_tarski(X1,X2),X3))=X1&(k18_mcart_1(k4_tarski(k4_tarski(X1,X2),X3))=X2&(k19_mcart_1(k4_tarski(X6,k4_tarski(X4,X5)))=X4&k20_mcart_1(k4_tarski(X6,k4_tarski(X4,X5)))=X5)))), inference(assume_negation,[status(cth)],[t89_mcart_1])).
fof(c_0_7, negated_conjecture, (k17_mcart_1(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0))!=esk1_0|(k18_mcart_1(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0))!=esk2_0|(k19_mcart_1(k4_tarski(esk6_0,k4_tarski(esk4_0,esk5_0)))!=esk4_0|k20_mcart_1(k4_tarski(esk6_0,k4_tarski(esk4_0,esk5_0)))!=esk5_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
fof(c_0_8, plain, ![X13]:k17_mcart_1(X13)=k1_mcart_1(k1_mcart_1(X13)), inference(variable_rename,[status(thm)],[d14_mcart_1])).
fof(c_0_9, plain, ![X17]:k18_mcart_1(X17)=k2_mcart_1(k1_mcart_1(X17)), inference(variable_rename,[status(thm)],[d15_mcart_1])).
fof(c_0_10, plain, ![X18]:k19_mcart_1(X18)=k1_mcart_1(k2_mcart_1(X18)), inference(variable_rename,[status(thm)],[d16_mcart_1])).
fof(c_0_11, plain, ![X16]:k20_mcart_1(X16)=k2_mcart_1(k2_mcart_1(X16)), inference(variable_rename,[status(thm)],[d17_mcart_1])).
cnf(c_0_12, negated_conjecture, (k17_mcart_1(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0))!=esk1_0|k18_mcart_1(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0))!=esk2_0|k19_mcart_1(k4_tarski(esk6_0,k4_tarski(esk4_0,esk5_0)))!=esk4_0|k20_mcart_1(k4_tarski(esk6_0,k4_tarski(esk4_0,esk5_0)))!=esk5_0), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (k17_mcart_1(X1)=k1_mcart_1(k1_mcart_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (k18_mcart_1(X1)=k2_mcart_1(k1_mcart_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (k19_mcart_1(X1)=k1_mcart_1(k2_mcart_1(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (k20_mcart_1(X1)=k2_mcart_1(k2_mcart_1(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, plain, ![X14, X15]:(k1_mcart_1(k4_tarski(X14,X15))=X14&k2_mcart_1(k4_tarski(X14,X15))=X15), inference(variable_rename,[status(thm)],[t7_mcart_1])).
cnf(c_0_18, negated_conjecture, (k1_mcart_1(k2_mcart_1(k4_tarski(esk6_0,k4_tarski(esk4_0,esk5_0))))!=esk4_0|k2_mcart_1(k2_mcart_1(k4_tarski(esk6_0,k4_tarski(esk4_0,esk5_0))))!=esk5_0|k1_mcart_1(k1_mcart_1(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0)))!=esk1_0|k2_mcart_1(k1_mcart_1(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0)))!=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_14]), c_0_15]), c_0_16])).
cnf(c_0_19, plain, (k1_mcart_1(k4_tarski(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, negated_conjecture, (k1_mcart_1(k2_mcart_1(k4_tarski(esk6_0,k4_tarski(esk4_0,esk5_0))))!=esk4_0|k2_mcart_1(k2_mcart_1(k4_tarski(esk6_0,k4_tarski(esk4_0,esk5_0))))!=esk5_0|k2_mcart_1(k4_tarski(esk1_0,esk2_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_19]), c_0_19])])).
cnf(c_0_21, plain, (k2_mcart_1(k4_tarski(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_21]), c_0_19]), c_0_21]), c_0_21]), c_0_21])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 10
# Proof object formula steps           : 13
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 0
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 7
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 3
# Processed clauses                    : 3
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 3
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 0
# ...of the previous two non-trivial   : 0
# Contextual simplify-reflections      : 0
# Paramodulations                      : 0
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 5
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 337

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.000 s
# Total time               : 0.008 s
# Maximum resident set size: 28716 pages
