# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', d2_subset_1)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', t7_boole)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', d3_tarski)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', t3_subset)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', existence_m1_subset_1)).
fof(l13_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', l13_xboole_0)).
fof(fc3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>v1_xboole_0(k1_struct_0(X1))), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', fc3_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', dt_l1_orders_2)).
fof(t45_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k1_struct_0(X1))=u1_struct_0(X1)), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', t45_orders_2)).
fof(dt_k2_orders_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', dt_k2_orders_2)).
fof(fraenkel_a_2_1_orders_2, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&(v3_orders_2(X2)&(v4_orders_2(X2)&(v5_orders_2(X2)&l1_orders_2(X2)))))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))=>(r2_hidden(X1,a_2_1_orders_2(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r2_hidden(X5,X3)=>r2_orders_2(X2,X4,X5))))))), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', fraenkel_a_2_1_orders_2)).
fof(d13_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2))), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', d13_orders_2)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', fc1_xboole_0)).
fof(d2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k1_struct_0(X1)=k1_xboole_0), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', d2_struct_0)).
fof(t49_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(![X3]:(m1_subset_1(X3,X1)=>r2_hidden(X3,X2))=>X1=X2)), file('data/MPTP/proofs/dc6fa40df23e4926b59e73bb0fbc61a3.E_input', t49_subset_1)).
fof(c_0_15, plain, ![X14, X15]:(((~m1_subset_1(X15,X14)|r2_hidden(X15,X14)|v1_xboole_0(X14))&(~r2_hidden(X15,X14)|m1_subset_1(X15,X14)|v1_xboole_0(X14)))&((~m1_subset_1(X15,X14)|v1_xboole_0(X15)|~v1_xboole_0(X14))&(~v1_xboole_0(X15)|m1_subset_1(X15,X14)|~v1_xboole_0(X14)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_16, plain, ![X11, X12]:(~r2_hidden(X11,X12)|~v1_xboole_0(X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_17, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, ![X33, X34, X35, X36, X37]:((~r1_tarski(X33,X34)|(~r2_hidden(X35,X33)|r2_hidden(X35,X34)))&((r2_hidden(esk6_2(X36,X37),X36)|r1_tarski(X36,X37))&(~r2_hidden(esk6_2(X36,X37),X37)|r1_tarski(X36,X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_20, plain, ![X30, X31]:((~m1_subset_1(X30,k1_zfmisc_1(X31))|r1_tarski(X30,X31))&(~r1_tarski(X30,X31)|m1_subset_1(X30,k1_zfmisc_1(X31)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_21, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_22, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (r1_tarski(X1,X2)|m1_subset_1(esk6_2(X1,X2),X1)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (r1_tarski(esk6_2(k1_zfmisc_1(X1),X2),X1)|r1_tarski(k1_zfmisc_1(X1),X2)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|r2_hidden(X3,X1)|~r2_hidden(X3,esk6_2(k1_zfmisc_1(X1),X2))), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_28, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_29, plain, ![X18]:m1_subset_1(esk2_1(X18),X18), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_30, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|v1_xboole_0(esk6_2(k1_zfmisc_1(X1),X2))|r2_hidden(X3,X1)|~m1_subset_1(X3,esk6_2(k1_zfmisc_1(X1),X2))), inference(spm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_31, plain, (m1_subset_1(esk2_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_32, plain, ![X32]:(~v1_xboole_0(X32)|X32=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l13_xboole_0])])).
cnf(c_0_33, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|v1_xboole_0(esk6_2(k1_zfmisc_1(X1),X2))|r2_hidden(esk2_1(esk6_2(k1_zfmisc_1(X1),X2)),X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_35, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|v1_xboole_0(esk6_2(k1_zfmisc_1(X1),X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_18, c_0_33])).
cnf(c_0_36, plain, (esk6_2(k1_zfmisc_1(X1),X2)=k1_xboole_0|r1_tarski(k1_zfmisc_1(X1),X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_37, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|~v1_xboole_0(X1)|~r2_hidden(X3,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_36]), c_0_18])).
fof(c_0_38, plain, ![X16]:(~l1_struct_0(X16)|v1_xboole_0(k1_struct_0(X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_struct_0])])).
cnf(c_0_39, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|r1_tarski(k1_xboole_0,X3)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_37, c_0_22])).
cnf(c_0_40, plain, (v1_xboole_0(k1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_41, plain, ![X17]:(~l1_orders_2(X17)|l1_struct_0(X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_42, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>k2_orders_2(X1,k1_struct_0(X1))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t45_orders_2])).
cnf(c_0_43, plain, (r1_tarski(k1_zfmisc_1(k1_struct_0(X1)),X2)|r1_tarski(k1_xboole_0,X3)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_44, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_45, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&k2_orders_2(esk1_0,k1_struct_0(esk1_0))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_42])])])])).
cnf(c_0_46, plain, (r1_tarski(k1_zfmisc_1(k1_struct_0(X1)),X2)|r1_tarski(k1_xboole_0,X3)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_48, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_struct_0(esk1_0)),X1)|r1_tarski(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_49, negated_conjecture, (r1_tarski(k1_xboole_0,X1)|r2_hidden(X2,X3)|~r2_hidden(X2,k1_zfmisc_1(k1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_25, c_0_48])).
cnf(c_0_50, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_51, negated_conjecture, (r1_tarski(k1_xboole_0,X1)|v1_xboole_0(k1_zfmisc_1(k1_struct_0(esk1_0)))|r2_hidden(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_49, c_0_28])).
cnf(c_0_52, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|~r2_hidden(esk6_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_54, plain, (v1_xboole_0(esk2_1(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_50, c_0_31])).
cnf(c_0_55, negated_conjecture, (r1_tarski(k1_xboole_0,X1)|v1_xboole_0(k1_zfmisc_1(k1_struct_0(esk1_0)))|r2_hidden(X2,X3)|~r1_tarski(X2,k1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_56, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_53, c_0_22])).
fof(c_0_57, plain, ![X9, X10]:(v2_struct_0(X9)|(~v3_orders_2(X9)|(~v4_orders_2(X9)|(~v5_orders_2(X9)|~l1_orders_2(X9))))|~m1_subset_1(X10,k1_zfmisc_1(u1_struct_0(X9)))|m1_subset_1(k2_orders_2(X9,X10),k1_zfmisc_1(u1_struct_0(X9)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_orders_2])])])).
fof(c_0_58, plain, ![X20, X21, X22, X24, X25]:(((m1_subset_1(esk3_3(X20,X21,X22),u1_struct_0(X21))|~r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&((X20=esk3_3(X20,X21,X22)|~r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&(~m1_subset_1(X24,u1_struct_0(X21))|(~r2_hidden(X24,X22)|r2_orders_2(X21,esk3_3(X20,X21,X22),X24))|~r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))))&((m1_subset_1(esk4_4(X20,X21,X22,X25),u1_struct_0(X21))|X20!=X25|~m1_subset_1(X25,u1_struct_0(X21))|r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&((r2_hidden(esk4_4(X20,X21,X22,X25),X22)|X20!=X25|~m1_subset_1(X25,u1_struct_0(X21))|r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))&(~r2_orders_2(X21,X25,esk4_4(X20,X21,X22,X25))|X20!=X25|~m1_subset_1(X25,u1_struct_0(X21))|r2_hidden(X20,a_2_1_orders_2(X21,X22))|(v2_struct_0(X21)|(~v3_orders_2(X21)|(~v4_orders_2(X21)|(~v5_orders_2(X21)|~l1_orders_2(X21))))|~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X21)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_orders_2])])])])])])).
cnf(c_0_59, plain, (esk2_1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_34, c_0_54])).
cnf(c_0_60, negated_conjecture, (r1_tarski(k1_xboole_0,X1)|v1_xboole_0(k1_zfmisc_1(k1_struct_0(esk1_0)))|r2_hidden(k1_struct_0(esk1_0),X2)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_61, plain, (v2_struct_0(X1)|m1_subset_1(k2_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_62, plain, ![X7, X8]:(v2_struct_0(X7)|(~v3_orders_2(X7)|(~v4_orders_2(X7)|(~v5_orders_2(X7)|~l1_orders_2(X7))))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X7)))|k2_orders_2(X7,X8)=a_2_1_orders_2(X7,X8))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_orders_2])])])])).
cnf(c_0_63, plain, (r1_tarski(esk2_1(k1_zfmisc_1(X1)),X1)), inference(spm,[status(thm)],[c_0_23, c_0_31])).
cnf(c_0_64, plain, (r2_hidden(esk4_4(X1,X2,X3,X4),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|X1!=X4|~m1_subset_1(X4,u1_struct_0(X2))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65, plain, (m1_subset_1(k1_xboole_0,X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_31, c_0_59])).
cnf(c_0_66, negated_conjecture, (r1_tarski(k1_xboole_0,X1)|v1_xboole_0(k1_zfmisc_1(k1_struct_0(esk1_0)))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_18, c_0_60])).
cnf(c_0_67, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_68, plain, (r1_tarski(k2_orders_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_23, c_0_61])).
cnf(c_0_69, plain, (v2_struct_0(X1)|k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_70, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk2_1(k1_zfmisc_1(X2)))), inference(spm,[status(thm)],[c_0_25, c_0_63])).
cnf(c_0_71, plain, (r2_hidden(esk4_4(X1,X2,X3,X1),X3)|r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(er,[status(thm)],[c_0_64])).
cnf(c_0_72, plain, (r1_tarski(k1_xboole_0,X1)|~v1_xboole_0(k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_23, c_0_65])).
cnf(c_0_73, negated_conjecture, (r1_tarski(k1_xboole_0,X1)|v1_xboole_0(k1_zfmisc_1(k1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_74, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k2_orders_2(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_25, c_0_68])).
cnf(c_0_75, plain, (k2_orders_2(X1,esk2_1(k1_zfmisc_1(u1_struct_0(X1))))=a_2_1_orders_2(X1,esk2_1(k1_zfmisc_1(u1_struct_0(X1))))|v2_struct_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_69, c_0_31])).
cnf(c_0_76, plain, (r2_hidden(esk4_4(X1,X2,esk2_1(k1_zfmisc_1(X3)),X1),X3)|r2_hidden(X1,a_2_1_orders_2(X2,esk2_1(k1_zfmisc_1(X3))))|v2_struct_0(X2)|~m1_subset_1(esk2_1(k1_zfmisc_1(X3)),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_77, negated_conjecture, (r1_tarski(k1_xboole_0,k1_struct_0(esk1_0))|r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_78, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,esk2_1(k1_zfmisc_1(u1_struct_0(X2)))))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_31])])).
cnf(c_0_79, plain, (r2_hidden(X1,a_2_1_orders_2(X2,esk2_1(k1_zfmisc_1(X3))))|v2_struct_0(X2)|~v1_xboole_0(X3)|~m1_subset_1(esk2_1(k1_zfmisc_1(X3)),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_18, c_0_76])).
fof(c_0_80, plain, ![X13]:(~l1_struct_0(X13)|k1_struct_0(X13)=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_struct_0])])).
cnf(c_0_81, negated_conjecture, (r1_tarski(k1_xboole_0,k1_struct_0(esk1_0))), inference(ef,[status(thm)],[c_0_77])).
fof(c_0_82, plain, ![X27, X28]:((m1_subset_1(esk5_2(X27,X28),X27)|X27=X28|~m1_subset_1(X28,k1_zfmisc_1(X27)))&(~r2_hidden(esk5_2(X27,X28),X28)|X27=X28|~m1_subset_1(X28,k1_zfmisc_1(X27)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t49_subset_1])])])])).
cnf(c_0_83, plain, (X1=esk3_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_84, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_31])]), c_0_28])).
cnf(c_0_85, plain, (k1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_86, negated_conjecture, (r2_hidden(X1,k1_struct_0(esk1_0))|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_25, c_0_81])).
cnf(c_0_87, plain, (X1=X2|~r2_hidden(esk5_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_88, plain, (r2_hidden(X1,a_2_1_orders_2(X2,X3))|v2_struct_0(X2)|~v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_18, c_0_71])).
cnf(c_0_89, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_orders_2(X2,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_90, plain, (esk3_3(esk6_2(a_2_1_orders_2(X1,X2),X3),X1,X2)=esk6_2(a_2_1_orders_2(X1,X2),X3)|r1_tarski(a_2_1_orders_2(X1,X2),X3)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_83, c_0_22])).
cnf(c_0_91, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_18, c_0_84])).
cnf(c_0_92, negated_conjecture, (k2_orders_2(esk1_0,k1_struct_0(esk1_0))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_93, plain, (k1_struct_0(X1)=k1_xboole_0|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_85, c_0_44])).
cnf(c_0_94, negated_conjecture, (~v1_xboole_0(k1_struct_0(esk1_0))|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_18, c_0_86])).
cnf(c_0_95, plain, (X1=a_2_1_orders_2(X2,X3)|v2_struct_0(X2)|~v1_xboole_0(X3)|~m1_subset_1(esk5_2(X1,a_2_1_orders_2(X2,X3)),u1_struct_0(X2))|~m1_subset_1(a_2_1_orders_2(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_96, plain, (m1_subset_1(esk5_2(X1,X2),X1)|X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_97, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk6_2(X1,X2),X2)), inference(spm,[status(thm)],[c_0_53, c_0_28])).
cnf(c_0_98, plain, (r1_tarski(a_2_1_orders_2(X1,X2),X3)|m1_subset_1(esk6_2(a_2_1_orders_2(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_22])).
cnf(c_0_99, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_91, c_0_31])).
cnf(c_0_100, negated_conjecture, (k2_orders_2(esk1_0,k1_xboole_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_47])])).
cnf(c_0_101, plain, (k2_orders_2(X1,X2)=a_2_1_orders_2(X1,X2)|v2_struct_0(X1)|~r1_tarski(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_69, c_0_52])).
cnf(c_0_102, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_103, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_104, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_105, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_106, negated_conjecture, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_93]), c_0_67]), c_0_47])])).
cnf(c_0_107, plain, (a_2_1_orders_2(X1,X2)=u1_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(a_2_1_orders_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_108, plain, (r1_tarski(a_2_1_orders_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])).
cnf(c_0_109, negated_conjecture, (a_2_1_orders_2(esk1_0,k1_xboole_0)!=u1_struct_0(esk1_0)|~r1_tarski(k1_xboole_0,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_47]), c_0_102]), c_0_103]), c_0_104])]), c_0_105])).
cnf(c_0_110, negated_conjecture, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_106, c_0_22])).
cnf(c_0_111, plain, (a_2_1_orders_2(X1,X2)=u1_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_52]), c_0_108])).
cnf(c_0_112, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_18, c_0_22])).
cnf(c_0_113, negated_conjecture, (a_2_1_orders_2(esk1_0,k1_xboole_0)!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_114, plain, (a_2_1_orders_2(X1,X2)=u1_struct_0(X1)|v2_struct_0(X1)|~v1_xboole_0(X2)|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_52]), c_0_112])).
cnf(c_0_115, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_67]), c_0_47]), c_0_102]), c_0_103]), c_0_104])]), c_0_105]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 116
# Proof object clause steps            : 86
# Proof object formula steps           : 30
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 55
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 32
# Processed clauses                    : 21142
# ...of these trivial                  : 175
# ...subsumed                          : 14589
# ...remaining for further processing  : 6378
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 342
# Backward-rewritten                   : 403
# Generated clauses                    : 584432
# ...of the previous two non-trivial   : 383419
# Contextual simplify-reflections      : 252
# Paramodulations                      : 584425
# Factorizations                       : 4
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 5630
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 5610
# Current number of unprocessed clauses: 352881
# ...number of literals in the above   : 2223877
# Current number of archived formulas  : 0
# Current number of archived clauses   : 745
# Clause-clause subsumption calls (NU) : 3898243
# Rec. Clause-clause subsumption calls : 1288026
# Non-unit clause-clause subsumptions  : 12666
# Unit Clause-clause subsumption calls : 1565
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 26
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 16668598

# -------------------------------------------------
# User time                : 8.064 s
# System time              : 0.106 s
# Total time               : 8.170 s
# Maximum resident set size: 278168 pages
