# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t204_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k11_relat_1(X3,X1)))), file('data/MPTP/proofs/dc89a68957af456790e9d24d34a191d9.E_input', t204_relat_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/dc89a68957af456790e9d24d34a191d9.E_input', t3_xboole_0)).
fof(l24_zfmisc_1, axiom, ![X1, X2]:~((r1_xboole_0(k1_tarski(X1),X2)&r2_hidden(X1,X2))), file('data/MPTP/proofs/dc89a68957af456790e9d24d34a191d9.E_input', l24_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/dc89a68957af456790e9d24d34a191d9.E_input', t69_enumset1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/dc89a68957af456790e9d24d34a191d9.E_input', d4_relat_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/dc89a68957af456790e9d24d34a191d9.E_input', t65_xboole_1)).
fof(t205_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k1_relat_1(X2))<=>~(k11_relat_1(X2,X1)=k1_xboole_0))), file('data/MPTP/proofs/dc89a68957af456790e9d24d34a191d9.E_input', t205_relat_1)).
fof(d16_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))), file('data/MPTP/proofs/dc89a68957af456790e9d24d34a191d9.E_input', d16_relat_1)).
fof(c_0_8, plain, ![X10, X11, X12]:((~r2_hidden(k4_tarski(X10,X11),X12)|r2_hidden(X11,k11_relat_1(X12,X10))|~v1_relat_1(X12))&(~r2_hidden(X11,k11_relat_1(X12,X10))|r2_hidden(k4_tarski(X10,X11),X12)|~v1_relat_1(X12))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t204_relat_1])])])).
fof(c_0_9, plain, ![X47, X48, X50, X51, X52]:(((r2_hidden(esk8_2(X47,X48),X47)|r1_xboole_0(X47,X48))&(r2_hidden(esk8_2(X47,X48),X48)|r1_xboole_0(X47,X48)))&(~r2_hidden(X52,X50)|~r2_hidden(X52,X51)|~r1_xboole_0(X50,X51))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_10, plain, ![X45, X46]:(~r1_xboole_0(k1_tarski(X45),X46)|~r2_hidden(X45,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l24_zfmisc_1])])).
fof(c_0_11, plain, ![X7]:k2_tarski(X7,X7)=k1_tarski(X7), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_12, plain, ![X32, X33, X34, X36, X37, X38, X39, X41]:(((~r2_hidden(X34,X33)|r2_hidden(k4_tarski(X34,esk5_3(X32,X33,X34)),X32)|X33!=k1_relat_1(X32))&(~r2_hidden(k4_tarski(X36,X37),X32)|r2_hidden(X36,X33)|X33!=k1_relat_1(X32)))&((~r2_hidden(esk6_2(X38,X39),X39)|~r2_hidden(k4_tarski(esk6_2(X38,X39),X41),X38)|X39=k1_relat_1(X38))&(r2_hidden(esk6_2(X38,X39),X39)|r2_hidden(k4_tarski(esk6_2(X38,X39),esk7_2(X38,X39)),X38)|X39=k1_relat_1(X38)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_13, plain, (r2_hidden(k4_tarski(X3,X1),X2)|~r2_hidden(X1,k11_relat_1(X2,X3))|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r2_hidden(esk8_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (~r1_xboole_0(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_17, plain, ![X31]:r1_xboole_0(X31,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_18, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, plain, (r1_xboole_0(X1,k11_relat_1(X2,X3))|r2_hidden(k4_tarski(X3,esk8_2(X1,k11_relat_1(X2,X3))),X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_20, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(r2_hidden(X1,k1_relat_1(X2))<=>~(k11_relat_1(X2,X1)=k1_xboole_0)))), inference(assume_negation,[status(cth)],[t205_relat_1])).
cnf(c_0_21, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(k2_tarski(X1,X1),X2)), inference(rw,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_22, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_23, plain, ![X43, X44]:(~v1_relat_1(X43)|k11_relat_1(X43,X44)=k9_relat_1(X43,k1_tarski(X44))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_1])])])).
cnf(c_0_24, plain, (r1_xboole_0(X1,k11_relat_1(X2,X3))|r2_hidden(X3,X4)|X4!=k1_relat_1(X2)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_25, negated_conjecture, (v1_relat_1(esk2_0)&((~r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)=k1_xboole_0)&(r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)!=k1_xboole_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_20])])])])).
cnf(c_0_26, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (r2_hidden(esk6_2(X1,X2),X2)|r2_hidden(k4_tarski(esk6_2(X1,X2),esk7_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_28, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (r1_xboole_0(X1,k11_relat_1(X2,X3))|r2_hidden(X3,k1_relat_1(X2))|~v1_relat_1(X2)), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_31, plain, (X1=k1_relat_1(k1_xboole_0)|r2_hidden(esk6_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k2_tarski(X2,X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_28, c_0_16])).
cnf(c_0_33, negated_conjecture, (r1_xboole_0(X1,k11_relat_1(esk2_0,X2))|r2_hidden(X2,k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_26, c_0_31])).
cnf(c_0_35, negated_conjecture, (k11_relat_1(esk2_0,esk1_0)=k1_xboole_0|~r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_36, negated_conjecture, (k11_relat_1(esk2_0,X1)=k9_relat_1(esk2_0,k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_32, c_0_30])).
cnf(c_0_37, negated_conjecture, (r2_hidden(X1,k1_relat_1(esk2_0))|~r2_hidden(X2,k11_relat_1(esk2_0,X1))), inference(spm,[status(thm)],[c_0_21, c_0_33])).
cnf(c_0_38, plain, (X1=k1_xboole_0|r2_hidden(esk6_2(k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[c_0_31, c_0_34])).
cnf(c_0_39, plain, (r2_hidden(X2,k11_relat_1(X3,X1))|~r2_hidden(k4_tarski(X1,X2),X3)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_40, negated_conjecture, (k9_relat_1(esk2_0,k2_tarski(esk1_0,esk1_0))=k1_xboole_0|~r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_41, negated_conjecture, (k11_relat_1(esk2_0,X1)=k1_xboole_0|r2_hidden(X1,k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_42, negated_conjecture, (r2_hidden(X1,k9_relat_1(esk2_0,k2_tarski(X2,X2)))|~r2_hidden(k4_tarski(X2,X1),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_36]), c_0_30])])).
cnf(c_0_43, negated_conjecture, (k9_relat_1(esk2_0,k2_tarski(esk1_0,esk1_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_36])])).
cnf(c_0_44, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk2_0))|k11_relat_1(esk2_0,esk1_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_45, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,X1),esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_26])).
cnf(c_0_46, plain, (r2_hidden(k4_tarski(X1,esk5_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_47, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk2_0))|k9_relat_1(esk2_0,k2_tarski(esk1_0,esk1_0))!=k1_xboole_0), inference(rw,[status(thm)],[c_0_44, c_0_36])).
cnf(c_0_48, negated_conjecture, (X1!=k1_relat_1(esk2_0)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_49, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_43])])).
cnf(c_0_50, negated_conjecture, ($false), inference(spm,[status(thm)],[c_0_48, c_0_49]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 51
# Proof object clause steps            : 34
# Proof object formula steps           : 17
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 29
# Processed clauses                    : 2807
# ...of these trivial                  : 104
# ...subsumed                          : 2130
# ...remaining for further processing  : 573
# Other redundant clauses eliminated   : 56
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 19
# Backward-rewritten                   : 31
# Generated clauses                    : 16969
# ...of the previous two non-trivial   : 14862
# Contextual simplify-reflections      : 9
# Paramodulations                      : 16756
# Factorizations                       : 36
# Equation resolutions                 : 174
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 491
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 456
# Current number of unprocessed clauses: 12030
# ...number of literals in the above   : 51161
# Current number of archived formulas  : 0
# Current number of archived clauses   : 80
# Clause-clause subsumption calls (NU) : 97525
# Rec. Clause-clause subsumption calls : 56410
# Non-unit clause-clause subsumptions  : 1795
# Unit Clause-clause subsumption calls : 2198
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 72
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 199305

# -------------------------------------------------
# User time                : 0.161 s
# System time              : 0.003 s
# Total time               : 0.165 s
# Maximum resident set size: 28644 pages
