# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', d2_tarski)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', t73_enumset1)).
fof(t143_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k9_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k1_relat_1(X3))&(r2_hidden(k4_tarski(X4,X1),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', t143_relat_1)).
fof(d16_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', d16_relat_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', t69_enumset1)).
fof(t204_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k11_relat_1(X3,X1)))), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', t204_relat_1)).
fof(t20_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)=>(r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X2,k2_relat_1(X3))))), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', t20_relat_1)).
fof(d4_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:(X7=k4_enumset1(X1,X2,X3,X4,X5,X6)<=>![X8]:(r2_hidden(X8,X7)<=>~((~(X8=X1)&(~(X8=X2)&(~(X8=X3)&(~(X8=X4)&(~(X8=X5)&~(X8=X6))))))))), file('data/MPTP/proofs/dcde2a4571064a49a4792bf6a2b004f9.E_input', d4_enumset1)).
fof(c_0_11, plain, ![X24, X25, X26, X27, X28, X29, X30, X31]:(((~r2_hidden(X27,X26)|(X27=X24|X27=X25)|X26!=k2_tarski(X24,X25))&((X28!=X24|r2_hidden(X28,X26)|X26!=k2_tarski(X24,X25))&(X28!=X25|r2_hidden(X28,X26)|X26!=k2_tarski(X24,X25))))&(((esk5_3(X29,X30,X31)!=X29|~r2_hidden(esk5_3(X29,X30,X31),X31)|X31=k2_tarski(X29,X30))&(esk5_3(X29,X30,X31)!=X30|~r2_hidden(esk5_3(X29,X30,X31),X31)|X31=k2_tarski(X29,X30)))&(r2_hidden(esk5_3(X29,X30,X31),X31)|(esk5_3(X29,X30,X31)=X29|esk5_3(X29,X30,X31)=X30)|X31=k2_tarski(X29,X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_12, plain, ![X33, X34]:k1_enumset1(X33,X33,X34)=k2_tarski(X33,X34), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_13, plain, ![X16, X17, X18]:k2_enumset1(X16,X16,X17,X18)=k1_enumset1(X16,X17,X18), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_14, plain, ![X12, X13, X14, X15]:k3_enumset1(X12,X12,X13,X14,X15)=k2_enumset1(X12,X13,X14,X15), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_15, plain, ![X41, X42, X43, X44, X45]:k4_enumset1(X41,X41,X42,X43,X44,X45)=k3_enumset1(X41,X42,X43,X44,X45), inference(variable_rename,[status(thm)],[t73_enumset1])).
cnf(c_0_16, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (X1=X4|X1=X3|X2!=k4_enumset1(X3,X3,X3,X3,X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19]), c_0_20])).
fof(c_0_22, plain, ![X19, X20, X21, X23]:(((r2_hidden(esk4_3(X19,X20,X21),k1_relat_1(X21))|~r2_hidden(X19,k9_relat_1(X21,X20))|~v1_relat_1(X21))&((r2_hidden(k4_tarski(esk4_3(X19,X20,X21),X19),X21)|~r2_hidden(X19,k9_relat_1(X21,X20))|~v1_relat_1(X21))&(r2_hidden(esk4_3(X19,X20,X21),X20)|~r2_hidden(X19,k9_relat_1(X21,X20))|~v1_relat_1(X21))))&(~r2_hidden(X23,k1_relat_1(X21))|(~r2_hidden(k4_tarski(X23,X19),X21)|~r2_hidden(X23,X20))|r2_hidden(X19,k9_relat_1(X21,X20))|~v1_relat_1(X21))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t143_relat_1])])])])])).
fof(c_0_23, plain, ![X38, X39]:(~v1_relat_1(X38)|k11_relat_1(X38,X39)=k9_relat_1(X38,k1_tarski(X39))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d16_relat_1])])])).
fof(c_0_24, plain, ![X40]:k2_tarski(X40,X40)=k1_tarski(X40), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_25, plain, (X1=X2|X1=X3|~r2_hidden(X1,k4_enumset1(X3,X3,X3,X3,X3,X2))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_26, plain, (r2_hidden(esk4_3(X1,X2,X3),X2)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k1_tarski(X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (esk4_3(X1,k4_enumset1(X2,X2,X2,X2,X2,X3),X4)=X3|esk4_3(X1,k4_enumset1(X2,X2,X2,X2,X2,X3),X4)=X2|~r2_hidden(X1,k9_relat_1(X4,k4_enumset1(X2,X2,X2,X2,X2,X3)))|~v1_relat_1(X4)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30, plain, (k11_relat_1(X1,X2)=k9_relat_1(X1,k4_enumset1(X2,X2,X2,X2,X2,X2))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_17]), c_0_18]), c_0_19]), c_0_20])).
cnf(c_0_31, plain, (r2_hidden(k4_tarski(esk4_3(X1,X2,X3),X1),X3)|~r2_hidden(X1,k9_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (esk4_3(X1,k4_enumset1(X2,X2,X2,X2,X2,X2),X3)=X2|~r2_hidden(X1,k11_relat_1(X3,X2))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_33, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(k4_tarski(X1,X2),X3)<=>r2_hidden(X2,k11_relat_1(X3,X1))))), inference(assume_negation,[status(cth)],[t204_relat_1])).
fof(c_0_34, plain, ![X35, X36, X37]:((r2_hidden(X35,k1_relat_1(X37))|~r2_hidden(k4_tarski(X35,X36),X37)|~v1_relat_1(X37))&(r2_hidden(X36,k2_relat_1(X37))|~r2_hidden(k4_tarski(X35,X36),X37)|~v1_relat_1(X37))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t20_relat_1])])])).
cnf(c_0_35, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,k9_relat_1(X3,k4_enumset1(X1,X1,X1,X1,X1,X1)))|~r2_hidden(X2,k11_relat_1(X3,X1))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
fof(c_0_36, negated_conjecture, (v1_relat_1(esk3_0)&((~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0)))&(r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_33])])])).
cnf(c_0_37, plain, (r2_hidden(X3,k9_relat_1(X2,X4))|~r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~r2_hidden(X1,X4)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_38, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_39, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,k11_relat_1(X3,X1))|~v1_relat_1(X3)), inference(spm,[status(thm)],[c_0_35, c_0_30])).
cnf(c_0_40, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_41, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_42, plain, (r2_hidden(X1,k9_relat_1(X2,X3))|~r2_hidden(k4_tarski(X4,X1),X2)|~r2_hidden(X4,X3)|~v1_relat_1(X2)), inference(csr,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_43, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
fof(c_0_44, plain, ![X46, X47, X48, X49, X50, X51, X52, X53, X54, X55, X56, X57, X58, X59, X60, X61]:(((~r2_hidden(X53,X52)|(X53=X46|(X53=X47|(X53=X48|(X53=X49|(X53=X50|X53=X51)))))|X52!=k4_enumset1(X46,X47,X48,X49,X50,X51))&((X54!=X46|r2_hidden(X54,X52)|X52!=k4_enumset1(X46,X47,X48,X49,X50,X51))&((X54!=X47|r2_hidden(X54,X52)|X52!=k4_enumset1(X46,X47,X48,X49,X50,X51))&((X54!=X48|r2_hidden(X54,X52)|X52!=k4_enumset1(X46,X47,X48,X49,X50,X51))&((X54!=X49|r2_hidden(X54,X52)|X52!=k4_enumset1(X46,X47,X48,X49,X50,X51))&((X54!=X50|r2_hidden(X54,X52)|X52!=k4_enumset1(X46,X47,X48,X49,X50,X51))&(X54!=X51|r2_hidden(X54,X52)|X52!=k4_enumset1(X46,X47,X48,X49,X50,X51))))))))&(((esk6_7(X55,X56,X57,X58,X59,X60,X61)!=X55|~r2_hidden(esk6_7(X55,X56,X57,X58,X59,X60,X61),X61)|X61=k4_enumset1(X55,X56,X57,X58,X59,X60))&((esk6_7(X55,X56,X57,X58,X59,X60,X61)!=X56|~r2_hidden(esk6_7(X55,X56,X57,X58,X59,X60,X61),X61)|X61=k4_enumset1(X55,X56,X57,X58,X59,X60))&((esk6_7(X55,X56,X57,X58,X59,X60,X61)!=X57|~r2_hidden(esk6_7(X55,X56,X57,X58,X59,X60,X61),X61)|X61=k4_enumset1(X55,X56,X57,X58,X59,X60))&((esk6_7(X55,X56,X57,X58,X59,X60,X61)!=X58|~r2_hidden(esk6_7(X55,X56,X57,X58,X59,X60,X61),X61)|X61=k4_enumset1(X55,X56,X57,X58,X59,X60))&((esk6_7(X55,X56,X57,X58,X59,X60,X61)!=X59|~r2_hidden(esk6_7(X55,X56,X57,X58,X59,X60,X61),X61)|X61=k4_enumset1(X55,X56,X57,X58,X59,X60))&(esk6_7(X55,X56,X57,X58,X59,X60,X61)!=X60|~r2_hidden(esk6_7(X55,X56,X57,X58,X59,X60,X61),X61)|X61=k4_enumset1(X55,X56,X57,X58,X59,X60)))))))&(r2_hidden(esk6_7(X55,X56,X57,X58,X59,X60,X61),X61)|(esk6_7(X55,X56,X57,X58,X59,X60,X61)=X55|(esk6_7(X55,X56,X57,X58,X59,X60,X61)=X56|(esk6_7(X55,X56,X57,X58,X59,X60,X61)=X57|(esk6_7(X55,X56,X57,X58,X59,X60,X61)=X58|(esk6_7(X55,X56,X57,X58,X59,X60,X61)=X59|esk6_7(X55,X56,X57,X58,X59,X60,X61)=X60)))))|X61=k4_enumset1(X55,X56,X57,X58,X59,X60)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_enumset1])])])])])])])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk2_0,k9_relat_1(esk3_0,X1))|~r2_hidden(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_41])])).
cnf(c_0_46, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k4_enumset1(X4,X5,X6,X2,X7,X8)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_47, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk2_0,k11_relat_1(esk3_0,X1))|~r2_hidden(esk1_0,k4_enumset1(X1,X1,X1,X1,X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_30]), c_0_41])])).
cnf(c_0_49, plain, (r2_hidden(X1,k4_enumset1(X2,X3,X4,X1,X5,X6))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_46])])).
cnf(c_0_50, negated_conjecture, (~r2_hidden(esk2_0,k11_relat_1(esk3_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_43])])).
cnf(c_0_51, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 52
# Proof object clause steps            : 29
# Proof object formula steps           : 23
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 11
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 35
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 30
# Processed clauses                    : 422
# ...of these trivial                  : 0
# ...subsumed                          : 157
# ...remaining for further processing  : 265
# Other redundant clauses eliminated   : 144
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 43
# Backward-rewritten                   : 57
# Generated clauses                    : 2626
# ...of the previous two non-trivial   : 2405
# Contextual simplify-reflections      : 7
# Paramodulations                      : 2406
# Factorizations                       : 84
# Equation resolutions                 : 144
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 127
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 115
# Current number of unprocessed clauses: 2006
# ...number of literals in the above   : 13427
# Current number of archived formulas  : 0
# Current number of archived clauses   : 133
# Clause-clause subsumption calls (NU) : 8590
# Rec. Clause-clause subsumption calls : 1585
# Non-unit clause-clause subsumptions  : 207
# Unit Clause-clause subsumption calls : 48
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 39
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 72400

# -------------------------------------------------
# User time                : 0.063 s
# System time              : 0.004 s
# Total time               : 0.067 s
# Maximum resident set size: 28852 pages
