# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(free_g1_orders_2, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3, X4]:(g1_orders_2(X1,X2)=g1_orders_2(X3,X4)=>(X1=X3&X2=X4))), file('data/MPTP/proofs/dee36f0073a84813a968664639338f65.E_input', free_g1_orders_2)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/dee36f0073a84813a968664639338f65.E_input', dt_u1_orders_2)).
fof(t14_yellow_0, conjecture, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:((r1_yellow_0(X1,X3)=>r1_yellow_0(X2,X3))&(r2_yellow_0(X1,X3)=>r2_yellow_0(X2,X3)))))), file('data/MPTP/proofs/dee36f0073a84813a968664639338f65.E_input', t14_yellow_0)).
fof(t2_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3, X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(X4=X5=>((r2_lattice3(X1,X3,X4)=>r2_lattice3(X2,X3,X5))&(r1_lattice3(X1,X3,X4)=>r1_lattice3(X2,X3,X5))))))))), file('data/MPTP/proofs/dee36f0073a84813a968664639338f65.E_input', t2_yellow_0)).
fof(d8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X5)=>r1_orders_2(X1,X5,X4))))=>X4=X3))))))), file('data/MPTP/proofs/dee36f0073a84813a968664639338f65.E_input', d8_yellow_0)).
fof(d7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r1_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r2_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X4)=>r1_orders_2(X1,X3,X4)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r2_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r2_lattice3(X1,X2,X5)=>r1_orders_2(X1,X4,X5))))=>X4=X3))))))), file('data/MPTP/proofs/dee36f0073a84813a968664639338f65.E_input', d7_yellow_0)).
fof(t1_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>![X6]:(m1_subset_1(X6,u1_struct_0(X2))=>((X3=X5&X4=X6)=>((r1_orders_2(X1,X3,X4)=>r1_orders_2(X2,X5,X6))&(r2_orders_2(X1,X3,X4)=>r2_orders_2(X2,X5,X6))))))))))), file('data/MPTP/proofs/dee36f0073a84813a968664639338f65.E_input', t1_yellow_0)).
fof(c_0_7, plain, ![X7, X8, X9, X10]:((X7=X9|g1_orders_2(X7,X8)!=g1_orders_2(X9,X10)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X7,X7))))&(X8=X10|g1_orders_2(X7,X8)!=g1_orders_2(X9,X10)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X7,X7))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_orders_2])])])])).
fof(c_0_8, plain, ![X11]:(~l1_orders_2(X11)|m1_subset_1(u1_orders_2(X11),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X11))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
fof(c_0_9, negated_conjecture, ~(![X1]:(l1_orders_2(X1)=>![X2]:(l1_orders_2(X2)=>(g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))=g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))=>![X3]:((r1_yellow_0(X1,X3)=>r1_yellow_0(X2,X3))&(r2_yellow_0(X1,X3)=>r2_yellow_0(X2,X3))))))), inference(assume_negation,[status(cth)],[t14_yellow_0])).
cnf(c_0_10, plain, (X1=X2|g1_orders_2(X3,X1)!=g1_orders_2(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_12, negated_conjecture, (l1_orders_2(esk1_0)&(l1_orders_2(esk2_0)&(g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))&(((r2_yellow_0(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk3_0))&(~r2_yellow_0(esk2_0,esk3_0)|r1_yellow_0(esk1_0,esk3_0)))&((r2_yellow_0(esk1_0,esk3_0)|~r1_yellow_0(esk2_0,esk3_0))&(~r2_yellow_0(esk2_0,esk3_0)|~r1_yellow_0(esk2_0,esk3_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])])).
cnf(c_0_13, plain, (u1_orders_2(X1)=X2|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(X3,X2)|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk1_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_15, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, negated_conjecture, (u1_orders_2(esk1_0)=X1|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_17, negated_conjecture, (u1_orders_2(esk1_0)=u1_orders_2(esk2_0)), inference(er,[status(thm)],[c_0_16])).
fof(c_0_18, plain, ![X15, X16, X17, X18, X19]:((~r2_lattice3(X15,X17,X18)|r2_lattice3(X16,X17,X19)|X18!=X19|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X15))|g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))!=g1_orders_2(u1_struct_0(X16),u1_orders_2(X16))|~l1_orders_2(X16)|~l1_orders_2(X15))&(~r1_lattice3(X15,X17,X18)|r1_lattice3(X16,X17,X19)|X18!=X19|~m1_subset_1(X19,u1_struct_0(X16))|~m1_subset_1(X18,u1_struct_0(X15))|g1_orders_2(u1_struct_0(X15),u1_orders_2(X15))!=g1_orders_2(u1_struct_0(X16),u1_orders_2(X16))|~l1_orders_2(X16)|~l1_orders_2(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_yellow_0])])])])).
cnf(c_0_19, plain, (X1=X2|g1_orders_2(X1,X3)!=g1_orders_2(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20, negated_conjecture, (m1_subset_1(u1_orders_2(esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_17]), c_0_15])])).
cnf(c_0_21, negated_conjecture, (g1_orders_2(u1_struct_0(esk1_0),u1_orders_2(esk2_0))=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))), inference(rw,[status(thm)],[c_0_14, c_0_17])).
cnf(c_0_22, plain, (r1_lattice3(X4,X2,X5)|~r1_lattice3(X1,X2,X3)|X3!=X5|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X20, X21, X23, X24, X26, X27, X30]:(((m1_subset_1(esk4_2(X20,X21),u1_struct_0(X20))|~r2_yellow_0(X20,X21)|~l1_orders_2(X20))&((r1_lattice3(X20,X21,esk4_2(X20,X21))|~r2_yellow_0(X20,X21)|~l1_orders_2(X20))&((~m1_subset_1(X23,u1_struct_0(X20))|(~r1_lattice3(X20,X21,X23)|r1_orders_2(X20,X23,esk4_2(X20,X21)))|~r2_yellow_0(X20,X21)|~l1_orders_2(X20))&((m1_subset_1(esk5_3(X20,X21,X24),u1_struct_0(X20))|~r1_lattice3(X20,X21,X24)|X24=esk4_2(X20,X21)|~m1_subset_1(X24,u1_struct_0(X20))|~r2_yellow_0(X20,X21)|~l1_orders_2(X20))&((r1_lattice3(X20,X21,esk5_3(X20,X21,X24))|~r1_lattice3(X20,X21,X24)|X24=esk4_2(X20,X21)|~m1_subset_1(X24,u1_struct_0(X20))|~r2_yellow_0(X20,X21)|~l1_orders_2(X20))&(~r1_orders_2(X20,esk5_3(X20,X21,X24),X24)|~r1_lattice3(X20,X21,X24)|X24=esk4_2(X20,X21)|~m1_subset_1(X24,u1_struct_0(X20))|~r2_yellow_0(X20,X21)|~l1_orders_2(X20)))))))&(((m1_subset_1(esk7_3(X20,X26,X27),u1_struct_0(X20))|m1_subset_1(esk6_3(X20,X26,X27),u1_struct_0(X20))|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))&(((r1_lattice3(X20,X26,esk7_3(X20,X26,X27))|m1_subset_1(esk6_3(X20,X26,X27),u1_struct_0(X20))|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))&(~m1_subset_1(X30,u1_struct_0(X20))|(~r1_lattice3(X20,X26,X30)|r1_orders_2(X20,X30,esk7_3(X20,X26,X27)))|m1_subset_1(esk6_3(X20,X26,X27),u1_struct_0(X20))|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20)))&(esk7_3(X20,X26,X27)!=X27|m1_subset_1(esk6_3(X20,X26,X27),u1_struct_0(X20))|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))))&(((m1_subset_1(esk7_3(X20,X26,X27),u1_struct_0(X20))|r1_lattice3(X20,X26,esk6_3(X20,X26,X27))|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))&(((r1_lattice3(X20,X26,esk7_3(X20,X26,X27))|r1_lattice3(X20,X26,esk6_3(X20,X26,X27))|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))&(~m1_subset_1(X30,u1_struct_0(X20))|(~r1_lattice3(X20,X26,X30)|r1_orders_2(X20,X30,esk7_3(X20,X26,X27)))|r1_lattice3(X20,X26,esk6_3(X20,X26,X27))|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20)))&(esk7_3(X20,X26,X27)!=X27|r1_lattice3(X20,X26,esk6_3(X20,X26,X27))|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))))&((m1_subset_1(esk7_3(X20,X26,X27),u1_struct_0(X20))|~r1_orders_2(X20,esk6_3(X20,X26,X27),X27)|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))&(((r1_lattice3(X20,X26,esk7_3(X20,X26,X27))|~r1_orders_2(X20,esk6_3(X20,X26,X27),X27)|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))&(~m1_subset_1(X30,u1_struct_0(X20))|(~r1_lattice3(X20,X26,X30)|r1_orders_2(X20,X30,esk7_3(X20,X26,X27)))|~r1_orders_2(X20,esk6_3(X20,X26,X27),X27)|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20)))&(esk7_3(X20,X26,X27)!=X27|~r1_orders_2(X20,esk6_3(X20,X26,X27),X27)|~r1_lattice3(X20,X26,X27)|~m1_subset_1(X27,u1_struct_0(X20))|r2_yellow_0(X20,X26)|~l1_orders_2(X20))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_0])])])])])])).
cnf(c_0_24, plain, (r2_lattice3(X4,X2,X5)|~r2_lattice3(X1,X2,X3)|X3!=X5|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_25, plain, ![X31, X32, X34, X35, X37, X38, X41]:(((m1_subset_1(esk8_2(X31,X32),u1_struct_0(X31))|~r1_yellow_0(X31,X32)|~l1_orders_2(X31))&((r2_lattice3(X31,X32,esk8_2(X31,X32))|~r1_yellow_0(X31,X32)|~l1_orders_2(X31))&((~m1_subset_1(X34,u1_struct_0(X31))|(~r2_lattice3(X31,X32,X34)|r1_orders_2(X31,esk8_2(X31,X32),X34))|~r1_yellow_0(X31,X32)|~l1_orders_2(X31))&((m1_subset_1(esk9_3(X31,X32,X35),u1_struct_0(X31))|~r2_lattice3(X31,X32,X35)|X35=esk8_2(X31,X32)|~m1_subset_1(X35,u1_struct_0(X31))|~r1_yellow_0(X31,X32)|~l1_orders_2(X31))&((r2_lattice3(X31,X32,esk9_3(X31,X32,X35))|~r2_lattice3(X31,X32,X35)|X35=esk8_2(X31,X32)|~m1_subset_1(X35,u1_struct_0(X31))|~r1_yellow_0(X31,X32)|~l1_orders_2(X31))&(~r1_orders_2(X31,X35,esk9_3(X31,X32,X35))|~r2_lattice3(X31,X32,X35)|X35=esk8_2(X31,X32)|~m1_subset_1(X35,u1_struct_0(X31))|~r1_yellow_0(X31,X32)|~l1_orders_2(X31)))))))&(((m1_subset_1(esk11_3(X31,X37,X38),u1_struct_0(X31))|m1_subset_1(esk10_3(X31,X37,X38),u1_struct_0(X31))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))&(((r2_lattice3(X31,X37,esk11_3(X31,X37,X38))|m1_subset_1(esk10_3(X31,X37,X38),u1_struct_0(X31))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))&(~m1_subset_1(X41,u1_struct_0(X31))|(~r2_lattice3(X31,X37,X41)|r1_orders_2(X31,esk11_3(X31,X37,X38),X41))|m1_subset_1(esk10_3(X31,X37,X38),u1_struct_0(X31))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31)))&(esk11_3(X31,X37,X38)!=X38|m1_subset_1(esk10_3(X31,X37,X38),u1_struct_0(X31))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))))&(((m1_subset_1(esk11_3(X31,X37,X38),u1_struct_0(X31))|r2_lattice3(X31,X37,esk10_3(X31,X37,X38))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))&(((r2_lattice3(X31,X37,esk11_3(X31,X37,X38))|r2_lattice3(X31,X37,esk10_3(X31,X37,X38))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))&(~m1_subset_1(X41,u1_struct_0(X31))|(~r2_lattice3(X31,X37,X41)|r1_orders_2(X31,esk11_3(X31,X37,X38),X41))|r2_lattice3(X31,X37,esk10_3(X31,X37,X38))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31)))&(esk11_3(X31,X37,X38)!=X38|r2_lattice3(X31,X37,esk10_3(X31,X37,X38))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))))&((m1_subset_1(esk11_3(X31,X37,X38),u1_struct_0(X31))|~r1_orders_2(X31,X38,esk10_3(X31,X37,X38))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))&(((r2_lattice3(X31,X37,esk11_3(X31,X37,X38))|~r1_orders_2(X31,X38,esk10_3(X31,X37,X38))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))&(~m1_subset_1(X41,u1_struct_0(X31))|(~r2_lattice3(X31,X37,X41)|r1_orders_2(X31,esk11_3(X31,X37,X38),X41))|~r1_orders_2(X31,X38,esk10_3(X31,X37,X38))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31)))&(esk11_3(X31,X37,X38)!=X38|~r1_orders_2(X31,X38,esk10_3(X31,X37,X38))|~r2_lattice3(X31,X37,X38)|~m1_subset_1(X38,u1_struct_0(X31))|r1_yellow_0(X31,X37)|~l1_orders_2(X31))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_yellow_0])])])])])])).
cnf(c_0_26, negated_conjecture, (u1_struct_0(esk1_0)=X1|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_27, plain, (r1_lattice3(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_lattice3(X4,X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r1_lattice3(X1,X2,esk4_2(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (m1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, ![X42, X43, X44, X45, X46, X47]:((~r1_orders_2(X42,X44,X45)|r1_orders_2(X43,X46,X47)|(X44!=X46|X45!=X47)|~m1_subset_1(X47,u1_struct_0(X43))|~m1_subset_1(X46,u1_struct_0(X43))|~m1_subset_1(X45,u1_struct_0(X42))|~m1_subset_1(X44,u1_struct_0(X42))|g1_orders_2(u1_struct_0(X42),u1_orders_2(X42))!=g1_orders_2(u1_struct_0(X43),u1_orders_2(X43))|~l1_orders_2(X43)|~l1_orders_2(X42))&(~r2_orders_2(X42,X44,X45)|r2_orders_2(X43,X46,X47)|(X44!=X46|X45!=X47)|~m1_subset_1(X47,u1_struct_0(X43))|~m1_subset_1(X46,u1_struct_0(X43))|~m1_subset_1(X45,u1_struct_0(X42))|~m1_subset_1(X44,u1_struct_0(X42))|g1_orders_2(u1_struct_0(X42),u1_orders_2(X42))!=g1_orders_2(u1_struct_0(X43),u1_orders_2(X43))|~l1_orders_2(X43)|~l1_orders_2(X42))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_yellow_0])])])])).
cnf(c_0_31, plain, (r2_lattice3(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X4,X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (r2_lattice3(X1,X2,esk8_2(X1,X2))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, negated_conjecture, (u1_struct_0(esk1_0)=u1_struct_0(esk2_0)), inference(er,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (r1_lattice3(X1,X2,esk4_2(X3,X2))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_yellow_0(X3,X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk4_2(X3,X2),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_36, plain, (r1_orders_2(X4,X5,X6)|~r1_orders_2(X1,X2,X3)|X2!=X5|X3!=X6|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~l1_orders_2(X4)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (r2_lattice3(X1,X2,esk8_2(X3,X2))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_yellow_0(X3,X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk8_2(X3,X2),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_38, plain, (r1_orders_2(X2,X1,esk7_3(X2,X3,X4))|r1_lattice3(X2,X3,esk6_3(X2,X3,X4))|r2_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r1_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_39, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,X1),u1_struct_0(esk2_0))|~r2_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_34]), c_0_15])])).
cnf(c_0_40, negated_conjecture, (l1_orders_2(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_41, negated_conjecture, (r1_lattice3(X1,X2,esk4_2(esk1_0,X2))|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk4_2(esk1_0,X2),u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_17]), c_0_34]), c_0_15])])).
cnf(c_0_42, plain, (r1_lattice3(X1,X2,esk5_3(X1,X2,X3))|X3=esk4_2(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_43, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|X3=esk4_2(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_44, plain, (r1_orders_2(X1,X2,X3)|g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_orders_2(X4,X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_36])])).
cnf(c_0_45, plain, (r1_orders_2(X2,X1,esk4_2(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_46, plain, (r1_orders_2(X2,esk11_3(X2,X3,X4),X1)|r2_lattice3(X2,X3,esk10_3(X2,X3,X4))|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk2_0))|~r1_yellow_0(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_15])])).
cnf(c_0_48, negated_conjecture, (r2_lattice3(X1,X2,esk8_2(esk1_0,X2))|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk8_2(esk1_0,X2),u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_17]), c_0_34]), c_0_15])])).
cnf(c_0_49, plain, (r2_lattice3(X1,X2,esk9_3(X1,X2,X3))|X3=esk8_2(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_50, plain, (m1_subset_1(esk9_3(X1,X2,X3),u1_struct_0(X1))|X3=esk8_2(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r1_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_51, plain, (r1_orders_2(X2,esk8_2(X2,X3),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_52, negated_conjecture, (r1_orders_2(esk2_0,X1,esk7_3(esk2_0,X2,esk4_2(esk1_0,X3)))|r1_lattice3(esk2_0,X2,esk6_3(esk2_0,X2,esk4_2(esk1_0,X3)))|r2_yellow_0(esk2_0,X2)|~r1_lattice3(esk2_0,X2,esk4_2(esk1_0,X3))|~r1_lattice3(esk2_0,X2,X1)|~r2_yellow_0(esk1_0,X3)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])])).
cnf(c_0_53, negated_conjecture, (r1_lattice3(esk2_0,X1,esk4_2(esk1_0,X1))|~r2_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_41]), c_0_40])]), c_0_39])).
cnf(c_0_54, plain, (X1=esk4_2(X2,X3)|r1_lattice3(X4,X3,esk5_3(X2,X3,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~l1_orders_2(X4)|~l1_orders_2(X2)|~m1_subset_1(esk5_3(X2,X3,X1),u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_42]), c_0_43])).
cnf(c_0_55, plain, (r1_orders_2(X1,X2,esk4_2(X3,X4))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_lattice3(X3,X4,X2)|~r2_yellow_0(X3,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk4_2(X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_29])).
cnf(c_0_56, plain, (r1_lattice3(X1,X2,esk6_3(X1,X2,X3))|r2_yellow_0(X1,X2)|esk7_3(X1,X2,X3)!=X3|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_57, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|esk7_3(X1,X2,X3)!=X3|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_58, negated_conjecture, (r1_orders_2(esk2_0,esk11_3(esk2_0,X1,esk8_2(esk1_0,X2)),X3)|r2_lattice3(esk2_0,X1,esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,esk8_2(esk1_0,X2))|~r2_lattice3(esk2_0,X1,X3)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(X3,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_40])])).
cnf(c_0_59, negated_conjecture, (r2_lattice3(esk2_0,X1,esk8_2(esk1_0,X1))|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_48]), c_0_40])]), c_0_47])).
cnf(c_0_60, plain, (X1=esk8_2(X2,X3)|r2_lattice3(X4,X3,esk9_3(X2,X3,X1))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~r2_lattice3(X2,X3,X1)|~r1_yellow_0(X2,X3)|~l1_orders_2(X4)|~l1_orders_2(X2)|~m1_subset_1(esk9_3(X2,X3,X1),u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_49]), c_0_50])).
cnf(c_0_61, plain, (r1_orders_2(X1,esk8_2(X2,X3),X4)|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X2,X3,X4)|~r1_yellow_0(X2,X3)|~l1_orders_2(X1)|~l1_orders_2(X2)|~m1_subset_1(esk8_2(X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_51]), c_0_33])).
cnf(c_0_62, plain, (r2_lattice3(X1,X2,esk10_3(X1,X2,X3))|r1_yellow_0(X1,X2)|esk11_3(X1,X2,X3)!=X3|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_63, plain, (m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|esk11_3(X1,X2,X3)!=X3|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_64, negated_conjecture, (r1_orders_2(esk2_0,X1,esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r1_lattice3(esk2_0,X2,esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X2)|~r1_lattice3(esk2_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_65, negated_conjecture, (X1=esk4_2(esk1_0,X2)|m1_subset_1(esk5_3(esk1_0,X2,X1),u1_struct_0(esk2_0))|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_34]), c_0_15])])).
cnf(c_0_66, negated_conjecture, (X1=esk4_2(esk1_0,X2)|r1_lattice3(X3,X2,esk5_3(esk1_0,X2,X1))|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~l1_orders_2(X3)|~m1_subset_1(esk5_3(esk1_0,X2,X1),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_67, negated_conjecture, (r1_orders_2(X1,X2,esk4_2(esk1_0,X3))|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_lattice3(esk1_0,X3,X2)|~r2_yellow_0(esk1_0,X3)|~l1_orders_2(X1)|~m1_subset_1(esk4_2(esk1_0,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_68, plain, (r1_lattice3(X1,X2,esk6_3(X3,X2,X4))|r2_yellow_0(X3,X2)|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|esk7_3(X3,X2,X4)!=X4|~r1_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk6_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_56]), c_0_57])).
cnf(c_0_69, negated_conjecture, (r1_orders_2(esk2_0,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)),X2)|r2_lattice3(esk2_0,X1,esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_70, negated_conjecture, (X1=esk8_2(esk1_0,X2)|m1_subset_1(esk9_3(esk1_0,X2,X1),u1_struct_0(esk2_0))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_34]), c_0_15])])).
cnf(c_0_71, negated_conjecture, (X1=esk8_2(esk1_0,X2)|r2_lattice3(X3,X2,esk9_3(esk1_0,X2,X1))|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X3)|~m1_subset_1(esk9_3(esk1_0,X2,X1),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_72, negated_conjecture, (r1_orders_2(X1,esk8_2(esk1_0,X2),X3)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk1_0,X2,X3)|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk8_2(esk1_0,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_73, plain, (r2_lattice3(X1,X2,esk10_3(X3,X2,X4))|r1_yellow_0(X3,X2)|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|esk11_3(X3,X2,X4)!=X4|~r2_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk10_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_62]), c_0_63])).
cnf(c_0_74, negated_conjecture, (X1=esk4_2(esk1_0,X2)|r1_orders_2(esk2_0,esk5_3(esk1_0,X2,X1),esk7_3(esk2_0,X3,esk4_2(esk1_0,X3)))|r1_lattice3(esk2_0,X3,esk6_3(esk2_0,X3,esk4_2(esk1_0,X3)))|r2_yellow_0(esk2_0,X3)|~r1_lattice3(esk2_0,X3,esk5_3(esk1_0,X2,X1))|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X3)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_75, negated_conjecture, (X1=esk4_2(esk1_0,X2)|r1_lattice3(esk2_0,X2,esk5_3(esk1_0,X2,X1))|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_66]), c_0_40])]), c_0_65])).
cnf(c_0_76, plain, (r2_yellow_0(X1,X2)|esk7_3(X1,X2,X3)!=X3|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_77, negated_conjecture, (r1_orders_2(esk2_0,X1,esk4_2(esk1_0,X2))|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_67]), c_0_40])]), c_0_39])).
cnf(c_0_78, negated_conjecture, (r1_lattice3(esk1_0,X1,esk6_3(X2,X1,X3))|r2_yellow_0(X2,X1)|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|esk7_3(X2,X1,X3)!=X3|~r1_lattice3(X2,X1,X3)|~l1_orders_2(X2)|~m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_79, negated_conjecture, (X1=esk8_2(esk1_0,X2)|r1_orders_2(esk2_0,esk11_3(esk2_0,X3,esk8_2(esk1_0,X3)),esk9_3(esk1_0,X2,X1))|r2_lattice3(esk2_0,X3,esk10_3(esk2_0,X3,esk8_2(esk1_0,X3)))|r1_yellow_0(esk2_0,X3)|~r2_lattice3(esk2_0,X3,esk9_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X3)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_80, negated_conjecture, (X1=esk8_2(esk1_0,X2)|r2_lattice3(esk2_0,X2,esk9_3(esk1_0,X2,X1))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_71]), c_0_40])]), c_0_70])).
cnf(c_0_81, plain, (r1_yellow_0(X1,X2)|esk11_3(X1,X2,X3)!=X3|~r1_orders_2(X1,X3,esk10_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_82, negated_conjecture, (r1_orders_2(esk2_0,esk8_2(esk1_0,X1),X2)|~r2_lattice3(esk1_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_72]), c_0_40])]), c_0_47])).
cnf(c_0_83, negated_conjecture, (r2_lattice3(esk1_0,X1,esk10_3(X2,X1,X3))|r1_yellow_0(X2,X1)|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|esk11_3(X2,X1,X3)!=X3|~r2_lattice3(X2,X1,X3)|~l1_orders_2(X2)|~m1_subset_1(esk10_3(X2,X1,X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_84, negated_conjecture, (X1=esk4_2(esk1_0,X2)|r1_orders_2(esk2_0,esk5_3(esk1_0,X2,X1),esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r1_lattice3(esk2_0,X2,esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X2)|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_85, negated_conjecture, (r2_yellow_0(esk2_0,X1)|esk7_3(esk2_0,X1,esk4_2(esk1_0,X2))!=esk4_2(esk1_0,X2)|~r1_lattice3(esk1_0,X2,esk6_3(esk2_0,X1,esk4_2(esk1_0,X2)))|~r1_lattice3(esk2_0,X1,esk4_2(esk1_0,X2))|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_40])]), c_0_39])).
cnf(c_0_86, negated_conjecture, (r1_lattice3(esk1_0,X1,esk6_3(esk2_0,X1,X2))|r2_yellow_0(esk2_0,X1)|esk7_3(esk2_0,X1,X2)!=X2|~r1_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_78]), c_0_40])])).
cnf(c_0_87, plain, (r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|r1_lattice3(X1,X2,esk6_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_88, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r1_lattice3(X1,X2,esk6_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_89, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_90, negated_conjecture, (X1=esk8_2(esk1_0,X2)|r1_orders_2(esk2_0,esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),esk9_3(esk1_0,X2,X1))|r2_lattice3(esk2_0,X2,esk10_3(esk2_0,X2,esk8_2(esk1_0,X2)))|r1_yellow_0(esk2_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_91, negated_conjecture, (r1_yellow_0(esk2_0,X1)|esk11_3(esk2_0,X1,esk8_2(esk1_0,X2))!=esk8_2(esk1_0,X2)|~r2_lattice3(esk1_0,X2,esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)))|~r2_lattice3(esk2_0,X1,esk8_2(esk1_0,X2))|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_40])]), c_0_47])).
cnf(c_0_92, negated_conjecture, (r2_lattice3(esk1_0,X1,esk10_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|esk11_3(esk2_0,X1,X2)!=X2|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk10_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_83]), c_0_40])])).
cnf(c_0_93, plain, (r2_lattice3(X1,X2,esk11_3(X1,X2,X3))|r2_lattice3(X1,X2,esk10_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_94, plain, (m1_subset_1(esk11_3(X1,X2,X3),u1_struct_0(X1))|r2_lattice3(X1,X2,esk10_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_95, plain, (m1_subset_1(esk11_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_96, negated_conjecture, (X1=esk4_2(esk1_0,X2)|r1_orders_2(X3,esk5_3(esk1_0,X2,X1),esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r1_lattice3(esk2_0,X2,esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~l1_orders_2(X3)|~m1_subset_1(esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(X3))|~m1_subset_1(esk5_3(esk1_0,X2,X1),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_84]), c_0_40])]), c_0_65])).
cnf(c_0_97, negated_conjecture, (r2_yellow_0(esk2_0,X1)|esk7_3(esk2_0,X1,esk4_2(esk1_0,X1))!=esk4_2(esk1_0,X1)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_39]), c_0_53])).
cnf(c_0_98, plain, (r1_lattice3(X1,X2,esk6_3(X1,X2,X3))|r1_lattice3(X4,X2,esk7_3(X1,X2,X3))|r2_yellow_0(X1,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~r1_lattice3(X1,X2,X3)|~l1_orders_2(X4)|~l1_orders_2(X1)|~m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_87]), c_0_88])).
cnf(c_0_99, plain, (r1_lattice3(X1,X2,esk6_3(X3,X2,X4))|r2_yellow_0(X3,X2)|m1_subset_1(esk7_3(X3,X2,X4),u1_struct_0(X3))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk6_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_88]), c_0_89])).
cnf(c_0_100, negated_conjecture, (X1=esk8_2(esk1_0,X2)|r1_orders_2(X3,esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),esk9_3(esk1_0,X2,X1))|r2_lattice3(esk2_0,X2,esk10_3(esk2_0,X2,esk8_2(esk1_0,X2)))|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X3)|~m1_subset_1(esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(X3))|~m1_subset_1(esk9_3(esk1_0,X2,X1),u1_struct_0(X3))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_90]), c_0_40])]), c_0_70])).
cnf(c_0_101, negated_conjecture, (r1_yellow_0(esk2_0,X1)|esk11_3(esk2_0,X1,esk8_2(esk1_0,X1))!=esk8_2(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_47]), c_0_59])).
cnf(c_0_102, plain, (r2_lattice3(X1,X2,esk10_3(X1,X2,X3))|r2_lattice3(X4,X2,esk11_3(X1,X2,X3))|r1_yellow_0(X1,X2)|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(X4),u1_orders_2(X4))|~r2_lattice3(X1,X2,X3)|~l1_orders_2(X4)|~l1_orders_2(X1)|~m1_subset_1(esk11_3(X1,X2,X3),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_93]), c_0_94])).
cnf(c_0_103, plain, (r2_lattice3(X1,X2,esk10_3(X3,X2,X4))|r1_yellow_0(X3,X2)|m1_subset_1(esk11_3(X3,X2,X4),u1_struct_0(X3))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk10_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_94]), c_0_95])).
cnf(c_0_104, plain, (X3=esk4_2(X1,X2)|~r1_orders_2(X1,esk5_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_105, negated_conjecture, (X1=esk4_2(esk1_0,X2)|r1_orders_2(esk1_0,esk5_3(esk1_0,X2,X1),esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r1_lattice3(esk2_0,X2,esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X2)|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_17]), c_0_34]), c_0_15]), c_0_34]), c_0_34])]), c_0_65])).
cnf(c_0_106, negated_conjecture, (r2_yellow_0(esk2_0,X1)|esk7_3(esk2_0,X1,esk4_2(esk1_0,X1))!=esk4_2(esk1_0,X1)|~r2_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_57]), c_0_40])]), c_0_39]), c_0_53])).
cnf(c_0_107, negated_conjecture, (r1_lattice3(esk1_0,X1,esk7_3(X2,X1,X3))|r1_lattice3(X2,X1,esk6_3(X2,X1,X3))|r2_yellow_0(X2,X1)|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r1_lattice3(X2,X1,X3)|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X2,X1,X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_108, plain, (m1_subset_1(esk7_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_109, negated_conjecture, (r1_lattice3(esk1_0,X1,esk6_3(X2,X1,X3))|r2_yellow_0(X2,X1)|m1_subset_1(esk7_3(X2,X1,X3),u1_struct_0(X2))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r1_lattice3(X2,X1,X3)|~l1_orders_2(X2)|~m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_110, plain, (X2=esk8_2(X1,X3)|~r1_orders_2(X1,X2,esk9_3(X1,X3,X2))|~r2_lattice3(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_yellow_0(X1,X3)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_111, negated_conjecture, (X1=esk8_2(esk1_0,X2)|r1_orders_2(esk1_0,esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),esk9_3(esk1_0,X2,X1))|r2_lattice3(esk2_0,X2,esk10_3(esk2_0,X2,esk8_2(esk1_0,X2)))|r1_yellow_0(esk2_0,X2)|~r2_lattice3(esk1_0,X2,X1)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_17]), c_0_34]), c_0_15]), c_0_34]), c_0_34])]), c_0_70])).
cnf(c_0_112, negated_conjecture, (r1_yellow_0(esk2_0,X1)|esk11_3(esk2_0,X1,esk8_2(esk1_0,X1))!=esk8_2(esk1_0,X1)|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_63]), c_0_40])]), c_0_47]), c_0_59])).
cnf(c_0_113, negated_conjecture, (r2_lattice3(esk1_0,X1,esk11_3(X2,X1,X3))|r2_lattice3(X2,X1,esk10_3(X2,X1,X3))|r1_yellow_0(X2,X1)|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r2_lattice3(X2,X1,X3)|~l1_orders_2(X2)|~m1_subset_1(esk11_3(X2,X1,X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_114, plain, (m1_subset_1(esk11_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk10_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_115, negated_conjecture, (r2_lattice3(esk1_0,X1,esk10_3(X2,X1,X3))|r1_yellow_0(X2,X1)|m1_subset_1(esk11_3(X2,X1,X3),u1_struct_0(X2))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r2_lattice3(X2,X1,X3)|~l1_orders_2(X2)|~m1_subset_1(esk10_3(X2,X1,X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_116, negated_conjecture, (r1_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)))|r2_yellow_0(esk2_0,X1)|~r1_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,esk4_2(esk1_0,X1)))|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk7_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_15]), c_0_34])]), c_0_106])).
cnf(c_0_117, negated_conjecture, (r1_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,X2))|r1_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,X2))|r2_yellow_0(esk2_0,X1)|~r1_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_107]), c_0_40])])).
cnf(c_0_118, negated_conjecture, (r2_yellow_0(esk2_0,X1)|m1_subset_1(esk7_3(esk2_0,X1,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))|~r1_lattice3(esk1_0,X2,esk6_3(esk2_0,X1,esk4_2(esk1_0,X2)))|~r1_lattice3(esk2_0,X1,esk4_2(esk1_0,X2))|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_77]), c_0_40])]), c_0_39])).
cnf(c_0_119, negated_conjecture, (r1_lattice3(esk1_0,X1,esk6_3(esk2_0,X1,X2))|r2_yellow_0(esk2_0,X1)|m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r1_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_109]), c_0_40])])).
cnf(c_0_120, negated_conjecture, (r2_lattice3(esk2_0,X1,esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk1_0,X1,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_15]), c_0_34])]), c_0_112])).
cnf(c_0_121, negated_conjecture, (r2_lattice3(esk2_0,X1,esk10_3(esk2_0,X1,X2))|r2_lattice3(esk1_0,X1,esk11_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk11_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_113]), c_0_40])])).
cnf(c_0_122, negated_conjecture, (r1_yellow_0(esk2_0,X1)|m1_subset_1(esk11_3(esk2_0,X1,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))|~r2_lattice3(esk1_0,X2,esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)))|~r2_lattice3(esk2_0,X1,esk8_2(esk1_0,X2))|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_82]), c_0_40])]), c_0_47])).
cnf(c_0_123, negated_conjecture, (r2_lattice3(esk1_0,X1,esk10_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk11_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk10_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_115]), c_0_40])])).
cnf(c_0_124, negated_conjecture, (r1_lattice3(esk2_0,X1,esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)))|r2_yellow_0(esk2_0,X1)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk7_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_39]), c_0_53])).
cnf(c_0_125, negated_conjecture, (r2_yellow_0(esk2_0,X1)|m1_subset_1(esk7_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119]), c_0_39]), c_0_53])).
cnf(c_0_126, negated_conjecture, (r2_lattice3(esk2_0,X1,esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_47]), c_0_59])).
cnf(c_0_127, negated_conjecture, (r1_yellow_0(esk2_0,X1)|m1_subset_1(esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_47]), c_0_59])).
cnf(c_0_128, plain, (r1_orders_2(X2,X1,esk7_3(X2,X3,X4))|r2_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r1_orders_2(X2,esk6_3(X2,X3,X4),X4)|~r1_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_129, negated_conjecture, (r1_lattice3(X1,X2,esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_124]), c_0_40])]), c_0_125])).
cnf(c_0_130, plain, (r1_orders_2(X2,esk11_3(X2,X3,X4),X1)|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r1_orders_2(X2,X4,esk10_3(X2,X3,X4))|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_131, negated_conjecture, (r2_lattice3(X1,X2,esk10_3(esk2_0,X2,esk8_2(esk1_0,X2)))|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk10_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(esk10_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_126]), c_0_40])]), c_0_127])).
cnf(c_0_132, negated_conjecture, (r1_orders_2(esk2_0,X1,esk7_3(esk2_0,X2,esk4_2(esk1_0,X3)))|r2_yellow_0(esk2_0,X2)|~r1_lattice3(esk1_0,X3,esk6_3(esk2_0,X2,esk4_2(esk1_0,X3)))|~r1_lattice3(esk2_0,X2,esk4_2(esk1_0,X3))|~r1_lattice3(esk2_0,X2,X1)|~r2_yellow_0(esk1_0,X3)|~m1_subset_1(esk6_3(esk2_0,X2,esk4_2(esk1_0,X3)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_77]), c_0_40])]), c_0_39])).
cnf(c_0_133, negated_conjecture, (r1_lattice3(esk1_0,X1,esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)))|r2_yellow_0(esk2_0,X1)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_134, plain, (r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk6_3(X1,X2,X3),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_135, plain, (r1_orders_2(X2,X1,esk7_3(X2,X3,X4))|m1_subset_1(esk6_3(X2,X3,X4),u1_struct_0(X2))|r2_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r1_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_136, negated_conjecture, (r1_orders_2(esk2_0,esk11_3(esk2_0,X1,esk8_2(esk1_0,X2)),X3)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk1_0,X2,esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)))|~r2_lattice3(esk2_0,X1,esk8_2(esk1_0,X2))|~r2_lattice3(esk2_0,X1,X3)|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_82]), c_0_40])]), c_0_47])).
cnf(c_0_137, negated_conjecture, (r2_lattice3(esk1_0,X1,esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_138, plain, (r2_lattice3(X1,X2,esk11_3(X1,X2,X3))|r1_yellow_0(X1,X2)|~r1_orders_2(X1,X3,esk10_3(X1,X2,X3))|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_139, plain, (r1_orders_2(X2,esk11_3(X2,X3,X4),X1)|m1_subset_1(esk10_3(X2,X3,X4),u1_struct_0(X2))|r1_yellow_0(X2,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r2_lattice3(X2,X3,X1)|~r2_lattice3(X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_140, negated_conjecture, (r1_orders_2(esk2_0,X1,esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X2)|~r1_lattice3(esk2_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_53])).
cnf(c_0_141, negated_conjecture, (r1_lattice3(esk2_0,X1,esk7_3(esk2_0,X1,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X1)|~r1_lattice3(esk1_0,X2,esk6_3(esk2_0,X1,esk4_2(esk1_0,X2)))|~r1_lattice3(esk2_0,X1,esk4_2(esk1_0,X2))|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_77]), c_0_40])]), c_0_39])).
cnf(c_0_142, plain, (r1_orders_2(X1,X2,esk7_3(X3,X4,X5))|r2_yellow_0(X3,X4)|m1_subset_1(esk6_3(X3,X4,X5),u1_struct_0(X3))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_lattice3(X3,X4,X5)|~r1_lattice3(X3,X4,X2)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk7_3(X3,X4,X5),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X3))|~m1_subset_1(X5,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_135]), c_0_89])).
cnf(c_0_143, plain, (r1_lattice3(X1,X2,esk7_3(X1,X2,X3))|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|r2_yellow_0(X1,X2)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_144, negated_conjecture, (r1_orders_2(esk2_0,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)),X2)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_137]), c_0_59])).
cnf(c_0_145, negated_conjecture, (r2_lattice3(esk2_0,X1,esk11_3(esk2_0,X1,esk8_2(esk1_0,X2)))|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk1_0,X2,esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)))|~r2_lattice3(esk2_0,X1,esk8_2(esk1_0,X2))|~r1_yellow_0(esk1_0,X2)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_82]), c_0_40])]), c_0_47])).
cnf(c_0_146, plain, (r1_orders_2(X1,esk11_3(X2,X3,X4),X5)|r1_yellow_0(X2,X3)|m1_subset_1(esk10_3(X2,X3,X4),u1_struct_0(X2))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X2,X3,X4)|~r2_lattice3(X2,X3,X5)|~l1_orders_2(X1)|~l1_orders_2(X2)|~m1_subset_1(esk11_3(X2,X3,X4),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_139]), c_0_95])).
cnf(c_0_147, plain, (r2_lattice3(X1,X2,esk11_3(X1,X2,X3))|m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X1))|r1_yellow_0(X1,X2)|~r2_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_148, negated_conjecture, (r1_orders_2(X1,X2,esk7_3(esk2_0,X3,esk4_2(esk1_0,X3)))|r2_yellow_0(esk2_0,X3)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_lattice3(esk2_0,X3,X2)|~r2_yellow_0(esk1_0,X3)|~l1_orders_2(X1)|~m1_subset_1(esk6_3(esk2_0,X3,esk4_2(esk1_0,X3)),u1_struct_0(esk2_0))|~m1_subset_1(esk7_3(esk2_0,X3,esk4_2(esk1_0,X3)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_140]), c_0_40])]), c_0_125])).
cnf(c_0_149, negated_conjecture, (r1_lattice3(esk2_0,X1,esk7_3(esk2_0,X1,esk4_2(esk1_0,X1)))|r2_yellow_0(esk2_0,X1)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_133]), c_0_53])).
cnf(c_0_150, negated_conjecture, (r1_orders_2(esk1_0,X1,esk7_3(X2,X3,X4))|r2_yellow_0(X2,X3)|m1_subset_1(esk6_3(X2,X3,X4),u1_struct_0(X2))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r1_lattice3(X2,X3,X4)|~r1_lattice3(X2,X3,X1)|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X2,X3,X4),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_17]), c_0_34]), c_0_15]), c_0_34]), c_0_34])])).
cnf(c_0_151, plain, (r1_lattice3(X1,X2,esk7_3(X3,X2,X4))|r2_yellow_0(X3,X2)|m1_subset_1(esk6_3(X3,X2,X4),u1_struct_0(X3))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk7_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_143]), c_0_89])).
cnf(c_0_152, negated_conjecture, (r1_orders_2(X1,esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),X3)|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(esk2_0,X2,X3)|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk10_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_144]), c_0_40])]), c_0_127])).
cnf(c_0_153, negated_conjecture, (r2_lattice3(esk2_0,X1,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_137]), c_0_59])).
cnf(c_0_154, negated_conjecture, (r1_orders_2(esk1_0,esk11_3(X1,X2,X3),X4)|r1_yellow_0(X1,X2)|m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X1))|g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r2_lattice3(X1,X2,X3)|~r2_lattice3(X1,X2,X4)|~l1_orders_2(X1)|~m1_subset_1(esk11_3(X1,X2,X3),u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(esk2_0))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_17]), c_0_34]), c_0_15]), c_0_34]), c_0_34])])).
cnf(c_0_155, plain, (r2_lattice3(X1,X2,esk11_3(X3,X2,X4))|r1_yellow_0(X3,X2)|m1_subset_1(esk10_3(X3,X2,X4),u1_struct_0(X3))|g1_orders_2(u1_struct_0(X3),u1_orders_2(X3))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_lattice3(X3,X2,X4)|~l1_orders_2(X1)|~l1_orders_2(X3)|~m1_subset_1(esk11_3(X3,X2,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_147]), c_0_95])).
cnf(c_0_156, negated_conjecture, (r1_orders_2(esk1_0,X1,esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X2)|~r1_lattice3(esk2_0,X2,X1)|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_17]), c_0_34]), c_0_15]), c_0_34]), c_0_34])]), c_0_125])).
cnf(c_0_157, negated_conjecture, (r1_lattice3(X1,X2,esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)))|r2_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r2_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk6_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(esk7_3(esk2_0,X2,esk4_2(esk1_0,X2)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_149]), c_0_40])]), c_0_125])).
cnf(c_0_158, negated_conjecture, (r1_orders_2(esk1_0,X1,esk7_3(esk2_0,X2,X3))|r2_yellow_0(esk2_0,X2)|m1_subset_1(esk6_3(esk2_0,X2,X3),u1_struct_0(esk2_0))|~r1_lattice3(esk2_0,X2,X3)|~r1_lattice3(esk2_0,X2,X1)|~m1_subset_1(esk7_3(esk2_0,X2,X3),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_150]), c_0_40])])).
cnf(c_0_159, negated_conjecture, (r1_lattice3(esk1_0,X1,esk7_3(X2,X1,X3))|r2_yellow_0(X2,X1)|m1_subset_1(esk6_3(X2,X1,X3),u1_struct_0(X2))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r1_lattice3(X2,X1,X3)|~l1_orders_2(X2)|~m1_subset_1(esk7_3(X2,X1,X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_160, negated_conjecture, (r1_orders_2(esk1_0,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)),X2)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_17]), c_0_34]), c_0_15]), c_0_34]), c_0_34])]), c_0_127])).
cnf(c_0_161, negated_conjecture, (r2_lattice3(X1,X2,esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)))|r1_yellow_0(esk2_0,X2)|g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))!=g1_orders_2(u1_struct_0(X1),u1_orders_2(X1))|~r1_yellow_0(esk1_0,X2)|~l1_orders_2(X1)|~m1_subset_1(esk10_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(esk2_0))|~m1_subset_1(esk11_3(esk2_0,X2,esk8_2(esk1_0,X2)),u1_struct_0(X1))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_153]), c_0_40])]), c_0_127])).
cnf(c_0_162, negated_conjecture, (r1_orders_2(esk1_0,esk11_3(esk2_0,X1,X2),X3)|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk10_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,X1,X2)|~r2_lattice3(esk2_0,X1,X3)|~m1_subset_1(esk11_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_154]), c_0_40])])).
cnf(c_0_163, negated_conjecture, (r2_lattice3(esk1_0,X1,esk11_3(X2,X1,X3))|r1_yellow_0(X2,X1)|m1_subset_1(esk10_3(X2,X1,X3),u1_struct_0(X2))|g1_orders_2(u1_struct_0(X2),u1_orders_2(X2))!=g1_orders_2(u1_struct_0(esk2_0),u1_orders_2(esk2_0))|~r2_lattice3(X2,X1,X3)|~l1_orders_2(X2)|~m1_subset_1(esk11_3(X2,X1,X3),u1_struct_0(esk2_0))|~m1_subset_1(X3,u1_struct_0(X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_17]), c_0_34]), c_0_15]), c_0_34])])).
cnf(c_0_164, negated_conjecture, (esk7_3(esk2_0,X1,esk4_2(esk1_0,X1))=esk4_2(esk1_0,X2)|r2_yellow_0(esk2_0,X1)|~r1_lattice3(esk2_0,X1,esk5_3(esk1_0,X2,esk7_3(esk2_0,X1,esk4_2(esk1_0,X1))))|~r1_lattice3(esk1_0,X2,esk7_3(esk2_0,X1,esk4_2(esk1_0,X1)))|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_3(esk1_0,X2,esk7_3(esk2_0,X1,esk4_2(esk1_0,X1))),u1_struct_0(esk2_0))|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_156]), c_0_15]), c_0_34])]), c_0_125])).
cnf(c_0_165, negated_conjecture, (r1_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,esk4_2(esk1_0,X1)))|r2_yellow_0(esk2_0,X1)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_17]), c_0_34]), c_0_15]), c_0_34])]), c_0_125])).
cnf(c_0_166, negated_conjecture, (esk7_3(esk2_0,X1,X2)=esk4_2(esk1_0,X3)|r2_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r1_lattice3(esk2_0,X1,esk5_3(esk1_0,X3,esk7_3(esk2_0,X1,X2)))|~r1_lattice3(esk1_0,X3,esk7_3(esk2_0,X1,X2))|~r1_lattice3(esk2_0,X1,X2)|~r2_yellow_0(esk1_0,X3)|~m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_158]), c_0_15]), c_0_34])]), c_0_65])).
cnf(c_0_167, negated_conjecture, (r1_lattice3(esk1_0,X1,esk7_3(esk2_0,X1,X2))|r2_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r1_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_159]), c_0_40])])).
cnf(c_0_168, negated_conjecture, (esk11_3(esk2_0,X1,esk8_2(esk1_0,X1))=esk8_2(esk1_0,X2)|r1_yellow_0(esk2_0,X1)|~r2_lattice3(esk2_0,X1,esk9_3(esk1_0,X2,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1))))|~r2_lattice3(esk1_0,X2,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)))|~r1_yellow_0(esk1_0,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk9_3(esk1_0,X2,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1))),u1_struct_0(esk2_0))|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_160]), c_0_15]), c_0_34])]), c_0_127])).
cnf(c_0_169, negated_conjecture, (r2_lattice3(esk1_0,X1,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1)))|r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_17]), c_0_34]), c_0_15]), c_0_34])]), c_0_127])).
cnf(c_0_170, negated_conjecture, (esk11_3(esk2_0,X1,X2)=esk8_2(esk1_0,X3)|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk10_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,X1,esk9_3(esk1_0,X3,esk11_3(esk2_0,X1,X2)))|~r2_lattice3(esk1_0,X3,esk11_3(esk2_0,X1,X2))|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X3)|~m1_subset_1(esk11_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_162]), c_0_15]), c_0_34])]), c_0_70])).
cnf(c_0_171, negated_conjecture, (r2_lattice3(esk1_0,X1,esk11_3(esk2_0,X1,X2))|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk10_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,X1,X2)|~m1_subset_1(esk11_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_163]), c_0_40])])).
cnf(c_0_172, negated_conjecture, (r2_yellow_0(esk2_0,X1)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_3(esk1_0,X1,esk7_3(esk2_0,X1,esk4_2(esk1_0,X1))),u1_struct_0(esk2_0))|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_75]), c_0_125]), c_0_165]), c_0_106])).
cnf(c_0_173, negated_conjecture, (esk7_3(esk2_0,X1,X2)=esk4_2(esk1_0,X1)|r2_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r1_lattice3(esk2_0,X1,X2)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk7_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_75]), c_0_167])).
cnf(c_0_174, negated_conjecture, (r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk9_3(esk1_0,X1,esk11_3(esk2_0,X1,esk8_2(esk1_0,X1))),u1_struct_0(esk2_0))|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_80]), c_0_127]), c_0_169]), c_0_112])).
cnf(c_0_175, negated_conjecture, (esk11_3(esk2_0,X1,X2)=esk8_2(esk1_0,X1)|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk10_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk11_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_80]), c_0_171])).
cnf(c_0_176, negated_conjecture, (r2_yellow_0(esk2_0,X1)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk6_3(esk2_0,X1,esk4_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_65]), c_0_125]), c_0_165]), c_0_106])).
cnf(c_0_177, negated_conjecture, (esk7_3(esk2_0,X1,X2)=esk4_2(esk1_0,X1)|r2_yellow_0(esk2_0,X1)|m1_subset_1(esk6_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r1_lattice3(esk2_0,X1,X2)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_89]), c_0_40])])).
cnf(c_0_178, negated_conjecture, (r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(esk10_3(esk2_0,X1,esk8_2(esk1_0,X1)),u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_70]), c_0_127]), c_0_169]), c_0_112])).
cnf(c_0_179, negated_conjecture, (esk11_3(esk2_0,X1,X2)=esk8_2(esk1_0,X1)|r1_yellow_0(esk2_0,X1)|m1_subset_1(esk10_3(esk2_0,X1,X2),u1_struct_0(esk2_0))|~r2_lattice3(esk2_0,X1,X2)|~r1_yellow_0(esk1_0,X1)|~m1_subset_1(X2,u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_95]), c_0_40])])).
cnf(c_0_180, negated_conjecture, (r2_yellow_0(esk2_0,X1)|~r2_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_39]), c_0_53]), c_0_106])).
cnf(c_0_181, negated_conjecture, (r2_yellow_0(esk1_0,esk3_0)|r1_yellow_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_182, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)|~r2_yellow_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_183, negated_conjecture, (r2_yellow_0(esk1_0,esk3_0)|~r1_yellow_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_184, negated_conjecture, (~r2_yellow_0(esk2_0,esk3_0)|~r1_yellow_0(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_185, negated_conjecture, (r1_yellow_0(esk2_0,X1)|~r1_yellow_0(esk1_0,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179]), c_0_47]), c_0_59]), c_0_112])).
cnf(c_0_186, negated_conjecture, (r1_yellow_0(esk1_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_181]), c_0_182])).
cnf(c_0_187, negated_conjecture, (~r1_yellow_0(esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_183]), c_0_184])).
cnf(c_0_188, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_186]), c_0_187]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 189
# Proof object clause steps            : 174
# Proof object formula steps           : 15
# Proof object conjectures             : 115
# Proof object clause conjectures      : 112
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 49
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 121
# Proof object simplifying inferences  : 305
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 50
# Processed clauses                    : 2859
# ...of these trivial                  : 171
# ...subsumed                          : 929
# ...remaining for further processing  : 1759
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 445
# Backward-rewritten                   : 7
# Generated clauses                    : 5328
# ...of the previous two non-trivial   : 4315
# Contextual simplify-reflections      : 368
# Paramodulations                      : 5220
# Factorizations                       : 0
# Equation resolutions                 : 110
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1253
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 1246
# Current number of unprocessed clauses: 1527
# ...number of literals in the above   : 19364
# Current number of archived formulas  : 0
# Current number of archived clauses   : 502
# Clause-clause subsumption calls (NU) : 973781
# Rec. Clause-clause subsumption calls : 13286
# Non-unit clause-clause subsumptions  : 1742
# Unit Clause-clause subsumption calls : 2790
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 357103

# -------------------------------------------------
# User time                : 0.514 s
# System time              : 0.000 s
# Total time               : 0.514 s
# Maximum resident set size: 28784 pages
