# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t7_boole)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', d3_tarski)).
fof(t135_zfmisc_1, axiom, ![X1, X2]:((r1_tarski(X1,k2_zfmisc_1(X1,X2))|r1_tarski(X1,k2_zfmisc_1(X2,X1)))=>X1=k1_xboole_0), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t135_zfmisc_1)).
fof(rc1_connsp_1, axiom, ![X1]:(l1_pre_topc(X1)=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&v1_xboole_0(X2))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', rc1_connsp_1)).
fof(fc4_zfmisc_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>v1_xboole_0(k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', fc4_zfmisc_1)).
fof(t96_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3))))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t96_tmap_1)).
fof(t25_relset_1, axiom, ![X1, X2]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t25_relset_1)).
fof(cc1_funct_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_funct_1(X1)), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', cc1_funct_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t5_subset)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', cc1_relset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', d1_zfmisc_1)).
fof(d9_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,X2)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(r2_funct_2(X1,X2,X3,X4)<=>![X5]:(m1_subset_1(X5,X1)=>k1_funct_1(X3,X5)=k1_funct_1(X4,X5))))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', d9_funct_2)).
fof(dt_k4_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_funct_1(k4_tmap_1(X1,X2))&(v1_funct_2(k4_tmap_1(X1,X2),u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(k4_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', dt_k4_tmap_1)).
fof(d4_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2, X3]:((r2_hidden(X2,k1_relat_1(X1))=>(X3=k1_funct_1(X1,X2)<=>r2_hidden(k4_tarski(X2,X3),X1)))&(~(r2_hidden(X2,k1_relat_1(X1)))=>(X3=k1_funct_1(X1,X2)<=>X3=k1_xboole_0)))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', d4_funct_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t1_subset)).
fof(t55_pre_topc, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1))))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t55_pre_topc)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', d2_subset_1)).
fof(t8_grfunc_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r1_tarski(X1,X2)<=>(r1_tarski(k1_relat_1(X1),k1_relat_1(X2))&![X3]:(r2_hidden(X3,k1_relat_1(X1))=>k1_funct_1(X1,X3)=k1_funct_1(X2,X3)))))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t8_grfunc_1)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', cc3_relset_1)).
fof(t95_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,u1_struct_0(X2))=>k1_funct_1(k4_tmap_1(X1,X2),X3)=X3)))), file('data/MPTP/proofs/def892b19d1e48da99bab8631cc53d4a.E_input', t95_tmap_1)).
fof(c_0_20, plain, ![X67, X68]:(~r2_hidden(X67,X68)|~v1_xboole_0(X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_21, plain, ![X14, X15, X16, X17, X18]:((~r1_tarski(X14,X15)|(~r2_hidden(X16,X14)|r2_hidden(X16,X15)))&((r2_hidden(esk4_2(X17,X18),X17)|r1_tarski(X17,X18))&(~r2_hidden(esk4_2(X17,X18),X18)|r1_tarski(X17,X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_22, plain, ![X43, X44]:((~r1_tarski(X43,k2_zfmisc_1(X43,X44))|X43=k1_xboole_0)&(~r1_tarski(X43,k2_zfmisc_1(X44,X43))|X43=k1_xboole_0)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t135_zfmisc_1])])])).
cnf(c_0_23, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (X1=k1_xboole_0|~r1_tarski(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_27, plain, ![X69]:((m1_subset_1(esk8_1(X69),k1_zfmisc_1(u1_struct_0(X69)))|~l1_pre_topc(X69))&(v1_xboole_0(esk8_1(X69))|~l1_pre_topc(X69))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[rc1_connsp_1])])])])).
fof(c_0_28, plain, ![X45, X46]:(~v1_xboole_0(X45)|v1_xboole_0(k2_zfmisc_1(X45,X46))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_zfmisc_1])])).
cnf(c_0_29, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_30, plain, (v1_xboole_0(esk8_1(X1))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_31, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))))=>(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,u1_struct_0(X2))=>X4=k1_funct_1(X3,X4)))=>r2_funct_2(u1_struct_0(X2),u1_struct_0(X1),k4_tmap_1(X1,X2),X3)))))), inference(assume_negation,[status(cth)],[t96_tmap_1])).
fof(c_0_32, plain, ![X12, X13]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X12,X13))), inference(variable_rename,[status(thm)],[t25_relset_1])).
cnf(c_0_33, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (esk8_1(X1)=k1_xboole_0|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_35, negated_conjecture, ![X9]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))))&((~m1_subset_1(X9,u1_struct_0(esk1_0))|(~r2_hidden(X9,u1_struct_0(esk2_0))|X9=k1_funct_1(esk3_0,X9)))&~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_31])])])])])).
fof(c_0_36, plain, ![X23]:(~v1_xboole_0(X23)|v1_funct_1(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_1])])).
cnf(c_0_37, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_29, c_0_33])).
cnf(c_0_39, plain, (v1_xboole_0(k1_xboole_0)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_30, c_0_34])).
cnf(c_0_40, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (v1_funct_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_42, plain, ![X55, X56, X57]:(~r2_hidden(X55,X56)|(~m1_subset_1(X56,k1_zfmisc_1(X57))|~v1_xboole_0(X57))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_43, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_44, negated_conjecture, (v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_45, plain, ![X50, X51, X52]:(~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X50,X51)))|v1_relat_1(X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_46, plain, (v1_funct_1(esk8_1(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_41, c_0_30])).
fof(c_0_47, plain, ![X60, X61, X62, X63, X64, X65]:(((~r2_hidden(X62,X61)|r1_tarski(X62,X60)|X61!=k1_zfmisc_1(X60))&(~r1_tarski(X63,X60)|r2_hidden(X63,X61)|X61!=k1_zfmisc_1(X60)))&((~r2_hidden(esk7_2(X64,X65),X65)|~r1_tarski(esk7_2(X64,X65),X64)|X65=k1_zfmisc_1(X64))&(r2_hidden(esk7_2(X64,X65),X65)|r1_tarski(esk7_2(X64,X65),X64)|X65=k1_zfmisc_1(X64)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_48, plain, ![X31, X32, X33, X34, X35]:((~r2_funct_2(X31,X32,X33,X34)|(~m1_subset_1(X35,X31)|k1_funct_1(X33,X35)=k1_funct_1(X34,X35))|(~v1_funct_1(X34)|(~v1_funct_2(X34,X31,X32)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))))|(~v1_funct_1(X33)|(~v1_funct_2(X33,X31,X32)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32))))))&((m1_subset_1(esk6_4(X31,X32,X33,X34),X31)|r2_funct_2(X31,X32,X33,X34)|(~v1_funct_1(X34)|(~v1_funct_2(X34,X31,X32)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))))|(~v1_funct_1(X33)|(~v1_funct_2(X33,X31,X32)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32))))))&(k1_funct_1(X33,esk6_4(X31,X32,X33,X34))!=k1_funct_1(X34,esk6_4(X31,X32,X33,X34))|r2_funct_2(X31,X32,X33,X34)|(~v1_funct_1(X34)|(~v1_funct_2(X34,X31,X32)|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))))|(~v1_funct_1(X33)|(~v1_funct_2(X33,X31,X32)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X31,X32)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_funct_2])])])])])).
fof(c_0_49, plain, ![X58, X59]:((v1_funct_1(k4_tmap_1(X58,X59))|(v2_struct_0(X58)|(~v2_pre_topc(X58)|~l1_pre_topc(X58))|~m1_pre_topc(X59,X58)))&((v1_funct_2(k4_tmap_1(X58,X59),u1_struct_0(X59),u1_struct_0(X58))|(v2_struct_0(X58)|(~v2_pre_topc(X58)|~l1_pre_topc(X58))|~m1_pre_topc(X59,X58)))&(m1_subset_1(k4_tmap_1(X58,X59),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X59),u1_struct_0(X58))))|(v2_struct_0(X58)|(~v2_pre_topc(X58)|~l1_pre_topc(X58))|~m1_pre_topc(X59,X58))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_tmap_1])])])])).
cnf(c_0_50, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
fof(c_0_52, plain, ![X37, X38, X39]:(((X39!=k1_funct_1(X37,X38)|r2_hidden(k4_tarski(X38,X39),X37)|~r2_hidden(X38,k1_relat_1(X37))|(~v1_relat_1(X37)|~v1_funct_1(X37)))&(~r2_hidden(k4_tarski(X38,X39),X37)|X39=k1_funct_1(X37,X38)|~r2_hidden(X38,k1_relat_1(X37))|(~v1_relat_1(X37)|~v1_funct_1(X37))))&((X39!=k1_funct_1(X37,X38)|X39=k1_xboole_0|r2_hidden(X38,k1_relat_1(X37))|(~v1_relat_1(X37)|~v1_funct_1(X37)))&(X39!=k1_xboole_0|X39=k1_funct_1(X37,X38)|r2_hidden(X38,k1_relat_1(X37))|(~v1_relat_1(X37)|~v1_funct_1(X37))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_funct_1])])])])])).
cnf(c_0_53, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, plain, (v1_funct_1(k1_xboole_0)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_46, c_0_34])).
fof(c_0_55, plain, ![X10, X11]:(~r2_hidden(X10,X11)|m1_subset_1(X10,X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_56, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, plain, (m1_subset_1(esk6_4(X1,X2,X3,X4),X1)|r2_funct_2(X1,X2,X3,X4)|~v1_funct_1(X4)|~v1_funct_2(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_59, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_60, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_61, plain, (v1_funct_2(k4_tmap_1(X1,X2),u1_struct_0(X2),u1_struct_0(X1))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_62, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_63, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_64, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_65, plain, (m1_subset_1(k4_tmap_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_66, plain, (v1_funct_1(k4_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_67, negated_conjecture, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_44])])).
cnf(c_0_68, plain, (r2_hidden(k4_tarski(X3,X1),X2)|X1!=k1_funct_1(X2,X3)|~r2_hidden(X3,k1_relat_1(X2))|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_69, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_53, c_0_37])).
cnf(c_0_70, negated_conjecture, (v1_funct_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_54, c_0_40])).
cnf(c_0_71, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_72, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_56])).
fof(c_0_73, plain, ![X24, X25, X26]:(v2_struct_0(X24)|~l1_pre_topc(X24)|(v2_struct_0(X25)|~m1_pre_topc(X25,X24)|(~m1_subset_1(X26,u1_struct_0(X25))|m1_subset_1(X26,u1_struct_0(X24))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_pre_topc])])])])).
cnf(c_0_74, negated_conjecture, (r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk3_0)|m1_subset_1(esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),X1,esk3_0),u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk1_0))|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_60])])).
cnf(c_0_75, negated_conjecture, (v1_funct_2(k4_tmap_1(esk1_0,esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_40]), c_0_63])]), c_0_64])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(k4_tmap_1(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk1_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_62]), c_0_40]), c_0_63])]), c_0_64])).
cnf(c_0_77, negated_conjecture, (v1_funct_1(k4_tmap_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_62]), c_0_40]), c_0_63])]), c_0_64])).
cnf(c_0_78, negated_conjecture, (~r2_funct_2(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_79, negated_conjecture, (X1!=k1_funct_1(k1_xboole_0,X2)|~r2_hidden(X2,k1_relat_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70])])).
cnf(c_0_80, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
fof(c_0_81, plain, ![X53, X54]:(((~m1_subset_1(X54,X53)|r2_hidden(X54,X53)|v1_xboole_0(X53))&(~r2_hidden(X54,X53)|m1_subset_1(X54,X53)|v1_xboole_0(X53)))&((~m1_subset_1(X54,X53)|v1_xboole_0(X54)|~v1_xboole_0(X53))&(~v1_xboole_0(X54)|m1_subset_1(X54,X53)|~v1_xboole_0(X53)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_82, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_83, negated_conjecture, (m1_subset_1(esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_76]), c_0_77])]), c_0_78])).
cnf(c_0_84, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_85, negated_conjecture, (~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(er,[status(thm)],[c_0_79])).
cnf(c_0_86, plain, (X1!=k1_funct_1(X2,X3)|~v1_xboole_0(X2)|~v1_relat_1(X2)|~r2_hidden(X3,k1_relat_1(X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_68]), c_0_41])).
cnf(c_0_87, plain, (X1=k1_funct_1(X2,X3)|r2_hidden(X3,k1_relat_1(X2))|X1!=k1_xboole_0|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_88, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_80, c_0_26])).
cnf(c_0_89, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_90, negated_conjecture, (m1_subset_1(esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(X1))|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])).
cnf(c_0_91, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_92, plain, ![X27, X28, X29]:(((r1_tarski(k1_relat_1(X27),k1_relat_1(X28))|~r1_tarski(X27,X28)|(~v1_relat_1(X28)|~v1_funct_1(X28))|(~v1_relat_1(X27)|~v1_funct_1(X27)))&(~r2_hidden(X29,k1_relat_1(X27))|k1_funct_1(X27,X29)=k1_funct_1(X28,X29)|~r1_tarski(X27,X28)|(~v1_relat_1(X28)|~v1_funct_1(X28))|(~v1_relat_1(X27)|~v1_funct_1(X27))))&((r2_hidden(esk5_2(X27,X28),k1_relat_1(X27))|~r1_tarski(k1_relat_1(X27),k1_relat_1(X28))|r1_tarski(X27,X28)|(~v1_relat_1(X28)|~v1_funct_1(X28))|(~v1_relat_1(X27)|~v1_funct_1(X27)))&(k1_funct_1(X27,esk5_2(X27,X28))!=k1_funct_1(X28,esk5_2(X27,X28))|~r1_tarski(k1_relat_1(X27),k1_relat_1(X28))|r1_tarski(X27,X28)|(~v1_relat_1(X28)|~v1_funct_1(X28))|(~v1_relat_1(X27)|~v1_funct_1(X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_grfunc_1])])])])])).
cnf(c_0_93, negated_conjecture, (r1_tarski(k1_relat_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_85, c_0_24])).
cnf(c_0_94, plain, (~v1_xboole_0(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(er,[status(thm)],[c_0_86])).
cnf(c_0_95, plain, (k1_funct_1(X1,X2)=k1_xboole_0|r2_hidden(X2,k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[c_0_87])).
cnf(c_0_96, plain, (v1_relat_1(X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_53, c_0_88])).
cnf(c_0_97, negated_conjecture, (X1=k1_funct_1(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_98, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|r2_hidden(esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_89, c_0_83])).
cnf(c_0_99, negated_conjecture, (m1_subset_1(esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_62]), c_0_40])]), c_0_64])).
cnf(c_0_100, negated_conjecture, (v1_xboole_0(esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_91, c_0_83])).
cnf(c_0_101, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(X2))|~r1_tarski(X1,X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_102, negated_conjecture, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_93])).
cnf(c_0_103, plain, (k1_funct_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_41]), c_0_96])).
cnf(c_0_104, negated_conjecture, (k1_funct_1(esk3_0,esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))=esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)|v1_xboole_0(u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_99])])).
cnf(c_0_105, negated_conjecture, (esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_29, c_0_100])).
cnf(c_0_106, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_107, negated_conjecture, (r1_tarski(k1_relat_1(X1),k1_xboole_0)|~v1_relat_1(X1)|~r1_tarski(X1,k1_xboole_0)|~v1_funct_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_69]), c_0_70])])).
fof(c_0_108, plain, ![X40, X41, X42]:(~v1_xboole_0(X40)|(~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(X40,X41)))|v1_xboole_0(X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
cnf(c_0_109, plain, (r2_funct_2(X2,X3,X1,X4)|k1_funct_1(X1,esk6_4(X2,X3,X1,X4))!=k1_funct_1(X4,esk6_4(X2,X3,X1,X4))|~v1_funct_1(X4)|~v1_funct_2(X4,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_110, negated_conjecture, (esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)=k1_xboole_0|~v1_xboole_0(esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105])).
cnf(c_0_111, negated_conjecture, (~v1_relat_1(X1)|~r1_tarski(X1,k1_xboole_0)|~r2_hidden(X2,k1_relat_1(X1))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_67])).
cnf(c_0_112, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_113, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)!=k1_funct_1(esk3_0,k1_xboole_0)|~v1_xboole_0(esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_58]), c_0_76]), c_0_59]), c_0_75]), c_0_60]), c_0_77])]), c_0_78])).
cnf(c_0_114, negated_conjecture, (k1_funct_1(X1,X2)=k1_xboole_0|~v1_relat_1(X1)|~r1_tarski(X1,k1_xboole_0)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_111, c_0_95])).
cnf(c_0_115, negated_conjecture, (v1_relat_1(k4_tmap_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_53, c_0_76])).
cnf(c_0_116, negated_conjecture, (v1_xboole_0(k4_tmap_1(esk1_0,esk2_0))|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_112, c_0_76])).
cnf(c_0_117, negated_conjecture, (~v1_xboole_0(esk3_0)|~r1_tarski(k4_tmap_1(esk1_0,esk2_0),k1_xboole_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_115]), c_0_77])]), c_0_103])).
cnf(c_0_118, negated_conjecture, (k4_tmap_1(esk1_0,esk2_0)=k1_xboole_0|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_29, c_0_116])).
cnf(c_0_119, negated_conjecture, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_26, c_0_44])).
cnf(c_0_120, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_112, c_0_58])).
fof(c_0_121, plain, ![X47, X48, X49]:(v2_struct_0(X47)|(~v2_pre_topc(X47)|~l1_pre_topc(X47))|(v2_struct_0(X48)|~m1_pre_topc(X48,X47)|(~m1_subset_1(X49,u1_struct_0(X47))|(~r2_hidden(X49,u1_struct_0(X48))|k1_funct_1(k4_tmap_1(X47,X48),X49)=X49)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t95_tmap_1])])])])).
cnf(c_0_122, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])]), c_0_120])).
cnf(c_0_123, plain, (v2_struct_0(X1)|v2_struct_0(X2)|k1_funct_1(k4_tmap_1(X1,X2),X3)=X3|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_124, negated_conjecture, (r2_hidden(esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[c_0_98, c_0_122])).
cnf(c_0_125, negated_conjecture, (k1_funct_1(k4_tmap_1(X1,esk2_0),esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))=esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)|v2_struct_0(X1)|~m1_pre_topc(esk2_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_84]), c_0_90])).
cnf(c_0_126, negated_conjecture, (k1_funct_1(k4_tmap_1(esk1_0,esk2_0),esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))=esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_62]), c_0_40]), c_0_63])]), c_0_64])).
cnf(c_0_127, negated_conjecture, (k1_funct_1(esk3_0,esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0))=esk6_4(u1_struct_0(esk2_0),u1_struct_0(esk1_0),k4_tmap_1(esk1_0,esk2_0),esk3_0)), inference(sr,[status(thm)],[c_0_104, c_0_122])).
cnf(c_0_128, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_126]), c_0_127]), c_0_58]), c_0_76]), c_0_59]), c_0_75]), c_0_60]), c_0_77])]), c_0_78]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 129
# Proof object clause steps            : 88
# Proof object formula steps           : 41
# Proof object conjectures             : 49
# Proof object clause conjectures      : 46
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 51
# Proof object simplifying inferences  : 70
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 55
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 55
# Processed clauses                    : 9319
# ...of these trivial                  : 7
# ...subsumed                          : 7469
# ...remaining for further processing  : 1843
# Other redundant clauses eliminated   : 13
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 206
# Backward-rewritten                   : 29
# Generated clauses                    : 36187
# ...of the previous two non-trivial   : 31334
# Contextual simplify-reflections      : 181
# Paramodulations                      : 36108
# Factorizations                       : 10
# Equation resolutions                 : 59
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1546
#    Positive orientable unit clauses  : 38
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 1497
# Current number of unprocessed clauses: 21736
# ...number of literals in the above   : 146535
# Current number of archived formulas  : 0
# Current number of archived clauses   : 297
# Clause-clause subsumption calls (NU) : 678252
# Rec. Clause-clause subsumption calls : 244754
# Non-unit clause-clause subsumptions  : 4392
# Unit Clause-clause subsumption calls : 1639
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 51
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 819791

# -------------------------------------------------
# User time                : 0.683 s
# System time              : 0.007 s
# Total time               : 0.689 s
# Maximum resident set size: 28724 pages
