# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', d1_xboole_0)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', d4_relat_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', t7_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', fc1_xboole_0)).
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', d2_funct_2)).
fof(t6_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))=>~((r2_hidden(X1,X4)&![X5, X6]:~((X1=k4_tarski(X5,X6)&(r2_hidden(X5,X2)&r2_hidden(X6,X3))))))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', t6_relset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', d1_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', d10_xboole_0)).
fof(t121_funct_2, conjecture, ![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', t121_funct_2)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', t1_subset)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', t3_funct_2)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', t3_subset)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', t11_relset_1)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', cc5_funct_2)).
fof(d1_funct_1, axiom, ![X1]:(v1_funct_1(X1)<=>![X2, X3, X4]:((r2_hidden(k4_tarski(X2,X3),X1)&r2_hidden(k4_tarski(X2,X4),X1))=>X3=X4)), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', d1_funct_1)).
fof(t5_funct_2, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((k1_relat_1(X3)=X1&![X4]:(r2_hidden(X4,X1)=>r2_hidden(k1_funct_1(X3,X4),X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', t5_funct_2)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/e1839945b59b4149a2735fe382716a09.E_input', cc1_funct_2)).
fof(c_0_17, plain, ![X41, X42, X43]:((~v1_xboole_0(X41)|~r2_hidden(X42,X41))&(r2_hidden(esk11_1(X43),X43)|v1_xboole_0(X43))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_18, plain, ![X17, X18, X19, X21, X22, X23, X24, X26]:(((~r2_hidden(X19,X18)|r2_hidden(k4_tarski(X19,esk5_3(X17,X18,X19)),X17)|X18!=k1_relat_1(X17))&(~r2_hidden(k4_tarski(X21,X22),X17)|r2_hidden(X21,X18)|X18!=k1_relat_1(X17)))&((~r2_hidden(esk6_2(X23,X24),X24)|~r2_hidden(k4_tarski(esk6_2(X23,X24),X26),X23)|X24=k1_relat_1(X23))&(r2_hidden(esk6_2(X23,X24),X24)|r2_hidden(k4_tarski(esk6_2(X23,X24),esk7_2(X23,X24)),X23)|X24=k1_relat_1(X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_19, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, plain, (r2_hidden(k4_tarski(X1,esk5_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_21, plain, (X1!=k1_relat_1(X2)|~v1_xboole_0(X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_22, plain, ![X61]:(X61=k1_xboole_0|r2_hidden(esk14_1(X61),X61)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_23, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_24, plain, (X1=k1_xboole_0|r2_hidden(esk14_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_25, plain, (k1_relat_1(X1)=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_26, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_27, plain, ![X28, X29, X30, X31, X33, X34, X35, X36, X37, X39]:(((((v1_relat_1(esk8_4(X28,X29,X30,X31))|~r2_hidden(X31,X30)|X30!=k1_funct_2(X28,X29))&(v1_funct_1(esk8_4(X28,X29,X30,X31))|~r2_hidden(X31,X30)|X30!=k1_funct_2(X28,X29)))&((X31=esk8_4(X28,X29,X30,X31)|~r2_hidden(X31,X30)|X30!=k1_funct_2(X28,X29))&((k1_relat_1(esk8_4(X28,X29,X30,X31))=X28|~r2_hidden(X31,X30)|X30!=k1_funct_2(X28,X29))&(r1_tarski(k2_relat_1(esk8_4(X28,X29,X30,X31)),X29)|~r2_hidden(X31,X30)|X30!=k1_funct_2(X28,X29)))))&(~v1_relat_1(X34)|~v1_funct_1(X34)|(X33!=X34|(k1_relat_1(X34)!=X28|~r1_tarski(k2_relat_1(X34),X29)))|r2_hidden(X33,X30)|X30!=k1_funct_2(X28,X29)))&((~r2_hidden(esk9_3(X35,X36,X37),X37)|(~v1_relat_1(X39)|~v1_funct_1(X39)|(esk9_3(X35,X36,X37)!=X39|(k1_relat_1(X39)!=X35|~r1_tarski(k2_relat_1(X39),X36))))|X37=k1_funct_2(X35,X36))&(((v1_relat_1(esk10_3(X35,X36,X37))|r2_hidden(esk9_3(X35,X36,X37),X37)|X37=k1_funct_2(X35,X36))&(v1_funct_1(esk10_3(X35,X36,X37))|r2_hidden(esk9_3(X35,X36,X37),X37)|X37=k1_funct_2(X35,X36)))&((esk9_3(X35,X36,X37)=esk10_3(X35,X36,X37)|r2_hidden(esk9_3(X35,X36,X37),X37)|X37=k1_funct_2(X35,X36))&((k1_relat_1(esk10_3(X35,X36,X37))=X35|r2_hidden(esk9_3(X35,X36,X37),X37)|X37=k1_funct_2(X35,X36))&(r1_tarski(k2_relat_1(esk10_3(X35,X36,X37)),X36)|r2_hidden(esk9_3(X35,X36,X37),X37)|X37=k1_funct_2(X35,X36))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
cnf(c_0_28, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_29, plain, ![X48, X49, X50, X51]:((X48=k4_tarski(esk12_4(X48,X49,X50,X51),esk13_4(X48,X49,X50,X51))|~r2_hidden(X48,X51)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X49,X50))))&((r2_hidden(esk12_4(X48,X49,X50,X51),X49)|~r2_hidden(X48,X51)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X49,X50))))&(r2_hidden(esk13_4(X48,X49,X50,X51),X50)|~r2_hidden(X48,X51)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X49,X50)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_relset_1])])])])).
fof(c_0_30, plain, ![X10, X11, X12, X13, X14, X15]:(((~r2_hidden(X12,X11)|r1_tarski(X12,X10)|X11!=k1_zfmisc_1(X10))&(~r1_tarski(X13,X10)|r2_hidden(X13,X11)|X11!=k1_zfmisc_1(X10)))&((~r2_hidden(esk4_2(X14,X15),X15)|~r1_tarski(esk4_2(X14,X15),X14)|X15=k1_zfmisc_1(X14))&(r2_hidden(esk4_2(X14,X15),X15)|r1_tarski(esk4_2(X14,X15),X14)|X15=k1_zfmisc_1(X14)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_31, plain, ![X76, X77]:(((r1_tarski(X76,X77)|X76!=X77)&(r1_tarski(X77,X76)|X76!=X77))&(~r1_tarski(X76,X77)|~r1_tarski(X77,X76)|X76=X77)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_32, plain, (k1_relat_1(esk8_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (X1=esk8_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_34, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t121_funct_2])).
cnf(c_0_35, plain, (v1_relat_1(esk8_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (v1_funct_1(esk8_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (~r2_hidden(X1,k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_28]), c_0_26])])).
cnf(c_0_38, plain, (r2_hidden(esk12_4(X1,X2,X3,X4),X2)|~r2_hidden(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_39, plain, ![X56, X57]:(~r2_hidden(X56,X57)|m1_subset_1(X56,X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_40, plain, (r2_hidden(X1,X3)|~r1_tarski(X1,X2)|X3!=k1_zfmisc_1(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_42, plain, (k1_relat_1(X1)=X2|X3!=k1_funct_2(X2,X4)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
fof(c_0_43, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))&(~v1_funct_1(esk3_0)|(~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_34])])])).
cnf(c_0_44, plain, (v1_relat_1(X1)|X2!=k1_funct_2(X3,X4)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_35, c_0_33])).
cnf(c_0_45, plain, (v1_funct_1(X1)|X2!=k1_funct_2(X3,X4)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_36, c_0_33])).
cnf(c_0_46, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_28]), c_0_26])])).
cnf(c_0_47, plain, (r2_hidden(esk11_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_48, plain, (~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X2)))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_49, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, plain, (r2_hidden(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_41])).
fof(c_0_52, plain, ![X85]:((v1_funct_1(X85)|(~v1_relat_1(X85)|~v1_funct_1(X85)))&((v1_funct_2(X85,k1_relat_1(X85),k2_relat_1(X85))|(~v1_relat_1(X85)|~v1_funct_1(X85)))&(m1_subset_1(X85,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X85),k2_relat_1(X85))))|(~v1_relat_1(X85)|~v1_funct_1(X85))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_53, plain, (k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_54, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, plain, (v1_relat_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_56, plain, (v1_funct_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_57, plain, (v1_xboole_0(X1)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_58, plain, (~r2_hidden(X1,k1_zfmisc_1(k2_zfmisc_1(k1_xboole_0,X2)))|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_59, plain, (r2_hidden(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_60, plain, ![X74, X75]:((~m1_subset_1(X74,k1_zfmisc_1(X75))|r1_tarski(X74,X75))&(~r1_tarski(X74,X75)|m1_subset_1(X74,k1_zfmisc_1(X75)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_61, plain, ![X82, X83, X84]:(~v1_relat_1(X84)|(~r1_tarski(k1_relat_1(X84),X82)|~r1_tarski(k2_relat_1(X84),X83)|m1_subset_1(X84,k1_zfmisc_1(k2_zfmisc_1(X82,X83))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_62, plain, (r2_hidden(esk13_4(X1,X2,X3,X4),X3)|~r2_hidden(X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_63, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_65, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_55, c_0_54])).
cnf(c_0_66, negated_conjecture, (v1_funct_1(esk3_0)), inference(spm,[status(thm)],[c_0_56, c_0_54])).
fof(c_0_67, plain, ![X45, X46, X47]:((v1_funct_1(X47)|(~v1_funct_1(X47)|~v1_funct_2(X47,X45,X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X45,X46)))|v1_xboole_0(X46))&(v1_partfun1(X47,X45)|(~v1_funct_1(X47)|~v1_funct_2(X47,X45,X46))|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X45,X46)))|v1_xboole_0(X46))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
cnf(c_0_68, plain, (k1_relat_1(X1)=k1_xboole_0|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_25, c_0_57])).
cnf(c_0_69, plain, (~r2_hidden(X1,k2_zfmisc_1(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
fof(c_0_70, plain, ![X66, X67, X68, X69, X70]:((~v1_funct_1(X66)|(~r2_hidden(k4_tarski(X67,X68),X66)|~r2_hidden(k4_tarski(X67,X69),X66)|X68=X69))&(((r2_hidden(k4_tarski(esk15_1(X70),esk16_1(X70)),X70)|v1_funct_1(X70))&(r2_hidden(k4_tarski(esk15_1(X70),esk17_1(X70)),X70)|v1_funct_1(X70)))&(esk16_1(X70)!=esk17_1(X70)|v1_funct_1(X70)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_1])])])])])])).
cnf(c_0_71, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_72, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_73, plain, (r1_tarski(k2_relat_1(esk8_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_74, plain, (~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~r2_hidden(X4,X2)), inference(spm,[status(thm)],[c_0_19, c_0_62])).
cnf(c_0_75, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65]), c_0_66])])).
cnf(c_0_76, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_77, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_78, negated_conjecture, (~v1_funct_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_79, negated_conjecture, (esk1_0=k1_xboole_0|esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_68, c_0_64])).
cnf(c_0_80, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_69, c_0_24])).
fof(c_0_81, plain, ![X78, X79, X80]:(((v1_funct_1(X80)|(r2_hidden(esk18_3(X78,X79,X80),X78)|k1_relat_1(X80)!=X78)|(~v1_relat_1(X80)|~v1_funct_1(X80)))&((v1_funct_2(X80,X78,X79)|(r2_hidden(esk18_3(X78,X79,X80),X78)|k1_relat_1(X80)!=X78)|(~v1_relat_1(X80)|~v1_funct_1(X80)))&(m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(X78,X79)))|(r2_hidden(esk18_3(X78,X79,X80),X78)|k1_relat_1(X80)!=X78)|(~v1_relat_1(X80)|~v1_funct_1(X80)))))&((v1_funct_1(X80)|(~r2_hidden(k1_funct_1(X80,esk18_3(X78,X79,X80)),X79)|k1_relat_1(X80)!=X78)|(~v1_relat_1(X80)|~v1_funct_1(X80)))&((v1_funct_2(X80,X78,X79)|(~r2_hidden(k1_funct_1(X80,esk18_3(X78,X79,X80)),X79)|k1_relat_1(X80)!=X78)|(~v1_relat_1(X80)|~v1_funct_1(X80)))&(m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(X78,X79)))|(~r2_hidden(k1_funct_1(X80,esk18_3(X78,X79,X80)),X79)|k1_relat_1(X80)!=X78)|(~v1_relat_1(X80)|~v1_funct_1(X80)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_funct_2])])])])).
cnf(c_0_82, plain, (r2_hidden(k4_tarski(esk15_1(X1),esk17_1(X1)),X1)|v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_83, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_84, plain, (r1_tarski(k2_relat_1(X1),X2)|X3!=k1_funct_2(X4,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_73, c_0_33])).
cnf(c_0_85, negated_conjecture, (~v1_xboole_0(k2_relat_1(esk3_0))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_86, plain, (v1_partfun1(X1,k1_relat_1(X1))|v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_63]), c_0_77])).
cnf(c_0_87, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_66])])).
cnf(c_0_88, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_xboole_0))|esk3_0!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_79]), c_0_80])).
cnf(c_0_89, plain, (v1_funct_2(X1,X2,X3)|r2_hidden(esk18_3(X2,X3,X1),X2)|k1_relat_1(X1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_90, plain, (v1_funct_1(X1)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_82])).
fof(c_0_91, plain, ![X63, X64, X65]:((v1_funct_1(X65)|(~v1_funct_1(X65)|~v1_partfun1(X65,X63))|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64))))&(v1_funct_2(X65,X63,X64)|(~v1_funct_1(X65)|~v1_partfun1(X65,X63))|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_92, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(X1,X2))|~r1_tarski(k2_relat_1(esk3_0),X2)|~r1_tarski(esk1_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_64]), c_0_65])])).
cnf(c_0_93, plain, (r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(er,[status(thm)],[c_0_84])).
cnf(c_0_94, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_64]), c_0_65]), c_0_66])])).
cnf(c_0_95, negated_conjecture, (esk3_0!=k1_xboole_0|~v1_funct_2(esk3_0,k1_xboole_0,esk2_0)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_79]), c_0_80]), c_0_88])).
cnf(c_0_96, plain, (v1_funct_2(X1,X2,X3)|k1_xboole_0!=X2|X1!=k1_xboole_0|~v1_relat_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_68]), c_0_90]), c_0_46])).
cnf(c_0_97, negated_conjecture, (~r1_tarski(k1_relat_1(esk3_0),esk1_0)|~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_72]), c_0_65])])).
cnf(c_0_98, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_99, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_100, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,X1))|~r1_tarski(k2_relat_1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_92, c_0_51])).
cnf(c_0_101, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_93, c_0_54])).
cnf(c_0_102, negated_conjecture, (esk3_0=k1_xboole_0|v1_partfun1(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_94, c_0_24])).
cnf(c_0_103, negated_conjecture, (esk3_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_65])])).
cnf(c_0_104, negated_conjecture, (~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_64]), c_0_51])])).
cnf(c_0_105, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_106, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_107, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)), inference(sr,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_108, negated_conjecture, (~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_101])])).
cnf(c_0_109, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107]), c_0_66])]), c_0_108]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 110
# Proof object clause steps            : 76
# Proof object formula steps           : 34
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 57
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 52
# Processed clauses                    : 11486
# ...of these trivial                  : 24
# ...subsumed                          : 9792
# ...remaining for further processing  : 1670
# Other redundant clauses eliminated   : 284
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 62
# Backward-rewritten                   : 21
# Generated clauses                    : 69892
# ...of the previous two non-trivial   : 50787
# Contextual simplify-reflections      : 187
# Paramodulations                      : 69441
# Factorizations                       : 4
# Equation resolutions                 : 444
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1530
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 36
#    Non-unit-clauses                  : 1467
# Current number of unprocessed clauses: 39176
# ...number of literals in the above   : 179437
# Current number of archived formulas  : 0
# Current number of archived clauses   : 137
# Clause-clause subsumption calls (NU) : 256621
# Rec. Clause-clause subsumption calls : 114740
# Non-unit clause-clause subsumptions  : 3697
# Unit Clause-clause subsumption calls : 3750
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 15
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 933631

# -------------------------------------------------
# User time                : 0.611 s
# System time              : 0.007 s
# Total time               : 0.617 s
# Maximum resident set size: 34444 pages
