# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/e37782cb2264421daaa294d528cc4c8a.E_input', t1_tsep_1)).
fof(t114_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>(u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)&![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>u1_pre_topc(k8_tmap_1(X1,X2))=k5_tmap_1(X1,X3)))))), file('data/MPTP/proofs/e37782cb2264421daaa294d528cc4c8a.E_input', t114_tmap_1)).
fof(dt_k8_tmap_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_pre_topc(X2,X1))=>(v1_pre_topc(k8_tmap_1(X1,X2))&(v2_pre_topc(k8_tmap_1(X1,X2))&l1_pre_topc(k8_tmap_1(X1,X2))))), file('data/MPTP/proofs/e37782cb2264421daaa294d528cc4c8a.E_input', dt_k8_tmap_1)).
fof(t115_tmap_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_pre_topc(X3,k8_tmap_1(X1,X2))=>(u1_struct_0(X3)=u1_struct_0(X2)=>(v1_tsep_1(X3,k8_tmap_1(X1,X2))&m1_pre_topc(X3,k8_tmap_1(X1,X2))))))), file('data/MPTP/proofs/e37782cb2264421daaa294d528cc4c8a.E_input', t115_tmap_1)).
fof(t102_tmap_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>r2_hidden(X2,k5_tmap_1(X1,X2)))), file('data/MPTP/proofs/e37782cb2264421daaa294d528cc4c8a.E_input', t102_tmap_1)).
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v3_pre_topc(X2,X1)<=>r2_hidden(X2,u1_pre_topc(X1))))), file('data/MPTP/proofs/e37782cb2264421daaa294d528cc4c8a.E_input', d5_pre_topc)).
fof(t16_tsep_1, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>((v1_tsep_1(X2,X1)&m1_pre_topc(X2,X1))<=>v3_pre_topc(X3,X1)))))), file('data/MPTP/proofs/e37782cb2264421daaa294d528cc4c8a.E_input', t16_tsep_1)).
fof(c_0_7, plain, ![X4, X5]:(~l1_pre_topc(X4)|(~m1_pre_topc(X5,X4)|m1_subset_1(u1_struct_0(X5),k1_zfmisc_1(u1_struct_0(X4))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_8, plain, ![X6, X7, X8]:((u1_struct_0(k8_tmap_1(X6,X7))=u1_struct_0(X6)|(v2_struct_0(X7)|~m1_pre_topc(X7,X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))&(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|(X8!=u1_struct_0(X7)|u1_pre_topc(k8_tmap_1(X6,X7))=k5_tmap_1(X6,X8))|(v2_struct_0(X7)|~m1_pre_topc(X7,X6))|(v2_struct_0(X6)|(~v2_pre_topc(X6)|~l1_pre_topc(X6))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t114_tmap_1])])])])])).
fof(c_0_9, plain, ![X9, X10]:((v1_pre_topc(k8_tmap_1(X9,X10))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_pre_topc(X10,X9)))&((v2_pre_topc(k8_tmap_1(X9,X10))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_pre_topc(X10,X9)))&(l1_pre_topc(k8_tmap_1(X9,X10))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))|~m1_pre_topc(X10,X9))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k8_tmap_1])])])])).
fof(c_0_10, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_pre_topc(X3,k8_tmap_1(X1,X2))=>(u1_struct_0(X3)=u1_struct_0(X2)=>(v1_tsep_1(X3,k8_tmap_1(X1,X2))&m1_pre_topc(X3,k8_tmap_1(X1,X2)))))))), inference(assume_negation,[status(cth)],[t115_tmap_1])).
cnf(c_0_11, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (u1_struct_0(k8_tmap_1(X1,X2))=u1_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (l1_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&m1_pre_topc(esk2_0,esk1_0))&(m1_pre_topc(esk3_0,k8_tmap_1(esk1_0,esk2_0))&(u1_struct_0(esk3_0)=u1_struct_0(esk2_0)&(~v1_tsep_1(esk3_0,k8_tmap_1(esk1_0,esk2_0))|~m1_pre_topc(esk3_0,k8_tmap_1(esk1_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_10])])])])).
fof(c_0_15, plain, ![X14, X15]:(v2_struct_0(X14)|(~v2_pre_topc(X14)|~l1_pre_topc(X14))|(~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(X14)))|r2_hidden(X15,k5_tmap_1(X14,X15)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t102_tmap_1])])])])).
cnf(c_0_16, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~m1_pre_topc(X3,k8_tmap_1(X1,X2))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_11, c_0_12]), c_0_13])).
cnf(c_0_17, negated_conjecture, (m1_pre_topc(esk3_0,k8_tmap_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (m1_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_23, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_24, plain, ![X16, X17]:((~v3_pre_topc(X17,X16)|r2_hidden(X17,u1_pre_topc(X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_pre_topc(X16))&(~r2_hidden(X17,u1_pre_topc(X16))|v3_pre_topc(X17,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_pre_topc(X16))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
cnf(c_0_25, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_tmap_1(X1,X2))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19]), c_0_20]), c_0_21])]), c_0_22]), c_0_23])).
cnf(c_0_27, plain, (u1_pre_topc(k8_tmap_1(X2,X3))=k5_tmap_1(X2,X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|X1!=u1_struct_0(X3)|~m1_pre_topc(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_28, plain, ![X18, X19, X20]:((~v1_tsep_1(X19,X18)|~m1_pre_topc(X19,X18)|v3_pre_topc(X20,X18)|X20!=u1_struct_0(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|~m1_pre_topc(X19,X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18)))&((v1_tsep_1(X19,X18)|~v3_pre_topc(X20,X18)|X20!=u1_struct_0(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|~m1_pre_topc(X19,X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18)))&(m1_pre_topc(X19,X18)|~v3_pre_topc(X20,X18)|X20!=u1_struct_0(X19)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(X18)))|~m1_pre_topc(X19,X18)|(~v2_pre_topc(X18)|~l1_pre_topc(X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_tsep_1])])])])).
cnf(c_0_29, plain, (v3_pre_topc(X1,X2)|~r2_hidden(X1,u1_pre_topc(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_30, negated_conjecture, (r2_hidden(u1_struct_0(esk2_0),k5_tmap_1(esk1_0,u1_struct_0(esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_19]), c_0_21])]), c_0_22])).
cnf(c_0_31, negated_conjecture, (u1_pre_topc(k8_tmap_1(esk1_0,X1))=k5_tmap_1(esk1_0,u1_struct_0(esk2_0))|v2_struct_0(X1)|u1_struct_0(esk2_0)!=u1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_26]), c_0_19]), c_0_21])]), c_0_22])).
cnf(c_0_32, plain, (v1_tsep_1(X1,X2)|~v3_pre_topc(X3,X2)|X3!=u1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_33, plain, (v2_pre_topc(k8_tmap_1(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_34, plain, (v3_pre_topc(X1,k8_tmap_1(X2,X3))|v2_struct_0(X2)|v2_struct_0(X3)|~r2_hidden(X1,u1_pre_topc(k8_tmap_1(X2,X3)))|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_12]), c_0_13])).
cnf(c_0_35, negated_conjecture, (r2_hidden(u1_struct_0(esk2_0),u1_pre_topc(k8_tmap_1(esk1_0,X1)))|v2_struct_0(X1)|u1_struct_0(esk2_0)!=u1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, negated_conjecture, (~v1_tsep_1(esk3_0,k8_tmap_1(esk1_0,esk2_0))|~m1_pre_topc(esk3_0,k8_tmap_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_37, plain, (v1_tsep_1(X1,k8_tmap_1(X2,X3))|v2_struct_0(X2)|v2_struct_0(X3)|X4!=u1_struct_0(X1)|~v3_pre_topc(X4,k8_tmap_1(X2,X3))|~v2_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,k8_tmap_1(X2,X3))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_12]), c_0_13]), c_0_33])).
cnf(c_0_38, negated_conjecture, (v3_pre_topc(u1_struct_0(esk2_0),k8_tmap_1(esk1_0,X1))|v2_struct_0(X1)|u1_struct_0(esk2_0)!=u1_struct_0(X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_19]), c_0_26]), c_0_21])]), c_0_22])).
cnf(c_0_39, negated_conjecture, (~v1_tsep_1(esk3_0,k8_tmap_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_17])])).
cnf(c_0_40, negated_conjecture, (v1_tsep_1(X1,k8_tmap_1(esk1_0,X2))|v2_struct_0(X2)|u1_struct_0(esk2_0)!=u1_struct_0(X1)|u1_struct_0(esk2_0)!=u1_struct_0(X2)|~m1_pre_topc(X1,k8_tmap_1(esk1_0,X2))|~m1_pre_topc(X2,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_19]), c_0_26]), c_0_21])]), c_0_22])).
cnf(c_0_41, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_18]), c_0_17]), c_0_20])]), c_0_23]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 27
# Proof object formula steps           : 15
# Proof object conjectures             : 19
# Proof object clause conjectures      : 16
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 7
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 19
# Processed clauses                    : 118
# ...of these trivial                  : 0
# ...subsumed                          : 20
# ...remaining for further processing  : 98
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 209
# ...of the previous two non-trivial   : 174
# Contextual simplify-reflections      : 7
# Paramodulations                      : 208
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 76
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 66
# Current number of unprocessed clauses: 94
# ...number of literals in the above   : 859
# Current number of archived formulas  : 0
# Current number of archived clauses   : 22
# Clause-clause subsumption calls (NU) : 1412
# Rec. Clause-clause subsumption calls : 157
# Non-unit clause-clause subsumptions  : 30
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7603

# -------------------------------------------------
# User time                : 0.014 s
# System time              : 0.000 s
# Total time               : 0.014 s
# Maximum resident set size: 28712 pages
