# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', d1_xboole_0)).
fof(t17_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3))))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', t17_yellow19)).
fof(d10_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>(v7_struct_0(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>X2=X3)))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', d10_struct_0)).
fof(d1_tex_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>(v7_struct_0(X1)<=>?[X2]:(m1_subset_1(X2,u1_struct_0(X1))&u1_struct_0(X1)=k6_domain_1(u1_struct_0(X1),X2)))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', d1_tex_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', dt_l1_pre_topc)).
fof(t10_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(m2_connsp_2(X3,X1,k6_domain_1(u1_struct_0(X1),X2))<=>m1_connsp_2(X3,X1,X2))))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', t10_connsp_2)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', d4_subset_1)).
fof(t35_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>m2_connsp_2(k2_struct_0(X1),X1,X2))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', t35_connsp_2)).
fof(t6_connsp_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(m1_connsp_2(X2,X1,X3)=>r2_hidden(X3,X2))))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', t6_connsp_2)).
fof(dt_m1_connsp_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&m1_subset_1(X2,u1_struct_0(X1)))=>![X3]:(m1_connsp_2(X3,X1,X2)=>m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', dt_m1_connsp_2)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', d3_struct_0)).
fof(dt_k3_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&l1_waybel_0(k3_yellow19(X1,X2,X3),X1)))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', dt_k3_yellow19)).
fof(d2_yellow_1, axiom, ![X1]:k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1)), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', d2_yellow_1)).
fof(t15_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2)))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', t15_yellow19)).
fof(fc4_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2)))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v3_orders_2(k3_yellow19(X1,X2,X3))&(v4_orders_2(k3_yellow19(X1,X2,X3))&v6_waybel_0(k3_yellow19(X1,X2,X3),X1))))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', fc4_yellow19)).
fof(fc5_yellow19, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))&(~(v1_xboole_0(X3))&(v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))&(v2_waybel_0(X3,k3_yellow_1(X2))&(v13_waybel_0(X3,k3_yellow_1(X2))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))))))))=>(~(v2_struct_0(k3_yellow19(X1,X2,X3)))&(v6_waybel_0(k3_yellow19(X1,X2,X3),X1)&v7_waybel_0(k3_yellow19(X1,X2,X3))))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', fc5_yellow19)).
fof(t12_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,X2,X3)<=>r1_waybel_7(X1,k2_yellow19(X1,X2),X3))))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', t12_yellow19)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', t8_boole)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/e5d64bfdcdda4e0595bd2df7eb6a1189.E_input', fc1_xboole_0)).
fof(c_0_21, plain, ![X51, X52]:(((~m1_subset_1(X52,X51)|r2_hidden(X52,X51)|v1_xboole_0(X51))&(~r2_hidden(X52,X51)|m1_subset_1(X52,X51)|v1_xboole_0(X51)))&((~m1_subset_1(X52,X51)|v1_xboole_0(X52)|~v1_xboole_0(X51))&(~v1_xboole_0(X52)|m1_subset_1(X52,X51)|~v1_xboole_0(X51)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_22, plain, ![X16, X17, X18]:((~v1_xboole_0(X16)|~r2_hidden(X17,X16))&(r2_hidden(esk5_1(X18),X18)|v1_xboole_0(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_23, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v1_xboole_0(X2))&(v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))&(v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&(v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r3_waybel_9(X1,k3_yellow19(X1,k2_struct_0(X1),X2),X3)<=>r1_waybel_7(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t17_yellow19])).
cnf(c_0_24, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_26, plain, ![X39, X40, X41]:((~v7_struct_0(X39)|(~m1_subset_1(X40,u1_struct_0(X39))|(~m1_subset_1(X41,u1_struct_0(X39))|X40=X41))|~l1_struct_0(X39))&((m1_subset_1(esk6_1(X39),u1_struct_0(X39))|v7_struct_0(X39)|~l1_struct_0(X39))&((m1_subset_1(esk7_1(X39),u1_struct_0(X39))|v7_struct_0(X39)|~l1_struct_0(X39))&(esk6_1(X39)!=esk7_1(X39)|v7_struct_0(X39)|~l1_struct_0(X39))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_struct_0])])])])])).
fof(c_0_27, plain, ![X7, X9]:(((m1_subset_1(esk4_1(X7),u1_struct_0(X7))|~v7_struct_0(X7)|(v2_struct_0(X7)|~l1_struct_0(X7)))&(u1_struct_0(X7)=k6_domain_1(u1_struct_0(X7),esk4_1(X7))|~v7_struct_0(X7)|(v2_struct_0(X7)|~l1_struct_0(X7))))&(~m1_subset_1(X9,u1_struct_0(X7))|u1_struct_0(X7)!=k6_domain_1(u1_struct_0(X7),X9)|v7_struct_0(X7)|(v2_struct_0(X7)|~l1_struct_0(X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tex_1])])])])])])).
fof(c_0_28, plain, ![X28]:(~l1_pre_topc(X28)|l1_struct_0(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_29, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v1_xboole_0(esk2_0)&(v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))&(v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&(v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))&m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))))))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&((~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0))&(r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|r1_waybel_7(esk1_0,esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
fof(c_0_30, plain, ![X45, X46, X47]:((~m2_connsp_2(X47,X45,k6_domain_1(u1_struct_0(X45),X46))|m1_connsp_2(X47,X45,X46)|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X45)))|~m1_subset_1(X46,u1_struct_0(X45))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))))&(~m1_connsp_2(X47,X45,X46)|m2_connsp_2(X47,X45,k6_domain_1(u1_struct_0(X45),X46))|~m1_subset_1(X47,k1_zfmisc_1(u1_struct_0(X45)))|~m1_subset_1(X46,u1_struct_0(X45))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_connsp_2])])])])])).
cnf(c_0_31, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_33, plain, ![X20]:m1_subset_1(k2_subset_1(X20),k1_zfmisc_1(X20)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_34, plain, ![X29]:k2_subset_1(X29)=X29, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_35, plain, (X2=X3|~v7_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (m1_connsp_2(X1,X2,X3)|v2_struct_0(X2)|~m2_connsp_2(X1,X2,k6_domain_1(u1_struct_0(X2),X3))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (m1_subset_1(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_41, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_43, plain, ![X37, X38]:(v2_struct_0(X37)|(~v2_pre_topc(X37)|~l1_pre_topc(X37))|(~m1_subset_1(X38,k1_zfmisc_1(u1_struct_0(X37)))|m2_connsp_2(k2_struct_0(X37),X37,X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_connsp_2])])])])).
cnf(c_0_44, plain, (X1=esk4_1(X2)|v2_struct_0(X2)|~v7_struct_0(X2)|~l1_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_46, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_48, plain, ![X32, X33, X34]:(v2_struct_0(X32)|(~v2_pre_topc(X32)|~l1_pre_topc(X32))|(~m1_subset_1(X33,k1_zfmisc_1(u1_struct_0(X32)))|(~m1_subset_1(X34,u1_struct_0(X32))|(~m1_connsp_2(X33,X32,X34)|r2_hidden(X34,X33))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_connsp_2])])])])).
fof(c_0_49, plain, ![X48, X49, X50]:(v2_struct_0(X48)|(~v2_pre_topc(X48)|~l1_pre_topc(X48))|~m1_subset_1(X49,u1_struct_0(X48))|(~m1_connsp_2(X50,X48,X49)|m1_subset_1(X50,k1_zfmisc_1(u1_struct_0(X48))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_connsp_2])])])])).
cnf(c_0_50, plain, (m1_connsp_2(X1,X2,esk5_1(u1_struct_0(X2)))|v1_xboole_0(u1_struct_0(X2))|v2_struct_0(X2)|~m2_connsp_2(X1,X2,k6_domain_1(u1_struct_0(X2),esk5_1(u1_struct_0(X2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_51, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_52, plain, (X1=esk5_1(u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2))|~v7_struct_0(X2)|~l1_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_35, c_0_40])).
cnf(c_0_53, plain, (v2_struct_0(X1)|m2_connsp_2(k2_struct_0(X1),X1,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_54, plain, ![X21]:(~l1_struct_0(X21)|k2_struct_0(X21)=u1_struct_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_55, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_56, plain, (u1_struct_0(X1)=k6_domain_1(u1_struct_0(X1),esk4_1(X1))|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_57, negated_conjecture, (esk4_1(esk1_0)=esk3_0|~v7_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_46])]), c_0_47])).
cnf(c_0_58, plain, (v2_struct_0(X1)|r2_hidden(X3,X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_connsp_2(X2,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (v2_struct_0(X1)|m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_connsp_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_60, plain, (m1_connsp_2(u1_struct_0(X1),X1,esk5_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_connsp_2(u1_struct_0(X1),X1,k6_domain_1(u1_struct_0(X1),esk5_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_61, plain, (esk5_1(u1_struct_0(X1))=esk4_1(X1)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_struct_0(X1)), inference(spm,[status(thm)],[c_0_52, c_0_36])).
cnf(c_0_62, plain, (m2_connsp_2(k2_struct_0(X1),X1,u1_struct_0(X1))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_53, c_0_51])).
cnf(c_0_63, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_64, negated_conjecture, (m1_connsp_2(X1,esk1_0,esk3_0)|~m2_connsp_2(X1,esk1_0,k6_domain_1(u1_struct_0(esk1_0),esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_45]), c_0_38]), c_0_55])]), c_0_47])).
cnf(c_0_65, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk3_0)=u1_struct_0(esk1_0)|~v7_struct_0(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_46])]), c_0_47])).
fof(c_0_66, plain, ![X10, X11, X12]:((~v2_struct_0(k3_yellow19(X10,X11,X12))|(v2_struct_0(X10)|~l1_struct_0(X10)|(v1_xboole_0(X11)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v1_xboole_0(X12)|(~v2_waybel_0(X12,k3_yellow_1(X11))|(~v13_waybel_0(X12,k3_yellow_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X11))))))))))&((v6_waybel_0(k3_yellow19(X10,X11,X12),X10)|(v2_struct_0(X10)|~l1_struct_0(X10)|(v1_xboole_0(X11)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v1_xboole_0(X12)|(~v2_waybel_0(X12,k3_yellow_1(X11))|(~v13_waybel_0(X12,k3_yellow_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X11))))))))))&(l1_waybel_0(k3_yellow19(X10,X11,X12),X10)|(v2_struct_0(X10)|~l1_struct_0(X10)|(v1_xboole_0(X11)|~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(X10)))|(v1_xboole_0(X12)|(~v2_waybel_0(X12,k3_yellow_1(X11))|(~v13_waybel_0(X12,k3_yellow_1(X11))|~m1_subset_1(X12,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X11)))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_yellow19])])])])).
fof(c_0_67, plain, ![X44]:k3_yellow_1(X44)=k3_lattice3(k1_lattice3(X44)), inference(variable_rename,[status(thm)],[d2_yellow_1])).
fof(c_0_68, plain, ![X30, X31]:(v2_struct_0(X30)|~l1_struct_0(X30)|(v1_xboole_0(X31)|(~v1_subset_1(X31,u1_struct_0(k3_yellow_1(k2_struct_0(X30))))|(~v2_waybel_0(X31,k3_yellow_1(k2_struct_0(X30)))|(~v13_waybel_0(X31,k3_yellow_1(k2_struct_0(X30)))|~m1_subset_1(X31,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X30))))))))|X31=k2_yellow19(X30,k3_yellow19(X30,k2_struct_0(X30),X31)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_yellow19])])])])).
fof(c_0_69, plain, ![X25, X26, X27]:((~v2_struct_0(k3_yellow19(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v1_xboole_0(X27)|(~v2_waybel_0(X27,k3_yellow_1(X26))|(~v13_waybel_0(X27,k3_yellow_1(X26))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X26))))))))))&((v3_orders_2(k3_yellow19(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v1_xboole_0(X27)|(~v2_waybel_0(X27,k3_yellow_1(X26))|(~v13_waybel_0(X27,k3_yellow_1(X26))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X26))))))))))&((v4_orders_2(k3_yellow19(X25,X26,X27))|(v2_struct_0(X25)|~l1_struct_0(X25)|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v1_xboole_0(X27)|(~v2_waybel_0(X27,k3_yellow_1(X26))|(~v13_waybel_0(X27,k3_yellow_1(X26))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X26))))))))))&(v6_waybel_0(k3_yellow19(X25,X26,X27),X25)|(v2_struct_0(X25)|~l1_struct_0(X25)|(v1_xboole_0(X26)|~m1_subset_1(X26,k1_zfmisc_1(u1_struct_0(X25)))|(v1_xboole_0(X27)|(~v2_waybel_0(X27,k3_yellow_1(X26))|(~v13_waybel_0(X27,k3_yellow_1(X26))|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X26))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_yellow19])])])])).
cnf(c_0_70, plain, (r2_hidden(X1,X2)|v2_struct_0(X3)|~m1_connsp_2(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X3))|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(csr,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_71, plain, (m1_connsp_2(u1_struct_0(X1),X1,esk4_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m2_connsp_2(u1_struct_0(X1),X1,k6_domain_1(u1_struct_0(X1),esk4_1(X1)))|~v7_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_37])).
cnf(c_0_72, plain, (m2_connsp_2(u1_struct_0(X1),X1,u1_struct_0(X1))|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_37])).
cnf(c_0_73, negated_conjecture, (m1_connsp_2(X1,esk1_0,esk3_0)|~m2_connsp_2(X1,esk1_0,u1_struct_0(esk1_0))|~v7_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74, plain, (l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_75, plain, (k3_yellow_1(X1)=k3_lattice3(k1_lattice3(X1))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76, negated_conjecture, (v13_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_77, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0)))))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_78, negated_conjecture, (v2_waybel_0(esk2_0,k3_yellow_1(k2_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_79, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|~l1_struct_0(X1)|~v1_subset_1(X2,u1_struct_0(k3_yellow_1(k2_struct_0(X1))))|~v2_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~v13_waybel_0(X2,k3_yellow_1(k2_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_yellow_1(k2_struct_0(X1)))))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_80, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_yellow_1(k2_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_81, plain, (v4_orders_2(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_82, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_66])).
fof(c_0_83, plain, ![X13, X14, X15]:((~v2_struct_0(k3_yellow19(X13,X14,X15))|(v2_struct_0(X13)|~l1_struct_0(X13)|(v1_xboole_0(X14)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|(v1_xboole_0(X15)|(~v1_subset_1(X15,u1_struct_0(k3_yellow_1(X14)))|(~v2_waybel_0(X15,k3_yellow_1(X14))|(~v13_waybel_0(X15,k3_yellow_1(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X14)))))))))))&((v6_waybel_0(k3_yellow19(X13,X14,X15),X13)|(v2_struct_0(X13)|~l1_struct_0(X13)|(v1_xboole_0(X14)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|(v1_xboole_0(X15)|(~v1_subset_1(X15,u1_struct_0(k3_yellow_1(X14)))|(~v2_waybel_0(X15,k3_yellow_1(X14))|(~v13_waybel_0(X15,k3_yellow_1(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X14)))))))))))&(v7_waybel_0(k3_yellow19(X13,X14,X15))|(v2_struct_0(X13)|~l1_struct_0(X13)|(v1_xboole_0(X14)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X13)))|(v1_xboole_0(X15)|(~v1_subset_1(X15,u1_struct_0(k3_yellow_1(X14)))|(~v2_waybel_0(X15,k3_yellow_1(X14))|(~v13_waybel_0(X15,k3_yellow_1(X14))|~m1_subset_1(X15,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X14))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_yellow19])])])])).
cnf(c_0_84, plain, (r2_hidden(esk4_1(X1),X2)|v2_struct_0(X1)|~m1_connsp_2(X2,X1,esk4_1(X1))|~v7_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_36]), c_0_37])).
cnf(c_0_85, plain, (m1_connsp_2(u1_struct_0(X1),X1,esk4_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_56]), c_0_37]), c_0_72])).
cnf(c_0_86, negated_conjecture, (m1_connsp_2(k2_struct_0(esk1_0),esk1_0,esk3_0)|~v7_struct_0(esk1_0)|~m1_subset_1(k2_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_62]), c_0_38]), c_0_55])]), c_0_47])).
fof(c_0_87, plain, ![X22, X23, X24]:((~r3_waybel_9(X22,X23,X24)|r1_waybel_7(X22,k2_yellow19(X22,X23),X24)|~m1_subset_1(X24,u1_struct_0(X22))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))))&(~r1_waybel_7(X22,k2_yellow19(X22,X23),X24)|r3_waybel_9(X22,X23,X24)|~m1_subset_1(X24,u1_struct_0(X22))|(v2_struct_0(X23)|(~v4_orders_2(X23)|(~v7_waybel_0(X23)|~l1_waybel_0(X23,X22))))|(v2_struct_0(X22)|(~v2_pre_topc(X22)|~l1_pre_topc(X22))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_yellow19])])])])])).
cnf(c_0_88, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|l1_waybel_0(k3_yellow19(X1,X2,X3),X1)|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_75]), c_0_75]), c_0_75])).
cnf(c_0_89, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_76, c_0_75])).
cnf(c_0_90, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))))), inference(rw,[status(thm)],[c_0_77, c_0_75])).
cnf(c_0_91, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(k2_struct_0(esk1_0))))), inference(rw,[status(thm)],[c_0_78, c_0_75])).
cnf(c_0_92, plain, (X2=k2_yellow19(X1,k3_yellow19(X1,k2_struct_0(X1),X2))|v2_struct_0(X1)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(k2_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(X1))))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_75]), c_0_75]), c_0_75]), c_0_75])).
cnf(c_0_93, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(k2_struct_0(esk1_0)))))), inference(rw,[status(thm)],[c_0_80, c_0_75])).
cnf(c_0_94, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_95, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v4_orders_2(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_75]), c_0_75]), c_0_75])).
cnf(c_0_96, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(k3_yellow19(X1,X2,X3))|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_75]), c_0_75]), c_0_75])).
cnf(c_0_97, plain, (v7_waybel_0(k3_yellow19(X1,X2,X3))|v2_struct_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_yellow_1(X2)))|~v2_waybel_0(X3,k3_yellow_1(X2))|~v13_waybel_0(X3,k3_yellow_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_yellow_1(X2))))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_98, plain, (r2_hidden(esk4_1(X1),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_99, negated_conjecture, (r2_hidden(esk3_0,X1)|~m1_connsp_2(X1,esk1_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_45]), c_0_38]), c_0_55])]), c_0_47])).
cnf(c_0_100, negated_conjecture, (m1_connsp_2(u1_struct_0(esk1_0),esk1_0,esk3_0)|~v7_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_63]), c_0_51]), c_0_46])])).
cnf(c_0_101, plain, (r3_waybel_9(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_102, plain, (l1_waybel_0(k3_yellow19(X1,u1_struct_0(X1),X2),X1)|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))), inference(spm,[status(thm)],[c_0_88, c_0_51])).
cnf(c_0_103, negated_conjecture, (v13_waybel_0(esk2_0,k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_63]), c_0_46])])).
cnf(c_0_104, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_63]), c_0_46])])).
cnf(c_0_105, negated_conjecture, (v2_waybel_0(esk2_0,k3_lattice3(k1_lattice3(u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_63]), c_0_46])])).
cnf(c_0_106, negated_conjecture, (k2_yellow19(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0))=esk2_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_89]), c_0_46]), c_0_90]), c_0_91]), c_0_93])]), c_0_94]), c_0_47])).
cnf(c_0_107, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_108, plain, (v4_orders_2(k3_yellow19(X1,u1_struct_0(X1),X2))|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))), inference(spm,[status(thm)],[c_0_95, c_0_51])).
cnf(c_0_109, plain, (v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_struct_0(k3_yellow19(X1,u1_struct_0(X1),X2))), inference(spm,[status(thm)],[c_0_96, c_0_51])).
cnf(c_0_110, plain, (v2_struct_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X2)|v7_waybel_0(k3_yellow19(X1,X2,X3))|~l1_struct_0(X1)|~v2_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~v13_waybel_0(X3,k3_lattice3(k1_lattice3(X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_subset_1(X3,u1_struct_0(k3_lattice3(k1_lattice3(X2))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(X2)))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_75]), c_0_75]), c_0_75]), c_0_75])).
cnf(c_0_111, plain, (r1_waybel_7(X1,k2_yellow19(X1,X2),X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_112, plain, (m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_31, c_0_98])).
cnf(c_0_113, negated_conjecture, (r2_hidden(esk3_0,u1_struct_0(esk1_0))|~v7_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_114, negated_conjecture, (r3_waybel_9(esk1_0,X1,esk3_0)|v2_struct_0(X1)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~r1_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_45]), c_0_38]), c_0_55])]), c_0_47])).
cnf(c_0_115, negated_conjecture, (l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_46]), c_0_104]), c_0_105])]), c_0_94]), c_0_47])).
cnf(c_0_116, negated_conjecture, (k2_yellow19(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_63]), c_0_46])])).
cnf(c_0_117, negated_conjecture, (r1_waybel_7(esk1_0,esk2_0,esk3_0)|r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_63]), c_0_46])])).
cnf(c_0_118, negated_conjecture, (v4_orders_2(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_103]), c_0_46]), c_0_104]), c_0_105])]), c_0_94]), c_0_47])).
cnf(c_0_119, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|~v2_struct_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_103]), c_0_46]), c_0_104]), c_0_105])]), c_0_94]), c_0_47])).
cnf(c_0_120, plain, (v7_waybel_0(k3_yellow19(X1,u1_struct_0(X1),X2))|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(X2)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1))))))|~v13_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v2_waybel_0(X2,k3_lattice3(k1_lattice3(u1_struct_0(X1))))|~v1_subset_1(X2,u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(X1)))))), inference(spm,[status(thm)],[c_0_110, c_0_51])).
cnf(c_0_121, negated_conjecture, (v1_subset_1(esk2_0,u1_struct_0(k3_lattice3(k1_lattice3(u1_struct_0(esk1_0)))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_63]), c_0_46])])).
cnf(c_0_122, negated_conjecture, (r1_waybel_7(esk1_0,k2_yellow19(esk1_0,X1),esk3_0)|v2_struct_0(X1)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)|~r3_waybel_9(esk1_0,X1,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_45]), c_0_38]), c_0_55])]), c_0_47])).
cnf(c_0_123, plain, (r3_waybel_9(X1,X2,esk4_1(X1))|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|v2_struct_0(X2)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v7_struct_0(X1)|~r1_waybel_7(X1,k2_yellow19(X1,X2),esk4_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(spm,[status(thm)],[c_0_101, c_0_112])).
cnf(c_0_124, negated_conjecture, (~v7_struct_0(esk1_0)|~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_25, c_0_113])).
cnf(c_0_125, negated_conjecture, (~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,k2_struct_0(esk1_0),esk2_0),esk3_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_126, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_116]), c_0_117]), c_0_118]), c_0_119])).
cnf(c_0_127, negated_conjecture, (v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))|v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_103]), c_0_46]), c_0_104]), c_0_105]), c_0_121])]), c_0_94]), c_0_47])).
cnf(c_0_128, negated_conjecture, (r1_waybel_7(esk1_0,esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_115]), c_0_116]), c_0_117]), c_0_118]), c_0_119])).
cnf(c_0_129, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk4_1(esk1_0))|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk1_0)|~v7_struct_0(esk1_0)|~r1_waybel_7(esk1_0,esk2_0,esk4_1(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_116]), c_0_38]), c_0_55])]), c_0_47]), c_0_118]), c_0_119]), c_0_124])).
fof(c_0_130, plain, ![X35, X36]:(~v1_xboole_0(X35)|(X35=X36|~v1_xboole_0(X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
cnf(c_0_131, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_132, negated_conjecture, (~r1_waybel_7(esk1_0,esk2_0,esk3_0)|~r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_63]), c_0_46])])).
cnf(c_0_133, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_134, negated_conjecture, (r1_waybel_7(esk1_0,esk2_0,esk3_0)|v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_128, c_0_127])).
cnf(c_0_135, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))|~l1_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk1_0)|~v7_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_57]), c_0_117])).
cnf(c_0_136, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_137, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_138, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_131, c_0_45])).
cnf(c_0_139, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133]), c_0_134])).
cnf(c_0_140, negated_conjecture, (r3_waybel_9(esk1_0,k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0),esk3_0)|~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))|~v7_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_115]), c_0_124])).
cnf(c_0_141, plain, (k1_xboole_0=X1|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_142, negated_conjecture, (v1_xboole_0(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_138, c_0_139])])).
cnf(c_0_143, negated_conjecture, (~v7_waybel_0(k3_yellow19(esk1_0,u1_struct_0(esk1_0),esk2_0))|~v7_struct_0(esk1_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(spm,[status(thm)],[c_0_132, c_0_140])).
cnf(c_0_144, plain, (m1_subset_1(esk6_1(X1),u1_struct_0(X1))|v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_145, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_141, c_0_139])).
cnf(c_0_146, negated_conjecture, (k1_xboole_0=esk3_0), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_147, negated_conjecture, (~v7_struct_0(esk1_0)|~r1_waybel_7(esk1_0,esk2_0,esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_127]), c_0_124])).
cnf(c_0_148, plain, (v7_struct_0(X1)|esk6_1(X1)!=esk7_1(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_149, plain, (m1_subset_1(esk7_1(X1),u1_struct_0(X1))|v7_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_150, plain, (v7_struct_0(X1)|v1_xboole_0(esk6_1(X1))|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_131, c_0_144])).
cnf(c_0_151, negated_conjecture, (u1_struct_0(esk1_0)=esk3_0), inference(rw,[status(thm)],[c_0_145, c_0_146])).
cnf(c_0_152, negated_conjecture, (~v7_struct_0(esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_134]), c_0_124])).
cnf(c_0_153, plain, (v7_struct_0(X1)|X2!=esk6_1(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk7_1(X1))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_148, c_0_136])).
cnf(c_0_154, plain, (v7_struct_0(X1)|v1_xboole_0(esk7_1(X1))|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_131, c_0_149])).
cnf(c_0_155, negated_conjecture, (X1=esk3_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_136, c_0_142])).
cnf(c_0_156, negated_conjecture, (v1_xboole_0(esk6_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_46]), c_0_142])]), c_0_152])).
cnf(c_0_157, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk7_1(X1))|~v1_xboole_0(esk6_1(X1))), inference(er,[status(thm)],[c_0_153])).
cnf(c_0_158, negated_conjecture, (v1_xboole_0(esk7_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_151]), c_0_46]), c_0_142])]), c_0_152])).
cnf(c_0_159, negated_conjecture, (esk6_1(esk1_0)=esk3_0), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_160, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_46]), c_0_159]), c_0_142])]), c_0_152]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 161
# Proof object clause steps            : 119
# Proof object formula steps           : 42
# Proof object conjectures             : 63
# Proof object clause conjectures      : 60
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 39
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 66
# Proof object simplifying inferences  : 140
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 49
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 47
# Processed clauses                    : 1710
# ...of these trivial                  : 4
# ...subsumed                          : 1022
# ...remaining for further processing  : 684
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 55
# Backward-rewritten                   : 232
# Generated clauses                    : 2541
# ...of the previous two non-trivial   : 2434
# Contextual simplify-reflections      : 112
# Paramodulations                      : 2526
# Factorizations                       : 0
# Equation resolutions                 : 15
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 354
#    Positive orientable unit clauses  : 16
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 330
# Current number of unprocessed clauses: 784
# ...number of literals in the above   : 5216
# Current number of archived formulas  : 0
# Current number of archived clauses   : 332
# Clause-clause subsumption calls (NU) : 146324
# Rec. Clause-clause subsumption calls : 5717
# Non-unit clause-clause subsumptions  : 1167
# Unit Clause-clause subsumption calls : 1934
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 75991

# -------------------------------------------------
# User time                : 0.104 s
# System time              : 0.003 s
# Total time               : 0.107 s
# Maximum resident set size: 28728 pages
