# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t23_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t23_relat_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t69_enumset1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t71_enumset1)).
fof(t36_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))&k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t36_zfmisc_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t12_setfam_1)).
fof(t129_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k1_tarski(X4)))<=>(r2_hidden(X1,X3)&X2=X4)), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t129_zfmisc_1)).
fof(t34_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),k1_tarski(X4)))<=>(X1=X3&X2=X4)), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t34_zfmisc_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t100_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', d5_xboole_0)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', d4_relat_1)).
fof(t38_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_tarski(X1,X2),X3)<=>(r2_hidden(X1,X3)&r2_hidden(X2,X3))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t38_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', d10_xboole_0)).
fof(t11_setfam_1, axiom, ![X1]:k1_setfam_1(k1_tarski(X1))=X1, file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t11_setfam_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t3_boole)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t2_boole)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t113_zfmisc_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', d5_relat_1)).
fof(l44_zfmisc_1, axiom, ![X1, X2]:~((~(X1=k1_tarski(X2))&(~(X1=k1_xboole_0)&![X3]:~((r2_hidden(X3,X1)&~(X3=X2)))))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', l44_zfmisc_1)).
fof(t21_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', t21_relat_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/e6aba9287de44ca7a9fdb6e9e1b21a02.E_input', d3_tarski)).
fof(c_0_21, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t23_relat_1])).
fof(c_0_22, negated_conjecture, (v1_relat_1(esk3_0)&(esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))&(k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])).
fof(c_0_23, plain, ![X10]:k2_tarski(X10,X10)=k1_tarski(X10), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_24, plain, ![X14, X15]:k1_enumset1(X14,X14,X15)=k2_tarski(X14,X15), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_25, plain, ![X61, X62, X63]:k2_enumset1(X61,X61,X62,X63)=k1_enumset1(X61,X62,X63), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_26, plain, ![X11, X12, X13]:(k2_zfmisc_1(k1_tarski(X11),k2_tarski(X12,X13))=k2_tarski(k4_tarski(X11,X12),k4_tarski(X11,X13))&k2_zfmisc_1(k2_tarski(X11,X12),k1_tarski(X13))=k2_tarski(k4_tarski(X11,X13),k4_tarski(X12,X13))), inference(variable_rename,[status(thm)],[t36_zfmisc_1])).
fof(c_0_27, plain, ![X45, X46]:k1_setfam_1(k2_tarski(X45,X46))=k3_xboole_0(X45,X46), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_28, plain, ![X53, X54, X55, X56]:(((r2_hidden(X53,X55)|~r2_hidden(k4_tarski(X53,X54),k2_zfmisc_1(X55,k1_tarski(X56))))&(X54=X56|~r2_hidden(k4_tarski(X53,X54),k2_zfmisc_1(X55,k1_tarski(X56)))))&(~r2_hidden(X53,X55)|X54!=X56|r2_hidden(k4_tarski(X53,X54),k2_zfmisc_1(X55,k1_tarski(X56))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t129_zfmisc_1])])])).
fof(c_0_29, plain, ![X57, X58, X59, X60]:(((X57=X59|~r2_hidden(k4_tarski(X57,X58),k2_zfmisc_1(k1_tarski(X59),k1_tarski(X60))))&(X58=X60|~r2_hidden(k4_tarski(X57,X58),k2_zfmisc_1(k1_tarski(X59),k1_tarski(X60)))))&(X57!=X59|X58!=X60|r2_hidden(k4_tarski(X57,X58),k2_zfmisc_1(k1_tarski(X59),k1_tarski(X60))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_zfmisc_1])])])).
cnf(c_0_30, negated_conjecture, (esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_31, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_35, plain, ![X29, X30]:k4_xboole_0(X29,X30)=k5_xboole_0(X29,k3_xboole_0(X29,X30)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_36, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k1_tarski(X4)))|~r2_hidden(X1,X2)|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X4)))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, negated_conjecture, (esk3_0=k2_enumset1(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_40, plain, (k2_zfmisc_1(k2_enumset1(X1,X1,X1,X2),k2_enumset1(X3,X3,X3,X3))=k2_enumset1(k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_31]), c_0_32]), c_0_32]), c_0_32]), c_0_33]), c_0_33]), c_0_33])).
fof(c_0_41, plain, ![X71, X72, X73, X74, X75, X76, X77, X78]:((((r2_hidden(X74,X71)|~r2_hidden(X74,X73)|X73!=k4_xboole_0(X71,X72))&(~r2_hidden(X74,X72)|~r2_hidden(X74,X73)|X73!=k4_xboole_0(X71,X72)))&(~r2_hidden(X75,X71)|r2_hidden(X75,X72)|r2_hidden(X75,X73)|X73!=k4_xboole_0(X71,X72)))&((~r2_hidden(esk12_3(X76,X77,X78),X78)|(~r2_hidden(esk12_3(X76,X77,X78),X76)|r2_hidden(esk12_3(X76,X77,X78),X77))|X78=k4_xboole_0(X76,X77))&((r2_hidden(esk12_3(X76,X77,X78),X76)|r2_hidden(esk12_3(X76,X77,X78),X78)|X78=k4_xboole_0(X76,X77))&(~r2_hidden(esk12_3(X76,X77,X78),X77)|r2_hidden(esk12_3(X76,X77,X78),X78)|X78=k4_xboole_0(X76,X77))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_42, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_36, c_0_32])).
cnf(c_0_44, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,k2_enumset1(X4,X4,X4,X4)))|X3!=X4|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_45, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),k2_enumset1(X4,X4,X4,X4)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_31]), c_0_31]), c_0_32]), c_0_32]), c_0_33]), c_0_33])).
cnf(c_0_46, negated_conjecture, (k2_zfmisc_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=esk3_0), inference(rw,[status(thm)],[c_0_39, c_0_40])).
fof(c_0_47, plain, ![X18, X19, X20, X22, X23, X24, X25, X27]:(((~r2_hidden(X20,X19)|r2_hidden(k4_tarski(X20,esk4_3(X18,X19,X20)),X18)|X19!=k1_relat_1(X18))&(~r2_hidden(k4_tarski(X22,X23),X18)|r2_hidden(X22,X19)|X19!=k1_relat_1(X18)))&((~r2_hidden(esk5_2(X24,X25),X25)|~r2_hidden(k4_tarski(esk5_2(X24,X25),X27),X24)|X25=k1_relat_1(X24))&(r2_hidden(esk5_2(X24,X25),X25)|r2_hidden(k4_tarski(esk5_2(X24,X25),esk6_2(X24,X25)),X24)|X25=k1_relat_1(X24)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_48, plain, ![X47, X48, X49]:(((r2_hidden(X47,X49)|~r1_tarski(k2_tarski(X47,X48),X49))&(r2_hidden(X48,X49)|~r1_tarski(k2_tarski(X47,X48),X49)))&(~r2_hidden(X47,X49)|~r2_hidden(X48,X49)|r1_tarski(k2_tarski(X47,X48),X49))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_zfmisc_1])])])).
fof(c_0_49, plain, ![X8, X9]:(((r1_tarski(X8,X9)|X8!=X9)&(r1_tarski(X9,X8)|X8!=X9))&(~r1_tarski(X8,X9)|~r1_tarski(X9,X8)|X8=X9)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_50, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X4)))|X1!=X2|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_51, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_53, plain, ![X42]:k1_setfam_1(k1_tarski(X42))=X42, inference(variable_rename,[status(thm)],[t11_setfam_1])).
fof(c_0_54, plain, ![X64]:k4_xboole_0(X64,k1_xboole_0)=X64, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_55, plain, ![X17]:k3_xboole_0(X17,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_56, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k2_enumset1(X2,X2,X2,X2)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_44])).
cnf(c_0_57, negated_conjecture, (X1=esk1_0|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_58, plain, (r2_hidden(esk5_2(X1,X2),X2)|r2_hidden(k4_tarski(esk5_2(X1,X2),esk6_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_59, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_60, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_61, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),k2_enumset1(X4,X4,X4,X4)))|X3!=X4|X1!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_31]), c_0_31]), c_0_32]), c_0_32]), c_0_33]), c_0_33])).
fof(c_0_62, plain, ![X43, X44]:((k2_zfmisc_1(X43,X44)!=k1_xboole_0|(X43=k1_xboole_0|X44=k1_xboole_0))&((X43!=k1_xboole_0|k2_zfmisc_1(X43,X44)=k1_xboole_0)&(X44!=k1_xboole_0|k2_zfmisc_1(X43,X44)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_63, plain, (X3!=k5_xboole_0(X4,k1_setfam_1(k2_enumset1(X4,X4,X4,X2)))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52]), c_0_33])).
cnf(c_0_64, plain, (k1_setfam_1(k1_tarski(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_67, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_68, plain, ![X31, X32, X33, X35, X36, X37, X38, X40]:(((~r2_hidden(X33,X32)|r2_hidden(k4_tarski(esk7_3(X31,X32,X33),X33),X31)|X32!=k2_relat_1(X31))&(~r2_hidden(k4_tarski(X36,X35),X31)|r2_hidden(X35,X32)|X32!=k2_relat_1(X31)))&((~r2_hidden(esk8_2(X37,X38),X38)|~r2_hidden(k4_tarski(X40,esk8_2(X37,X38)),X37)|X38=k2_relat_1(X37))&(r2_hidden(esk8_2(X37,X38),X38)|r2_hidden(k4_tarski(esk9_2(X37,X38),esk8_2(X37,X38)),X37)|X38=k2_relat_1(X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_69, negated_conjecture, (r2_hidden(k4_tarski(X1,esk2_0),esk3_0)|~r2_hidden(X1,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_56, c_0_46])).
cnf(c_0_70, negated_conjecture, (esk5_2(esk3_0,X1)=esk1_0|X1=k1_relat_1(esk3_0)|r2_hidden(esk5_2(esk3_0,X1),X1)), inference(spm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_71, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X1,X1,X1,X3),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_32]), c_0_33])).
cnf(c_0_72, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_73, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_74, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X2,X2)))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_61])])).
cnf(c_0_75, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_tarski(X3,X1),X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_76, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_77, plain, (~r2_hidden(X1,k5_xboole_0(X2,k1_setfam_1(k2_enumset1(X2,X2,X2,X3))))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_63])).
cnf(c_0_78, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_79, plain, (k5_xboole_0(X1,k1_setfam_1(k2_enumset1(X1,X1,X1,k1_xboole_0)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_52]), c_0_33])).
cnf(c_0_80, plain, (k1_setfam_1(k2_enumset1(X1,X1,X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_43]), c_0_33])).
cnf(c_0_81, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k2_enumset1(X2,X2,X2,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_82, plain, (r2_hidden(k4_tarski(esk7_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_83, plain, ![X50, X51]:((r2_hidden(esk10_2(X50,X51),X50)|X50=k1_xboole_0|X50=k1_tarski(X51))&(esk10_2(X50,X51)!=X51|X50=k1_xboole_0|X50=k1_tarski(X51))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l44_zfmisc_1])])])])])).
cnf(c_0_84, plain, (X2=k1_relat_1(X1)|~r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(k4_tarski(esk5_2(X1,X2),X3),X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_85, negated_conjecture, (esk5_2(esk3_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))=esk1_0|k1_relat_1(esk3_0)=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_57])).
cnf(c_0_86, plain, (r2_hidden(X1,k2_enumset1(X1,X1,X1,X2))), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_87, plain, (r2_hidden(k4_tarski(X1,esk4_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_88, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_73])).
cnf(c_0_89, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_74, c_0_46])).
cnf(c_0_90, plain, (r2_hidden(X1,X2)|~r1_tarski(k2_enumset1(X3,X3,X3,X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_32]), c_0_33])).
cnf(c_0_91, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0|k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)=k1_xboole_0|esk3_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_76, c_0_46])).
cnf(c_0_92, plain, (~r2_hidden(X1,k5_xboole_0(X2,X2))|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_93, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_79, c_0_80])).
fof(c_0_94, plain, ![X16]:(~v1_relat_1(X16)|r1_tarski(X16,k2_zfmisc_1(k1_relat_1(X16),k2_relat_1(X16)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_relat_1])])).
cnf(c_0_95, negated_conjecture, (X1=esk2_0|~r2_hidden(k4_tarski(X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_81, c_0_46])).
cnf(c_0_96, plain, (r2_hidden(k4_tarski(esk7_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_82])).
cnf(c_0_97, plain, (r2_hidden(esk10_2(X1,X2),X1)|X1=k1_xboole_0|X1=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_98, negated_conjecture, (k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_99, negated_conjecture, (k1_relat_1(esk3_0)=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|~r2_hidden(k4_tarski(esk1_0,X1),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85]), c_0_86])])).
cnf(c_0_100, plain, (r2_hidden(k4_tarski(X1,esk4_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_87])).
cnf(c_0_101, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_102, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0|r2_hidden(esk2_0,X1)|esk3_0!=k1_xboole_0|~r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_103, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_104, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_105, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_106, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|esk10_2(X1,X2)!=X2), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_107, negated_conjecture, (X1=esk2_0|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_108, plain, (X1=k1_xboole_0|X1=k2_enumset1(X2,X2,X2,X2)|r2_hidden(esk10_2(X1,X2),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_109, negated_conjecture, (k1_relat_1(esk3_0)!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|k2_relat_1(esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_31]), c_0_31]), c_0_32]), c_0_32]), c_0_33]), c_0_33])).
cnf(c_0_110, negated_conjecture, (k1_relat_1(esk3_0)=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])])).
fof(c_0_111, plain, ![X65, X66, X67, X68, X69]:((~r1_tarski(X65,X66)|(~r2_hidden(X67,X65)|r2_hidden(X67,X66)))&((r2_hidden(esk11_2(X68,X69),X68)|r1_tarski(X68,X69))&(~r2_hidden(esk11_2(X68,X69),X69)|r1_tarski(X68,X69)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_112, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)=k1_xboole_0|esk3_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_72]), c_0_103])).
cnf(c_0_113, plain, (k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))=X1|~r1_tarski(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)),X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_114, negated_conjecture, (v1_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_115, plain, (X1=k1_xboole_0|X1=k2_enumset1(X2,X2,X2,X2)|esk10_2(X1,X2)!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_31]), c_0_32]), c_0_33])).
cnf(c_0_116, negated_conjecture, (k2_relat_1(esk3_0)=k2_enumset1(X1,X1,X1,X1)|esk10_2(k2_relat_1(esk3_0),X1)=esk2_0|k2_relat_1(esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_117, negated_conjecture, (k2_relat_1(esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_118, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_119, plain, (r2_hidden(esk11_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_120, negated_conjecture, (r2_hidden(esk1_0,X1)|esk3_0!=k1_xboole_0|~r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_90, c_0_112])).
cnf(c_0_121, negated_conjecture, (k2_zfmisc_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_relat_1(esk3_0))=esk3_0|~r1_tarski(k2_zfmisc_1(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k2_relat_1(esk3_0)),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_110]), c_0_114])])).
cnf(c_0_122, negated_conjecture, (k2_relat_1(esk3_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116])]), c_0_117])).
cnf(c_0_123, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_118])).
cnf(c_0_124, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_103, c_0_119])).
cnf(c_0_125, negated_conjecture, (esk3_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_72]), c_0_103])).
cnf(c_0_126, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_121, c_0_122]), c_0_123]), c_0_122]), c_0_123]), c_0_124])]), c_0_125]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 127
# Proof object clause steps            : 84
# Proof object formula steps           : 43
# Proof object conjectures             : 29
# Proof object clause conjectures      : 26
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 32
# Proof object initial formulas used   : 21
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 86
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 48
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 43
# Processed clauses                    : 1767
# ...of these trivial                  : 3
# ...subsumed                          : 1031
# ...remaining for further processing  : 733
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 29
# Backward-rewritten                   : 350
# Generated clauses                    : 15926
# ...of the previous two non-trivial   : 15412
# Contextual simplify-reflections      : 15
# Paramodulations                      : 15894
# Factorizations                       : 16
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 300
#    Positive orientable unit clauses  : 40
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 258
# Current number of unprocessed clauses: 13437
# ...number of literals in the above   : 48137
# Current number of archived formulas  : 0
# Current number of archived clauses   : 425
# Clause-clause subsumption calls (NU) : 27337
# Rec. Clause-clause subsumption calls : 16804
# Non-unit clause-clause subsumptions  : 552
# Unit Clause-clause subsumption calls : 266
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 80
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 326828

# -------------------------------------------------
# User time                : 0.196 s
# System time              : 0.010 s
# Total time               : 0.206 s
# Maximum resident set size: 26244 pages
