# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d15_waybel_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k3_waybel_0(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(m1_subset_1(X5,u1_struct_0(X1))&(r1_orders_2(X1,X4,X5)&r2_hidden(X5,X2))))))))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', d15_waybel_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', d1_xboole_0)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', d9_orders_2)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', cc4_relset_1)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', dt_u1_orders_2)).
fof(d17_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2)))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', d17_waybel_0)).
fof(dt_k5_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', dt_k5_waybel_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', d1_tarski)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', dt_k6_domain_1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', redefinition_k6_domain_1)).
fof(t17_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2))))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', t17_waybel_0)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/e7c4a9a0cf8b46d49ed4e0734cdbb5c6.E_input', t5_subset)).
fof(c_0_12, plain, ![X6, X7, X8, X9, X11, X13]:((((m1_subset_1(esk1_4(X6,X7,X8,X9),u1_struct_0(X6))|~r2_hidden(X9,X8)|~m1_subset_1(X9,u1_struct_0(X6))|X8!=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((r1_orders_2(X6,X9,esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,X8)|~m1_subset_1(X9,u1_struct_0(X6))|X8!=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&(r2_hidden(esk1_4(X6,X7,X8,X9),X7)|~r2_hidden(X9,X8)|~m1_subset_1(X9,u1_struct_0(X6))|X8!=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))))&(~m1_subset_1(X11,u1_struct_0(X6))|(~r1_orders_2(X6,X9,X11)|~r2_hidden(X11,X7))|r2_hidden(X9,X8)|~m1_subset_1(X9,u1_struct_0(X6))|X8!=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6)))&((m1_subset_1(esk2_3(X6,X7,X8),u1_struct_0(X6))|X8=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((~r2_hidden(esk2_3(X6,X7,X8),X8)|(~m1_subset_1(X13,u1_struct_0(X6))|(~r1_orders_2(X6,esk2_3(X6,X7,X8),X13)|~r2_hidden(X13,X7)))|X8=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((m1_subset_1(esk3_3(X6,X7,X8),u1_struct_0(X6))|r2_hidden(esk2_3(X6,X7,X8),X8)|X8=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&((r1_orders_2(X6,esk2_3(X6,X7,X8),esk3_3(X6,X7,X8))|r2_hidden(esk2_3(X6,X7,X8),X8)|X8=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))&(r2_hidden(esk3_3(X6,X7,X8),X7)|r2_hidden(esk2_3(X6,X7,X8),X8)|X8=k3_waybel_0(X6,X7)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X6)))|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X6)))|~l1_orders_2(X6))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d15_waybel_0])])])])])).
fof(c_0_13, plain, ![X33, X34, X35]:((~v1_xboole_0(X33)|~r2_hidden(X34,X33))&(r2_hidden(esk5_1(X35),X35)|v1_xboole_0(X35))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_14, plain, ![X40, X41, X42]:((~r1_orders_2(X40,X41,X42)|r2_hidden(k4_tarski(X41,X42),u1_orders_2(X40))|~m1_subset_1(X42,u1_struct_0(X40))|~m1_subset_1(X41,u1_struct_0(X40))|~l1_orders_2(X40))&(~r2_hidden(k4_tarski(X41,X42),u1_orders_2(X40))|r1_orders_2(X40,X41,X42)|~m1_subset_1(X42,u1_struct_0(X40))|~m1_subset_1(X41,u1_struct_0(X40))|~l1_orders_2(X40))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
fof(c_0_15, plain, ![X43, X44, X45]:(~v1_xboole_0(X43)|(~m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(X44,X43)))|v1_xboole_0(X45))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
fof(c_0_16, plain, ![X46]:(~l1_orders_2(X46)|m1_subset_1(u1_orders_2(X46),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X46),u1_struct_0(X46))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
cnf(c_0_17, plain, (r2_hidden(X3,X5)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,X3,X1)|~r2_hidden(X1,X4)|~m1_subset_1(X3,u1_struct_0(X2))|X5!=k3_waybel_0(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_18, plain, ![X15, X16]:(v2_struct_0(X15)|~l1_orders_2(X15)|(~m1_subset_1(X16,u1_struct_0(X15))|k5_waybel_0(X15,X16)=k3_waybel_0(X15,k6_domain_1(u1_struct_0(X15),X16)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_waybel_0])])])])).
fof(c_0_19, plain, ![X24, X25]:(v2_struct_0(X24)|~l1_orders_2(X24)|~m1_subset_1(X25,u1_struct_0(X24))|m1_subset_1(k5_waybel_0(X24,X25),k1_zfmisc_1(u1_struct_0(X24)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k5_waybel_0])])])).
cnf(c_0_20, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X26, X27, X28, X29, X30, X31]:(((~r2_hidden(X28,X27)|X28=X26|X27!=k1_tarski(X26))&(X29!=X26|r2_hidden(X29,X27)|X27!=k1_tarski(X26)))&((~r2_hidden(esk4_2(X30,X31),X31)|esk4_2(X30,X31)!=X30|X31=k1_tarski(X30))&(r2_hidden(esk4_2(X30,X31),X31)|esk4_2(X30,X31)=X30|X31=k1_tarski(X30)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_25, plain, (r2_hidden(X1,k3_waybel_0(X2,X3))|~r1_orders_2(X2,X1,X4)|~r2_hidden(X4,X3)|~m1_subset_1(k3_waybel_0(X2,X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (v2_struct_0(X1)|k5_waybel_0(X1,X2)=k3_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (v2_struct_0(X1)|m1_subset_1(k5_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_28, plain, ![X19, X20]:(v1_xboole_0(X19)|~m1_subset_1(X20,X19)|m1_subset_1(k6_domain_1(X19,X20),k1_zfmisc_1(X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_29, plain, (~v1_xboole_0(u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_30, plain, (v1_xboole_0(u1_orders_2(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X4)|~r2_hidden(X4,k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_27])).
cnf(c_0_33, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (~v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
fof(c_0_35, plain, ![X17, X18]:(v1_xboole_0(X17)|~m1_subset_1(X18,X17)|k6_domain_1(X17,X18)=k1_tarski(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_36, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_31])])).
cnf(c_0_38, plain, (r2_hidden(esk1_4(X1,X2,X3,X4),X2)|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k3_waybel_0(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_39, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k5_waybel_0(X1,X2))<=>r1_orders_2(X1,X3,X2)))))), inference(assume_negation,[status(cth)],[t17_waybel_0])).
cnf(c_0_40, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X4)|~r2_hidden(X4,k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_41, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_44, plain, (~v1_xboole_0(k1_tarski(X1))), inference(spm,[status(thm)],[c_0_20, c_0_37])).
cnf(c_0_45, plain, (r1_orders_2(X1,X2,esk1_4(X1,X3,X4,X2))|~r2_hidden(X2,X4)|~m1_subset_1(X2,u1_struct_0(X1))|X4!=k3_waybel_0(X1,X3)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_46, plain, (r2_hidden(esk1_4(X1,X2,k3_waybel_0(X1,X2),X3),X2)|~r2_hidden(X3,k3_waybel_0(X1,X2))|~m1_subset_1(k3_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_38])).
fof(c_0_47, plain, ![X21, X22, X23]:(~r2_hidden(X21,X22)|(~m1_subset_1(X22,k1_zfmisc_1(X23))|~v1_xboole_0(X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_48, negated_conjecture, ((~v2_struct_0(esk6_0)&l1_orders_2(esk6_0))&(m1_subset_1(esk7_0,u1_struct_0(esk6_0))&(m1_subset_1(esk8_0,u1_struct_0(esk6_0))&((~r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,esk8_0,esk7_0))&(r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))|r1_orders_2(esk6_0,esk8_0,esk7_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_39])])])])).
cnf(c_0_49, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X4)|~r2_hidden(X4,k1_tarski(X3))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_34])).
cnf(c_0_50, plain, (esk5_1(k1_tarski(X1))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])).
cnf(c_0_51, plain, (r1_orders_2(X1,X2,esk1_4(X1,X3,k3_waybel_0(X1,X3),X2))|~r2_hidden(X2,k3_waybel_0(X1,X3))|~m1_subset_1(k3_waybel_0(X1,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (esk1_4(X1,k1_tarski(X2),k3_waybel_0(X1,k1_tarski(X2)),X3)=X2|~r2_hidden(X3,k3_waybel_0(X1,k1_tarski(X2)))|~m1_subset_1(k3_waybel_0(X1,k1_tarski(X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_42, c_0_46])).
cnf(c_0_53, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54, negated_conjecture, (~r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))|~r1_orders_2(esk6_0,esk8_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (v2_struct_0(X1)|r2_hidden(X2,k5_waybel_0(X1,X3))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_43]), c_0_50]), c_0_50])]), c_0_44])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(esk8_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, negated_conjecture, (l1_orders_2(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, negated_conjecture, (~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_60, plain, (r1_orders_2(X1,X2,X3)|~r2_hidden(X2,k3_waybel_0(X1,k1_tarski(X3)))|~m1_subset_1(k3_waybel_0(X1,k1_tarski(X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_61, plain, (k3_waybel_0(X1,k1_tarski(X2))=k5_waybel_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_26, c_0_41])).
cnf(c_0_62, plain, (v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_33, c_0_41])).
cnf(c_0_63, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,k5_waybel_0(X1,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_53, c_0_27])).
cnf(c_0_64, negated_conjecture, (r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))|r1_orders_2(esk6_0,esk8_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_65, negated_conjecture, (~r1_orders_2(esk6_0,esk8_0,esk7_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_56]), c_0_57]), c_0_58])]), c_0_59])).
cnf(c_0_66, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~r2_hidden(X2,k5_waybel_0(X1,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_27]), c_0_63])).
cnf(c_0_67, negated_conjecture, (r2_hidden(esk8_0,k5_waybel_0(esk6_0,esk7_0))), inference(sr,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_68, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_56]), c_0_57]), c_0_58])]), c_0_59]), c_0_65]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 44
# Proof object formula steps           : 25
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 30
# Processed clauses                    : 128
# ...of these trivial                  : 0
# ...subsumed                          : 20
# ...remaining for further processing  : 108
# Other redundant clauses eliminated   : 11
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 118
# ...of the previous two non-trivial   : 94
# Contextual simplify-reflections      : 8
# Paramodulations                      : 105
# Factorizations                       : 2
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 68
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 57
# Current number of unprocessed clauses: 24
# ...number of literals in the above   : 198
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 3944
# Rec. Clause-clause subsumption calls : 176
# Non-unit clause-clause subsumptions  : 24
# Unit Clause-clause subsumption calls : 46
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 6805

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28660 pages
