# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C18_F1_AE_CS_SP_PI_PS_S0S
# and selection function SelectComplexG.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t33_tops_2, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_pre_topc(X3,X1)=>(v3_pre_topc(X2,X1)=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))=>(X4=X2=>v3_pre_topc(X4,X3))))))), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', t33_tops_2)).
fof(redefinition_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k3_xboole_0(X2,X3)), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', redefinition_k9_subset_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', t12_setfam_1)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', l3_subset_1)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', dt_m1_pre_topc)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', d3_tarski)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', dt_l1_pre_topc)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', t28_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', commutativity_k2_tarski)).
fof(commutativity_k9_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>k9_subset_1(X1,X2,X3)=k9_subset_1(X1,X3,X2)), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', commutativity_k9_subset_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', d3_struct_0)).
fof(t32_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))=>(v3_pre_topc(X3,X2)<=>?[X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))&(v3_pre_topc(X4,X1)&k9_subset_1(u1_struct_0(X2),X4,k2_struct_0(X2))=X3)))))), file('data/MPTP/proofs/e7e9c12455a541989da3a664ea6df4eb.E_input', t32_tops_2)).
fof(c_0_12, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_pre_topc(X3,X1)=>(v3_pre_topc(X2,X1)=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))=>(X4=X2=>v3_pre_topc(X4,X3)))))))), inference(assume_negation,[status(cth)],[t33_tops_2])).
fof(c_0_13, negated_conjecture, (l1_pre_topc(esk1_0)&(m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))&(m1_pre_topc(esk3_0,esk1_0)&(v3_pre_topc(esk2_0,esk1_0)&(m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))&(esk4_0=esk2_0&~v3_pre_topc(esk4_0,esk3_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
fof(c_0_14, plain, ![X25, X26, X27]:(~m1_subset_1(X27,k1_zfmisc_1(X25))|k9_subset_1(X25,X26,X27)=k3_xboole_0(X26,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_subset_1])])).
fof(c_0_15, plain, ![X11, X12]:k1_setfam_1(k2_tarski(X11,X12))=k3_xboole_0(X11,X12), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_16, plain, ![X33, X34, X35]:(~m1_subset_1(X34,k1_zfmisc_1(X33))|(~r2_hidden(X35,X34)|r2_hidden(X35,X33))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
cnf(c_0_17, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, negated_conjecture, (esk4_0=esk2_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (k9_subset_1(X2,X3,X1)=k3_xboole_0(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_21, plain, ![X37, X38]:(~l1_pre_topc(X37)|(~m1_pre_topc(X38,X37)|l1_pre_topc(X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_22, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk3_0)))), inference(rw,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_24, plain, ![X13, X14, X15, X16, X17]:((~r1_tarski(X13,X14)|(~r2_hidden(X15,X13)|r2_hidden(X15,X14)))&((r2_hidden(esk5_2(X16,X17),X16)|r1_tarski(X16,X17))&(~r2_hidden(esk5_2(X16,X17),X17)|r1_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_25, plain, (k9_subset_1(X2,X3,X1)=k1_setfam_1(k2_tarski(X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_27, plain, ![X36]:(~l1_pre_topc(X36)|l1_struct_0(X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_28, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, negated_conjecture, (m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_30, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_31, plain, ![X28, X29]:(~r1_tarski(X28,X29)|k3_xboole_0(X28,X29)=X28), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_32, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk3_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_33, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_34, plain, ![X9, X10]:k2_tarski(X9,X10)=k2_tarski(X10,X9), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_35, plain, ![X30, X31, X32]:(~m1_subset_1(X32,k1_zfmisc_1(X30))|k9_subset_1(X30,X31,X32)=k9_subset_1(X30,X32,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k9_subset_1])])).
cnf(c_0_36, negated_conjecture, (k1_setfam_1(k2_tarski(X1,esk2_0))=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
fof(c_0_37, plain, ![X24]:(~l1_struct_0(X24)|k2_struct_0(X24)=u1_struct_0(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_38, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_39, negated_conjecture, (l1_pre_topc(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])])).
cnf(c_0_40, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_42, negated_conjecture, (r2_hidden(esk5_2(esk2_0,X1),u1_struct_0(esk3_0))|r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_43, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (k9_subset_1(X2,X3,X1)=k9_subset_1(X2,X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_45, plain, ![X19, X20, X21, X23]:(((m1_subset_1(esk6_3(X19,X20,X21),k1_zfmisc_1(u1_struct_0(X19)))|~v3_pre_topc(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~m1_pre_topc(X20,X19)|~l1_pre_topc(X19))&((v3_pre_topc(esk6_3(X19,X20,X21),X19)|~v3_pre_topc(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~m1_pre_topc(X20,X19)|~l1_pre_topc(X19))&(k9_subset_1(u1_struct_0(X20),esk6_3(X19,X20,X21),k2_struct_0(X20))=X21|~v3_pre_topc(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~m1_pre_topc(X20,X19)|~l1_pre_topc(X19))))&(~m1_subset_1(X23,k1_zfmisc_1(u1_struct_0(X19)))|(~v3_pre_topc(X23,X19)|k9_subset_1(u1_struct_0(X20),X23,k2_struct_0(X20))!=X21)|v3_pre_topc(X21,X20)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X20)))|~m1_pre_topc(X20,X19)|~l1_pre_topc(X19))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t32_tops_2])])])])])).
cnf(c_0_46, negated_conjecture, (k9_subset_1(u1_struct_0(esk3_0),X1,esk2_0)=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_23]), c_0_36])).
cnf(c_0_47, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, negated_conjecture, (l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_49, plain, (k1_setfam_1(k2_tarski(X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_20])).
cnf(c_0_50, negated_conjecture, (r1_tarski(esk2_0,u1_struct_0(esk3_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51, negated_conjecture, (k1_setfam_1(k2_tarski(esk2_0,X1))=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(spm,[status(thm)],[c_0_36, c_0_43])).
cnf(c_0_52, negated_conjecture, (k9_subset_1(u1_struct_0(esk1_0),esk2_0,X1)=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(spm,[status(thm)],[c_0_44, c_0_26])).
cnf(c_0_53, plain, (v3_pre_topc(X4,X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v3_pre_topc(X1,X2)|k9_subset_1(u1_struct_0(X3),X1,k2_struct_0(X3))!=X4|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, negated_conjecture, (k9_subset_1(u1_struct_0(esk3_0),esk2_0,X1)=k9_subset_1(u1_struct_0(esk1_0),X1,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_23]), c_0_46])).
cnf(c_0_55, negated_conjecture, (k2_struct_0(esk3_0)=u1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_56, negated_conjecture, (k9_subset_1(u1_struct_0(esk1_0),esk2_0,u1_struct_0(esk3_0))=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_57, negated_conjecture, (~v3_pre_topc(esk4_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_58, negated_conjecture, (v3_pre_topc(X1,esk3_0)|esk2_0!=X1|~v3_pre_topc(esk2_0,X2)|~m1_pre_topc(esk3_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk3_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_52]), c_0_56])).
cnf(c_0_59, negated_conjecture, (~v3_pre_topc(esk2_0,esk3_0)), inference(rw,[status(thm)],[c_0_57, c_0_18])).
cnf(c_0_60, negated_conjecture, (~v3_pre_topc(esk2_0,X1)|~m1_pre_topc(esk3_0,X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_23]), c_0_59])).
cnf(c_0_61, negated_conjecture, (v3_pre_topc(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_62, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_26]), c_0_61]), c_0_29]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 38
# Proof object formula steps           : 25
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 12
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 23
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 22
# Processed clauses                    : 85
# ...of these trivial                  : 1
# ...subsumed                          : 8
# ...remaining for further processing  : 76
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 1
# Generated clauses                    : 55
# ...of the previous two non-trivial   : 47
# Contextual simplify-reflections      : 0
# Paramodulations                      : 53
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 52
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 27
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 36
# Current number of archived formulas  : 0
# Current number of archived clauses   : 25
# Clause-clause subsumption calls (NU) : 344
# Rec. Clause-clause subsumption calls : 120
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 51
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2659

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28796 pages
