# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', d10_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', t28_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', commutativity_k3_xboole_0)).
fof(t123_zfmisc_1, axiom, ![X1, X2, X3, X4]:k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4)), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', t123_zfmisc_1)).
fof(t134_zfmisc_1, conjecture, ![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4)))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', t134_zfmisc_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', t3_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', t65_xboole_1)).
fof(t6_xboole_0, axiom, ![X1, X2]:~((r2_xboole_0(X1,X2)&![X3]:~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', t6_xboole_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', d1_tarski)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', d1_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', d5_xboole_0)).
fof(d8_xboole_0, axiom, ![X1, X2]:(r2_xboole_0(X1,X2)<=>(r1_tarski(X1,X2)&~(X1=X2))), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', d8_xboole_0)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', t113_zfmisc_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/e879aced589a4fca94cdf05b038ee7cf.E_input', t2_boole)).
fof(c_0_16, plain, ![X48, X49, X50, X51, X52, X53, X54, X55]:((((r2_hidden(X51,X48)|~r2_hidden(X51,X50)|X50!=k3_xboole_0(X48,X49))&(r2_hidden(X51,X49)|~r2_hidden(X51,X50)|X50!=k3_xboole_0(X48,X49)))&(~r2_hidden(X52,X48)|~r2_hidden(X52,X49)|r2_hidden(X52,X50)|X50!=k3_xboole_0(X48,X49)))&((~r2_hidden(esk9_3(X53,X54,X55),X55)|(~r2_hidden(esk9_3(X53,X54,X55),X53)|~r2_hidden(esk9_3(X53,X54,X55),X54))|X55=k3_xboole_0(X53,X54))&((r2_hidden(esk9_3(X53,X54,X55),X53)|r2_hidden(esk9_3(X53,X54,X55),X55)|X55=k3_xboole_0(X53,X54))&(r2_hidden(esk9_3(X53,X54,X55),X54)|r2_hidden(esk9_3(X53,X54,X55),X55)|X55=k3_xboole_0(X53,X54))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_17, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_18, plain, ![X13, X14, X15, X16, X17]:((~r1_tarski(X13,X14)|(~r2_hidden(X15,X13)|r2_hidden(X15,X14)))&((r2_hidden(esk5_2(X16,X17),X16)|r1_tarski(X16,X17))&(~r2_hidden(esk5_2(X16,X17),X17)|r1_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_19, plain, ![X9, X10]:(((r1_tarski(X9,X10)|X9!=X10)&(r1_tarski(X10,X9)|X9!=X10))&(~r1_tarski(X9,X10)|~r1_tarski(X10,X9)|X9=X10)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_20, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_22, plain, ![X35, X36]:(~r1_tarski(X35,X36)|k3_xboole_0(X35,X36)=X35), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_23, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (r1_tarski(X1,X2)|~r2_hidden(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (r2_hidden(esk5_2(k3_xboole_0(X1,X2),X3),X2)|r1_tarski(k3_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_26, plain, ![X61, X62]:k3_xboole_0(X61,X62)=k3_xboole_0(X62,X61), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_27, plain, ![X19, X20, X21, X22]:k2_zfmisc_1(k3_xboole_0(X19,X20),k3_xboole_0(X21,X22))=k3_xboole_0(k2_zfmisc_1(X19,X21),k2_zfmisc_1(X20,X22)), inference(variable_rename,[status(thm)],[t123_zfmisc_1])).
cnf(c_0_28, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_23])).
fof(c_0_30, negated_conjecture, ~(![X1, X2, X3, X4]:(k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X1=X3&X2=X4))))), inference(assume_negation,[status(cth)],[t134_zfmisc_1])).
cnf(c_0_31, plain, (r1_tarski(k3_xboole_0(X1,X2),X2)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (k2_zfmisc_1(k3_xboole_0(X1,X2),k3_xboole_0(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_35, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk1_0!=esk3_0|esk2_0!=esk4_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_30])])])).
cnf(c_0_36, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X1,X4),k2_zfmisc_1(X3,X2))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_32]), c_0_33])).
cnf(c_0_38, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X1,X3))=k2_zfmisc_1(X1,k3_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, negated_conjecture, (k2_zfmisc_1(esk1_0,esk2_0)=k2_zfmisc_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_40, plain, (r1_tarski(k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4)),k2_zfmisc_1(X1,X4))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(esk1_0,X1))=k2_zfmisc_1(esk1_0,k3_xboole_0(X1,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_32])).
fof(c_0_42, plain, ![X63, X64, X66, X67, X68]:(((r2_hidden(esk11_2(X63,X64),X63)|r1_xboole_0(X63,X64))&(r2_hidden(esk11_2(X63,X64),X64)|r1_xboole_0(X63,X64)))&(~r2_hidden(X68,X66)|~r2_hidden(X68,X67)|~r1_xboole_0(X66,X67))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_43, plain, ![X46]:r1_xboole_0(X46,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_44, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,k3_xboole_0(X1,esk2_0)),k2_zfmisc_1(esk3_0,X1))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_46, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_47, plain, ![X32, X33]:((r2_hidden(esk7_2(X32,X33),X33)|~r2_xboole_0(X32,X33))&(~r2_hidden(esk7_2(X32,X33),X32)|~r2_xboole_0(X32,X33))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_xboole_0])])])])])).
fof(c_0_48, plain, ![X25, X26, X27, X28, X29, X30]:(((~r2_hidden(X27,X26)|X27=X25|X26!=k1_tarski(X25))&(X28!=X25|r2_hidden(X28,X26)|X26!=k1_tarski(X25)))&((~r2_hidden(esk6_2(X29,X30),X30)|esk6_2(X29,X30)!=X29|X30=k1_tarski(X29))&(r2_hidden(esk6_2(X29,X30),X30)|esk6_2(X29,X30)=X29|X30=k1_tarski(X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_49, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,X1)),k2_zfmisc_1(esk3_0,X1))), inference(spm,[status(thm)],[c_0_44, c_0_32])).
cnf(c_0_50, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_31]), c_0_32])).
fof(c_0_51, plain, ![X57, X58, X59]:((~v1_xboole_0(X57)|~r2_hidden(X58,X57))&(r2_hidden(esk10_1(X59),X59)|v1_xboole_0(X59))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_52, plain, ![X37, X38, X39, X40, X41, X42, X43, X44]:((((r2_hidden(X40,X37)|~r2_hidden(X40,X39)|X39!=k4_xboole_0(X37,X38))&(~r2_hidden(X40,X38)|~r2_hidden(X40,X39)|X39!=k4_xboole_0(X37,X38)))&(~r2_hidden(X41,X37)|r2_hidden(X41,X38)|r2_hidden(X41,X39)|X39!=k4_xboole_0(X37,X38)))&((~r2_hidden(esk8_3(X42,X43,X44),X44)|(~r2_hidden(esk8_3(X42,X43,X44),X42)|r2_hidden(esk8_3(X42,X43,X44),X43))|X44=k4_xboole_0(X42,X43))&((r2_hidden(esk8_3(X42,X43,X44),X42)|r2_hidden(esk8_3(X42,X43,X44),X44)|X44=k4_xboole_0(X42,X43))&(~r2_hidden(esk8_3(X42,X43,X44),X43)|r2_hidden(esk8_3(X42,X43,X44),X44)|X44=k4_xboole_0(X42,X43))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_53, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_54, plain, (r2_hidden(esk7_2(X1,X2),X2)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_55, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X2))=k2_zfmisc_1(k3_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_58, negated_conjecture, (r1_tarski(k2_zfmisc_1(esk1_0,k3_xboole_0(X1,esk2_0)),k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk2_0)))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_60, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_61, plain, (r2_hidden(esk10_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_63, plain, (~r2_hidden(esk7_2(X1,k1_xboole_0),X2)|~r2_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_64, plain, (r2_hidden(X1,k1_tarski(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_55])])).
fof(c_0_65, plain, ![X11, X12]:(((r1_tarski(X11,X12)|~r2_xboole_0(X11,X12))&(X11!=X12|~r2_xboole_0(X11,X12)))&(~r1_tarski(X11,X12)|X11=X12|r2_xboole_0(X11,X12))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_xboole_0])])])])).
cnf(c_0_66, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_56])).
cnf(c_0_67, plain, (r1_tarski(k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4)),k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_31, c_0_37])).
cnf(c_0_68, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(esk3_0,esk4_0),k2_zfmisc_1(X1,esk2_0))=k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_39])).
fof(c_0_69, plain, ![X23, X24]:((k2_zfmisc_1(X23,X24)!=k1_xboole_0|(X23=k1_xboole_0|X24=k1_xboole_0))&((X23!=k1_xboole_0|k2_zfmisc_1(X23,X24)=k1_xboole_0)&(X24!=k1_xboole_0|k2_zfmisc_1(X23,X24)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
cnf(c_0_70, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(X1,esk2_0))=k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_58]), c_0_57]), c_0_32]), c_0_33]), c_0_39]), c_0_38])).
fof(c_0_71, plain, ![X47]:k3_xboole_0(X47,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_72, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_59, c_0_21])).
cnf(c_0_73, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_60])).
cnf(c_0_74, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk10_1(k3_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_20, c_0_61])).
cnf(c_0_75, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_76, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_77, plain, (~r2_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_78, plain, (X1=X2|r2_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_79, plain, (esk5_2(k1_tarski(X1),X2)=X1|r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_66, c_0_21])).
cnf(c_0_80, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0),k2_zfmisc_1(X1,esk4_0))), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_81, plain, (X1=k1_xboole_0|X2=k1_xboole_0|k2_zfmisc_1(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_82, negated_conjecture, (k2_zfmisc_1(esk1_0,k3_xboole_0(esk2_0,X1))=k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))), inference(spm,[status(thm)],[c_0_70, c_0_32])).
cnf(c_0_83, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_84, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_85, plain, (k3_xboole_0(X1,X2)=X1|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_28, c_0_72])).
cnf(c_0_86, plain, (v1_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X3)))|~r2_hidden(esk10_1(k3_xboole_0(X1,k4_xboole_0(X2,X3))),X3)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_87, plain, (v1_xboole_0(k3_xboole_0(X1,X2))|r2_hidden(esk10_1(k3_xboole_0(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_75, c_0_61])).
cnf(c_0_88, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_76])).
cnf(c_0_89, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_90, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_24, c_0_79])).
cnf(c_0_91, negated_conjecture, (r1_tarski(k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0),k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk4_0))), inference(spm,[status(thm)],[c_0_80, c_0_50])).
cnf(c_0_92, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|k2_zfmisc_1(esk3_0,k3_xboole_0(X1,esk4_0))!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])).
cnf(c_0_93, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_94, plain, (v1_xboole_0(k3_xboole_0(X1,k4_xboole_0(X2,X1)))), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_95, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_96, plain, (r2_hidden(esk5_2(k4_xboole_0(X1,X2),X3),X1)|r1_tarski(k4_xboole_0(X1,X2),X3)), inference(spm,[status(thm)],[c_0_88, c_0_21])).
cnf(c_0_97, plain, (k1_tarski(X1)=k1_xboole_0|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_98, plain, (v1_xboole_0(k1_xboole_0)|~r2_hidden(esk10_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_53, c_0_61])).
cnf(c_0_99, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0)=k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_91]), c_0_38]), c_0_32]), c_0_33]), c_0_39]), c_0_57])).
cnf(c_0_100, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_101, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk2_0),k2_zfmisc_1(esk3_0,esk4_0))=k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0)), inference(spm,[status(thm)],[c_0_57, c_0_39])).
cnf(c_0_102, negated_conjecture, (k3_xboole_0(esk2_0,X1)=k1_xboole_0|k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,X1))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_92, c_0_32])).
cnf(c_0_103, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_104, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_95])).
cnf(c_0_105, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_24, c_0_96])).
cnf(c_0_106, plain, (~v1_xboole_0(k1_tarski(X1))), inference(spm,[status(thm)],[c_0_59, c_0_64])).
cnf(c_0_107, plain, (k1_tarski(esk5_2(k1_xboole_0,X1))=k1_xboole_0|r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_97, c_0_21])).
cnf(c_0_108, plain, (v1_xboole_0(k1_xboole_0)), inference(spm,[status(thm)],[c_0_98, c_0_64])).
cnf(c_0_109, negated_conjecture, (k3_xboole_0(X1,esk1_0)=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_99]), c_0_100])).
cnf(c_0_110, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_111, plain, (k3_xboole_0(k2_zfmisc_1(X1,X2),k2_zfmisc_1(X3,X4))=k3_xboole_0(k2_zfmisc_1(X3,X2),k2_zfmisc_1(X1,X4))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_32]), c_0_33])).
cnf(c_0_112, negated_conjecture, (k3_xboole_0(k2_zfmisc_1(X1,esk4_0),k2_zfmisc_1(esk3_0,esk2_0))=k2_zfmisc_1(k3_xboole_0(X1,esk1_0),esk2_0)), inference(rw,[status(thm)],[c_0_101, c_0_37])).
cnf(c_0_113, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_114, negated_conjecture, (k3_xboole_0(esk2_0,k4_xboole_0(X1,esk4_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])])).
cnf(c_0_115, plain, (k3_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_105]), c_0_32])).
cnf(c_0_116, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_117, plain, (r1_tarski(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_108])])).
cnf(c_0_118, negated_conjecture, (k3_xboole_0(X1,esk1_0)=k1_xboole_0|k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_109, c_0_32])).
cnf(c_0_119, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_110])).
cnf(c_0_120, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(X1,esk3_0),esk4_0)=k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_111]), c_0_112]), c_0_99])).
cnf(c_0_121, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_113])).
cnf(c_0_122, negated_conjecture, (k4_xboole_0(esk2_0,esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_123, plain, (~r2_hidden(X1,k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_53])).
cnf(c_0_124, negated_conjecture, (k3_xboole_0(esk1_0,k4_xboole_0(X1,esk3_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_103]), c_0_119])]), c_0_32])).
cnf(c_0_125, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk1_0,X1),esk2_0)=k2_zfmisc_1(k3_xboole_0(esk3_0,X1),esk4_0)), inference(spm,[status(thm)],[c_0_120, c_0_32])).
cnf(c_0_126, negated_conjecture, (r2_hidden(X1,esk4_0)|~r2_hidden(X1,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_123])).
cnf(c_0_127, negated_conjecture, (k4_xboole_0(esk1_0,esk3_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_124, c_0_115])).
cnf(c_0_128, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,k4_xboole_0(X1,esk1_0)),esk4_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_103]), c_0_119])).
cnf(c_0_129, negated_conjecture, (r2_hidden(esk5_2(esk2_0,X1),esk4_0)|r1_tarski(esk2_0,X1)), inference(spm,[status(thm)],[c_0_126, c_0_21])).
cnf(c_0_130, negated_conjecture, (r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_127]), c_0_123])).
cnf(c_0_131, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk3_0,esk1_0),esk4_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_128, c_0_115])).
cnf(c_0_132, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_84, c_0_32])).
cnf(c_0_133, negated_conjecture, (r1_tarski(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_24, c_0_129])).
cnf(c_0_134, negated_conjecture, (k2_zfmisc_1(esk3_0,k3_xboole_0(esk4_0,k4_xboole_0(X1,esk2_0)))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_103]), c_0_104]), c_0_32])).
cnf(c_0_135, negated_conjecture, (r2_hidden(esk5_2(esk1_0,X1),esk3_0)|r1_tarski(esk1_0,X1)), inference(spm,[status(thm)],[c_0_130, c_0_21])).
cnf(c_0_136, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk3_0,esk1_0),k3_xboole_0(esk4_0,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_131]), c_0_132])).
cnf(c_0_137, negated_conjecture, (k3_xboole_0(esk4_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_133]), c_0_32])).
cnf(c_0_138, negated_conjecture, (k2_zfmisc_1(esk3_0,k4_xboole_0(esk4_0,esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_134, c_0_115])).
cnf(c_0_139, negated_conjecture, (r1_tarski(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_24, c_0_135])).
cnf(c_0_140, negated_conjecture, (k2_zfmisc_1(k4_xboole_0(esk3_0,esk1_0),esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_136, c_0_137])).
cnf(c_0_141, negated_conjecture, (k2_zfmisc_1(k3_xboole_0(esk3_0,X1),k4_xboole_0(esk4_0,esk2_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_138]), c_0_132])).
cnf(c_0_142, negated_conjecture, (k3_xboole_0(esk3_0,esk1_0)=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_139]), c_0_32])).
cnf(c_0_143, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_140]), c_0_100])).
cnf(c_0_144, negated_conjecture, (k2_zfmisc_1(esk1_0,k4_xboole_0(esk4_0,esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_145, plain, (~r2_hidden(esk7_2(X1,X2),X1)|~r2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_146, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_143]), c_0_123])).
cnf(c_0_147, negated_conjecture, (k4_xboole_0(esk4_0,esk2_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_144]), c_0_83])).
cnf(c_0_148, negated_conjecture, (~r2_hidden(esk7_2(esk1_0,X1),esk3_0)|~r2_xboole_0(esk1_0,X1)), inference(spm,[status(thm)],[c_0_145, c_0_146])).
cnf(c_0_149, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk4_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_147]), c_0_123])).
cnf(c_0_150, negated_conjecture, (~r2_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_148, c_0_54])).
cnf(c_0_151, negated_conjecture, (~r2_hidden(esk7_2(esk2_0,X1),esk4_0)|~r2_xboole_0(esk2_0,X1)), inference(spm,[status(thm)],[c_0_145, c_0_149])).
cnf(c_0_152, negated_conjecture, (esk1_0!=esk3_0|esk2_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_153, negated_conjecture, (esk1_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_78]), c_0_139])])).
cnf(c_0_154, negated_conjecture, (~r2_xboole_0(esk2_0,esk4_0)), inference(spm,[status(thm)],[c_0_151, c_0_54])).
cnf(c_0_155, negated_conjecture, (esk2_0!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_153])])).
cnf(c_0_156, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_78]), c_0_133])]), c_0_155]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 157
# Proof object clause steps            : 124
# Proof object formula steps           : 33
# Proof object conjectures             : 56
# Proof object clause conjectures      : 53
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 29
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 83
# Proof object simplifying inferences  : 59
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 44
# Processed clauses                    : 3507
# ...of these trivial                  : 298
# ...subsumed                          : 2306
# ...remaining for further processing  : 903
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 30
# Backward-rewritten                   : 327
# Generated clauses                    : 32918
# ...of the previous two non-trivial   : 18788
# Contextual simplify-reflections      : 1
# Paramodulations                      : 32894
# Factorizations                       : 8
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 533
#    Positive orientable unit clauses  : 164
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 20
#    Non-unit-clauses                  : 346
# Current number of unprocessed clauses: 14952
# ...number of literals in the above   : 31836
# Current number of archived formulas  : 0
# Current number of archived clauses   : 357
# Clause-clause subsumption calls (NU) : 20956
# Rec. Clause-clause subsumption calls : 18253
# Non-unit clause-clause subsumptions  : 845
# Unit Clause-clause subsumption calls : 3702
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 396
# BW rewrite match successes           : 134
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 328153

# -------------------------------------------------
# User time                : 0.179 s
# System time              : 0.007 s
# Total time               : 0.185 s
# Maximum resident set size: 28540 pages
