# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d11_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(v1_relat_1(X3)=>(X3=k7_relat_1(X1,X2)<=>![X4, X5]:(r2_hidden(k4_tarski(X4,X5),X3)<=>(r2_hidden(X4,X2)&r2_hidden(k4_tarski(X4,X5),X1)))))), file('data/MPTP/proofs/e90cdc2d782a494ca6a12302ec3c0c50.E_input', d11_relat_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/e90cdc2d782a494ca6a12302ec3c0c50.E_input', d4_relat_1)).
fof(t73_funct_1, conjecture, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X1,X2))=>r2_hidden(k1_funct_1(X3,X1),k2_relat_1(k7_relat_1(X3,X2))))), file('data/MPTP/proofs/e90cdc2d782a494ca6a12302ec3c0c50.E_input', t73_funct_1)).
fof(t12_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X1),k2_relat_1(X2)))), file('data/MPTP/proofs/e90cdc2d782a494ca6a12302ec3c0c50.E_input', t12_funct_1)).
fof(t72_funct_1, axiom, ![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,X2)=>k1_funct_1(k7_relat_1(X3,X2),X1)=k1_funct_1(X3,X1))), file('data/MPTP/proofs/e90cdc2d782a494ca6a12302ec3c0c50.E_input', t72_funct_1)).
fof(fc8_funct_1, axiom, ![X1, X2]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k7_relat_1(X1,X2))&v1_funct_1(k7_relat_1(X1,X2)))), file('data/MPTP/proofs/e90cdc2d782a494ca6a12302ec3c0c50.E_input', fc8_funct_1)).
fof(c_0_6, plain, ![X17, X18, X19, X20, X21, X22, X23]:((((r2_hidden(X20,X18)|~r2_hidden(k4_tarski(X20,X21),X19)|X19!=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(X20,X21),X17)|~r2_hidden(k4_tarski(X20,X21),X19)|X19!=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17)))&(~r2_hidden(X22,X18)|~r2_hidden(k4_tarski(X22,X23),X17)|r2_hidden(k4_tarski(X22,X23),X19)|X19!=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17)))&((~r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X19)|(~r2_hidden(esk4_3(X17,X18,X19),X18)|~r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X17))|X19=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17))&((r2_hidden(esk4_3(X17,X18,X19),X18)|r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X19)|X19=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17))&(r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X17)|r2_hidden(k4_tarski(esk4_3(X17,X18,X19),esk5_3(X17,X18,X19)),X19)|X19=k7_relat_1(X17,X18)|~v1_relat_1(X19)|~v1_relat_1(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_relat_1])])])])])])).
fof(c_0_7, plain, ![X6, X7, X8, X10, X11, X12, X13, X15]:(((~r2_hidden(X8,X7)|r2_hidden(k4_tarski(X8,esk1_3(X6,X7,X8)),X6)|X7!=k1_relat_1(X6))&(~r2_hidden(k4_tarski(X10,X11),X6)|r2_hidden(X10,X7)|X7!=k1_relat_1(X6)))&((~r2_hidden(esk2_2(X12,X13),X13)|~r2_hidden(k4_tarski(esk2_2(X12,X13),X15),X12)|X13=k1_relat_1(X12))&(r2_hidden(esk2_2(X12,X13),X13)|r2_hidden(k4_tarski(esk2_2(X12,X13),esk3_2(X12,X13)),X12)|X13=k1_relat_1(X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_8, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X4!=k7_relat_1(X5,X2)|~v1_relat_1(X4)|~v1_relat_1(X5)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (r2_hidden(k4_tarski(X1,esk1_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, negated_conjecture, ~(![X1, X2, X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>((r2_hidden(X1,k1_relat_1(X3))&r2_hidden(X1,X2))=>r2_hidden(k1_funct_1(X3,X1),k2_relat_1(k7_relat_1(X3,X2)))))), inference(assume_negation,[status(cth)],[t73_funct_1])).
fof(c_0_11, plain, ![X29, X30]:(~v1_relat_1(X30)|~v1_funct_1(X30)|(~r2_hidden(X29,k1_relat_1(X30))|r2_hidden(k1_funct_1(X30,X29),k2_relat_1(X30)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_1])])).
fof(c_0_12, plain, ![X31, X32, X33]:(~v1_relat_1(X33)|~v1_funct_1(X33)|(~r2_hidden(X31,X32)|k1_funct_1(k7_relat_1(X33,X32),X31)=k1_funct_1(X33,X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t72_funct_1])])).
cnf(c_0_13, plain, (r2_hidden(X1,X2)|~v1_relat_1(k7_relat_1(X3,X2))|~v1_relat_1(X3)|~r2_hidden(k4_tarski(X1,X4),k7_relat_1(X3,X2))), inference(er,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r2_hidden(k4_tarski(X1,esk1_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_9])).
fof(c_0_15, plain, ![X34, X35]:((v1_relat_1(k7_relat_1(X34,X35))|(~v1_relat_1(X34)|~v1_funct_1(X34)))&(v1_funct_1(k7_relat_1(X34,X35))|(~v1_relat_1(X34)|~v1_funct_1(X34)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_funct_1])])])).
cnf(c_0_16, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_17, plain, (r2_hidden(k4_tarski(X1,X3),X5)|~r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),X4)|X5!=k7_relat_1(X4,X2)|~v1_relat_1(X5)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_18, negated_conjecture, ((v1_relat_1(esk8_0)&v1_funct_1(esk8_0))&((r2_hidden(esk6_0,k1_relat_1(esk8_0))&r2_hidden(esk6_0,esk7_0))&~r2_hidden(k1_funct_1(esk8_0,esk6_0),k2_relat_1(k7_relat_1(esk8_0,esk7_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_19, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_20, plain, (k1_funct_1(k7_relat_1(X1,X3),X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_21, plain, (r2_hidden(X1,X2)|~v1_relat_1(k7_relat_1(X3,X2))|~v1_relat_1(X3)|~r2_hidden(X1,k1_relat_1(k7_relat_1(X3,X2)))), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_22, plain, (v1_relat_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (v1_funct_1(k7_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_24, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(X1,X2),k7_relat_1(X3,X4))|~v1_relat_1(k7_relat_1(X3,X4))|~v1_relat_1(X3)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X1,X4)), inference(er,[status(thm)],[c_0_17])).
cnf(c_0_26, negated_conjecture, (~r2_hidden(k1_funct_1(esk8_0,esk6_0),k2_relat_1(k7_relat_1(esk8_0,esk7_0)))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(k7_relat_1(X1,X3)))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(k7_relat_1(X1,X3)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23])).
cnf(c_0_28, negated_conjecture, (v1_funct_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_29, negated_conjecture, (v1_relat_1(esk8_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_30, plain, (r2_hidden(X1,k1_relat_1(k7_relat_1(X2,X3)))|~v1_relat_1(k7_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(k4_tarski(X1,X4),X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_31, negated_conjecture, (~r2_hidden(esk6_0,k1_relat_1(k7_relat_1(esk8_0,esk7_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28]), c_0_29])])).
cnf(c_0_32, plain, (r2_hidden(X1,k1_relat_1(k7_relat_1(X2,X3)))|~v1_relat_1(k7_relat_1(X2,X3))|~v1_relat_1(X2)|~r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_30, c_0_14])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk6_0,k1_relat_1(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_34, negated_conjecture, (r2_hidden(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_35, negated_conjecture, (~v1_relat_1(k7_relat_1(esk8_0,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_32]), c_0_29]), c_0_33]), c_0_34])])).
cnf(c_0_36, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_22]), c_0_28]), c_0_29])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 24
# Proof object formula steps           : 13
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 53
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 53
# ...of the previous two non-trivial   : 43
# Contextual simplify-reflections      : 3
# Paramodulations                      : 46
# Factorizations                       : 2
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 22
# Current number of unprocessed clauses: 28
# ...number of literals in the above   : 168
# Current number of archived formulas  : 0
# Current number of archived clauses   : 19
# Clause-clause subsumption calls (NU) : 286
# Rec. Clause-clause subsumption calls : 66
# Non-unit clause-clause subsumptions  : 3
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2873

# -------------------------------------------------
# User time                : 0.008 s
# System time              : 0.003 s
# Total time               : 0.010 s
# Maximum resident set size: 28760 pages
