# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', d2_funct_2)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', d5_funct_1)).
fof(t121_funct_2, conjecture, ![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t121_funct_2)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', d1_xboole_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t2_subset)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', existence_m1_subset_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t5_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t3_subset)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', cc4_relset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', d10_xboole_0)).
fof(t21_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t21_relat_1)).
fof(t65_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k1_relat_1(X1)=k1_xboole_0<=>k2_relat_1(X1)=k1_xboole_0)), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t65_relat_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t113_zfmisc_1)).
fof(dt_k6_partfun1, axiom, ![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', dt_k6_partfun1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t7_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', d3_tarski)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', fc1_xboole_0)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t152_zfmisc_1)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t11_relset_1)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', t3_funct_2)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', cc1_funct_2)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/e963f9c115bf40b9872f2095947bc46e.E_input', cc5_funct_2)).
fof(c_0_22, plain, ![X6, X7, X8, X9, X11, X12, X13, X14, X15, X17]:(((((v1_relat_1(esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,X8)|X8!=k1_funct_2(X6,X7))&(v1_funct_1(esk1_4(X6,X7,X8,X9))|~r2_hidden(X9,X8)|X8!=k1_funct_2(X6,X7)))&((X9=esk1_4(X6,X7,X8,X9)|~r2_hidden(X9,X8)|X8!=k1_funct_2(X6,X7))&((k1_relat_1(esk1_4(X6,X7,X8,X9))=X6|~r2_hidden(X9,X8)|X8!=k1_funct_2(X6,X7))&(r1_tarski(k2_relat_1(esk1_4(X6,X7,X8,X9)),X7)|~r2_hidden(X9,X8)|X8!=k1_funct_2(X6,X7)))))&(~v1_relat_1(X12)|~v1_funct_1(X12)|(X11!=X12|(k1_relat_1(X12)!=X6|~r1_tarski(k2_relat_1(X12),X7)))|r2_hidden(X11,X8)|X8!=k1_funct_2(X6,X7)))&((~r2_hidden(esk2_3(X13,X14,X15),X15)|(~v1_relat_1(X17)|~v1_funct_1(X17)|(esk2_3(X13,X14,X15)!=X17|(k1_relat_1(X17)!=X13|~r1_tarski(k2_relat_1(X17),X14))))|X15=k1_funct_2(X13,X14))&(((v1_relat_1(esk3_3(X13,X14,X15))|r2_hidden(esk2_3(X13,X14,X15),X15)|X15=k1_funct_2(X13,X14))&(v1_funct_1(esk3_3(X13,X14,X15))|r2_hidden(esk2_3(X13,X14,X15),X15)|X15=k1_funct_2(X13,X14)))&((esk2_3(X13,X14,X15)=esk3_3(X13,X14,X15)|r2_hidden(esk2_3(X13,X14,X15),X15)|X15=k1_funct_2(X13,X14))&((k1_relat_1(esk3_3(X13,X14,X15))=X13|r2_hidden(esk2_3(X13,X14,X15),X15)|X15=k1_funct_2(X13,X14))&(r1_tarski(k2_relat_1(esk3_3(X13,X14,X15)),X14)|r2_hidden(esk2_3(X13,X14,X15),X15)|X15=k1_funct_2(X13,X14))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
fof(c_0_23, plain, ![X19, X20, X21, X23, X24, X25, X27]:((((r2_hidden(esk4_3(X19,X20,X21),k1_relat_1(X19))|~r2_hidden(X21,X20)|X20!=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))&(X21=k1_funct_1(X19,esk4_3(X19,X20,X21))|~r2_hidden(X21,X20)|X20!=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19))))&(~r2_hidden(X24,k1_relat_1(X19))|X23!=k1_funct_1(X19,X24)|r2_hidden(X23,X20)|X20!=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19))))&((~r2_hidden(esk5_2(X19,X25),X25)|(~r2_hidden(X27,k1_relat_1(X19))|esk5_2(X19,X25)!=k1_funct_1(X19,X27))|X25=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))&((r2_hidden(esk6_2(X19,X25),k1_relat_1(X19))|r2_hidden(esk5_2(X19,X25),X25)|X25=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))&(esk5_2(X19,X25)=k1_funct_1(X19,esk6_2(X19,X25))|r2_hidden(esk5_2(X19,X25),X25)|X25=k2_relat_1(X19)|(~v1_relat_1(X19)|~v1_funct_1(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
cnf(c_0_24, plain, (k1_relat_1(esk1_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_25, plain, (X1=esk1_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_26, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t121_funct_2])).
cnf(c_0_27, plain, (v1_relat_1(esk1_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, ![X32, X33, X34]:((~v1_xboole_0(X32)|~r2_hidden(X33,X32))&(r2_hidden(esk10_1(X34),X34)|v1_xboole_0(X34))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_29, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, ![X36, X37]:(~m1_subset_1(X36,X37)|(v1_xboole_0(X37)|r2_hidden(X36,X37))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_31, plain, ![X38]:m1_subset_1(esk11_1(X38),X38), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_32, plain, (v1_funct_1(esk1_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_33, plain, ![X42, X43, X44]:(~r2_hidden(X42,X43)|(~m1_subset_1(X43,k1_zfmisc_1(X44))|~v1_xboole_0(X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_34, plain, ![X45, X46]:((~m1_subset_1(X45,k1_zfmisc_1(X46))|r1_tarski(X45,X46))&(~r1_tarski(X45,X46)|m1_subset_1(X45,k1_zfmisc_1(X46)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_35, plain, (k1_relat_1(esk1_4(X1,X2,k1_funct_2(X1,X2),X3))=X1|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_36, plain, (esk1_4(X1,X2,k1_funct_2(X1,X2),X3)=X3|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_25])).
fof(c_0_37, negated_conjecture, (r2_hidden(esk9_0,k1_funct_2(esk7_0,esk8_0))&(~v1_funct_1(esk9_0)|(~v1_funct_2(esk9_0,esk7_0,esk8_0)|~m1_subset_1(esk9_0,k1_zfmisc_1(k2_zfmisc_1(esk7_0,esk8_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
cnf(c_0_38, plain, (v1_relat_1(esk1_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_27])).
fof(c_0_39, plain, ![X47, X48, X49]:(~v1_xboole_0(X47)|(~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X48,X47)))|v1_xboole_0(X49))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
fof(c_0_40, plain, ![X50, X51]:(((r1_tarski(X50,X51)|X50!=X51)&(r1_tarski(X51,X50)|X50!=X51))&(~r1_tarski(X50,X51)|~r1_tarski(X51,X50)|X50=X51)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_41, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_29])])).
cnf(c_0_43, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44, plain, (m1_subset_1(esk11_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_45, plain, (v1_funct_1(esk1_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_46, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_47, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_48, plain, ![X52]:(~v1_relat_1(X52)|r1_tarski(X52,k2_zfmisc_1(k1_relat_1(X52),k2_relat_1(X52)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_relat_1])])).
fof(c_0_49, plain, ![X53]:((k1_relat_1(X53)!=k1_xboole_0|k2_relat_1(X53)=k1_xboole_0|~v1_relat_1(X53))&(k2_relat_1(X53)!=k1_xboole_0|k1_relat_1(X53)=k1_xboole_0|~v1_relat_1(X53))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_relat_1])])])).
cnf(c_0_50, plain, (k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk9_0,k1_funct_2(esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_52, plain, (v1_relat_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_38, c_0_36])).
fof(c_0_53, plain, ![X40, X41]:((k2_zfmisc_1(X40,X41)!=k1_xboole_0|(X40=k1_xboole_0|X41=k1_xboole_0))&((X40!=k1_xboole_0|k2_zfmisc_1(X40,X41)=k1_xboole_0)&(X41!=k1_xboole_0|k2_zfmisc_1(X40,X41)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_54, plain, ![X54]:(v1_partfun1(k6_partfun1(X54),X54)&m1_subset_1(k6_partfun1(X54),k1_zfmisc_1(k2_zfmisc_1(X54,X54)))), inference(variable_rename,[status(thm)],[dt_k6_partfun1])).
fof(c_0_55, plain, ![X55]:(X55=k1_xboole_0|r2_hidden(esk12_1(X55),X55)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_56, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_58, plain, (~v1_xboole_0(k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_59, plain, (v1_xboole_0(X1)|r2_hidden(esk11_1(X1),X1)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_60, plain, (v1_funct_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_45, c_0_36])).
cnf(c_0_61, plain, (~v1_xboole_0(X1)|~r1_tarski(X2,X1)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_62, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, plain, (k2_relat_1(X1)=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_64, negated_conjecture, (k1_relat_1(esk9_0)=esk7_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_65, negated_conjecture, (v1_relat_1(esk9_0)), inference(spm,[status(thm)],[c_0_52, c_0_51])).
cnf(c_0_66, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_67, plain, ![X57, X58, X59, X60, X61]:((~r1_tarski(X57,X58)|(~r2_hidden(X59,X57)|r2_hidden(X59,X58)))&((r2_hidden(esk13_2(X60,X61),X60)|r1_tarski(X60,X61))&(~r2_hidden(esk13_2(X60,X61),X61)|r1_tarski(X60,X61)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_68, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_69, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_70, plain, (X1=k1_xboole_0|r2_hidden(esk12_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_71, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~r1_tarski(X1,k2_zfmisc_1(X3,X2))), inference(spm,[status(thm)],[c_0_56, c_0_47])).
cnf(c_0_72, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_73, plain, (v1_xboole_0(k1_relat_1(X1))|~v1_xboole_0(k2_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_74, negated_conjecture, (v1_funct_1(esk9_0)), inference(spm,[status(thm)],[c_0_60, c_0_51])).
cnf(c_0_75, plain, (~v1_xboole_0(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_76, negated_conjecture, (k2_relat_1(esk9_0)=k1_xboole_0|esk7_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])])).
cnf(c_0_77, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_66])).
cnf(c_0_78, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_79, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_80, plain, (r1_tarski(k6_partfun1(X1),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_81, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_41, c_0_70])).
cnf(c_0_82, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_83, negated_conjecture, (v1_xboole_0(esk7_0)|~v1_xboole_0(k2_relat_1(esk9_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_64]), c_0_74]), c_0_65])])).
fof(c_0_84, plain, ![X63, X64]:~r2_hidden(X63,k2_zfmisc_1(X63,X64)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
fof(c_0_85, plain, ![X65, X66, X67]:(~v1_relat_1(X67)|(~r1_tarski(k1_relat_1(X67),X65)|~r1_tarski(k2_relat_1(X67),X66)|m1_subset_1(X67,k1_zfmisc_1(k2_zfmisc_1(X65,X66))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_86, negated_conjecture, (esk7_0!=k1_xboole_0|~r2_hidden(X1,esk9_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_64]), c_0_77]), c_0_78]), c_0_65])])).
cnf(c_0_87, plain, (r2_hidden(esk13_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_88, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,k6_partfun1(X2))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_89, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_90, negated_conjecture, (v1_xboole_0(esk7_0)|esk7_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_76]), c_0_78])])).
cnf(c_0_91, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_92, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_77]), c_0_78])])).
cnf(c_0_93, plain, (r2_hidden(esk4_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_94, negated_conjecture, (~v1_funct_1(esk9_0)|~v1_funct_2(esk9_0,esk7_0,esk8_0)|~m1_subset_1(esk9_0,k1_zfmisc_1(k2_zfmisc_1(esk7_0,esk8_0)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_95, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_96, plain, (r1_tarski(k2_relat_1(esk1_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_97, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_98, negated_conjecture, (r1_tarski(esk9_0,X1)|esk7_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_99, plain, (r1_tarski(k6_partfun1(X1),X2)|r2_hidden(esk13_2(k6_partfun1(X1),X2),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_88, c_0_87])).
cnf(c_0_100, negated_conjecture, (k2_zfmisc_1(X1,esk7_0)=k1_xboole_0|esk7_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_101, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_91, c_0_77])).
fof(c_0_102, plain, ![X71]:((v1_funct_1(X71)|(~v1_relat_1(X71)|~v1_funct_1(X71)))&((v1_funct_2(X71,k1_relat_1(X71),k2_relat_1(X71))|(~v1_relat_1(X71)|~v1_funct_1(X71)))&(m1_subset_1(X71,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X71),k2_relat_1(X71))))|(~v1_relat_1(X71)|~v1_funct_1(X71))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_103, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_104, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_92, c_0_47])).
cnf(c_0_105, plain, (r2_hidden(esk4_3(X1,k2_relat_1(X1),X2),k1_relat_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_93])).
cnf(c_0_106, plain, (k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_107, negated_conjecture, (~v1_funct_2(esk9_0,esk7_0,esk8_0)|~r1_tarski(k1_relat_1(esk9_0),esk7_0)|~r1_tarski(k2_relat_1(esk9_0),esk8_0)|~v1_funct_1(esk9_0)|~v1_relat_1(esk9_0)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_108, plain, (r1_tarski(k2_relat_1(esk1_4(X1,X2,k1_funct_2(X1,X2),X3)),X2)|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_96])).
fof(c_0_109, plain, ![X68, X69, X70]:((v1_funct_1(X70)|(~v1_funct_1(X70)|~v1_partfun1(X70,X68))|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69))))&(v1_funct_2(X70,X68,X69)|(~v1_funct_1(X70)|~v1_partfun1(X70,X68))|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_110, negated_conjecture, (X1=esk9_0|esk7_0!=k1_xboole_0|~r1_tarski(X1,esk9_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_111, negated_conjecture, (r1_tarski(k6_partfun1(esk7_0),X1)|esk7_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])).
fof(c_0_112, plain, ![X72, X73, X74]:((v1_funct_1(X74)|(~v1_funct_1(X74)|~v1_funct_2(X74,X72,X73))|~m1_subset_1(X74,k1_zfmisc_1(k2_zfmisc_1(X72,X73)))|v1_xboole_0(X73))&(v1_partfun1(X74,X72)|(~v1_funct_1(X74)|~v1_funct_2(X74,X72,X73))|~m1_subset_1(X74,k1_zfmisc_1(k2_zfmisc_1(X72,X73)))|v1_xboole_0(X73))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
cnf(c_0_113, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_114, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_115, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_103])).
cnf(c_0_116, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_41, c_0_87])).
cnf(c_0_117, plain, (v1_xboole_0(k6_partfun1(X1))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_56, c_0_69])).
cnf(c_0_118, negated_conjecture, (v1_xboole_0(esk7_0)|~r1_tarski(k2_relat_1(esk9_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_83, c_0_104])).
cnf(c_0_119, plain, (k2_relat_1(X1)!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_101])).
cnf(c_0_120, negated_conjecture, (~v1_funct_2(esk9_0,esk7_0,esk8_0)|~r1_tarski(k1_relat_1(esk9_0),esk7_0)|~r1_tarski(k2_relat_1(esk9_0),esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_74])]), c_0_65])])).
cnf(c_0_121, plain, (r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(spm,[status(thm)],[c_0_108, c_0_36])).
cnf(c_0_122, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_123, plain, (v1_partfun1(k6_partfun1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_124, negated_conjecture, (k6_partfun1(esk7_0)=esk9_0|esk7_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_125, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_126, negated_conjecture, (m1_subset_1(esk9_0,k1_zfmisc_1(k2_zfmisc_1(esk7_0,k2_relat_1(esk9_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_64]), c_0_74]), c_0_65])])).
cnf(c_0_127, negated_conjecture, (v1_funct_2(esk9_0,esk7_0,k2_relat_1(esk9_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_64]), c_0_74]), c_0_65])])).
cnf(c_0_128, plain, (m1_subset_1(k6_partfun1(k1_xboole_0),k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_69, c_0_115])).
cnf(c_0_129, plain, (r1_tarski(k6_partfun1(X1),X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_130, negated_conjecture, (r1_tarski(esk7_0,X1)|~r1_tarski(k2_relat_1(esk9_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_116, c_0_118])).
cnf(c_0_131, plain, (r1_tarski(k2_relat_1(X1),X2)|k2_relat_1(X1)!=k1_xboole_0|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_119, c_0_87])).
cnf(c_0_132, negated_conjecture, (~v1_funct_2(esk9_0,esk7_0,esk8_0)|~r1_tarski(k2_relat_1(esk9_0),esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_120, c_0_64]), c_0_72])])).
cnf(c_0_133, negated_conjecture, (r1_tarski(k2_relat_1(esk9_0),esk8_0)), inference(spm,[status(thm)],[c_0_121, c_0_51])).
cnf(c_0_134, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_122, c_0_47])).
cnf(c_0_135, negated_conjecture, (v1_partfun1(esk9_0,esk7_0)|esk7_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_123, c_0_124])).
cnf(c_0_136, negated_conjecture, (~v1_funct_2(esk9_0,esk7_0,esk8_0)|~r1_tarski(esk9_0,k2_zfmisc_1(esk7_0,esk8_0))|~v1_funct_1(esk9_0)), inference(spm,[status(thm)],[c_0_94, c_0_47])).
cnf(c_0_137, negated_conjecture, (v1_partfun1(esk9_0,esk7_0)|v1_xboole_0(k2_relat_1(esk9_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_127]), c_0_74])])).
cnf(c_0_138, plain, (~r2_hidden(X1,k6_partfun1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_128]), c_0_78])])).
cnf(c_0_139, plain, (X1=k6_partfun1(X2)|~v1_xboole_0(X2)|~r1_tarski(X1,k6_partfun1(X2))), inference(spm,[status(thm)],[c_0_97, c_0_129])).
cnf(c_0_140, negated_conjecture, (r1_tarski(esk7_0,X1)|k2_relat_1(esk9_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131]), c_0_74]), c_0_65])])).
cnf(c_0_141, negated_conjecture, (~v1_funct_2(esk9_0,esk7_0,esk8_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_142, negated_conjecture, (v1_funct_2(esk9_0,esk7_0,X1)|esk7_0!=k1_xboole_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_74])]), c_0_98])).
cnf(c_0_143, negated_conjecture, (~v1_funct_2(esk9_0,esk7_0,esk8_0)|~r1_tarski(esk9_0,k2_zfmisc_1(esk7_0,esk8_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_136, c_0_74])])).
cnf(c_0_144, negated_conjecture, (v1_xboole_0(k2_relat_1(esk9_0))|v1_funct_2(esk9_0,esk7_0,X1)|~r1_tarski(esk9_0,k2_zfmisc_1(esk7_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_137]), c_0_74])])).
cnf(c_0_145, plain, (k6_partfun1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_138, c_0_70])).
cnf(c_0_146, negated_conjecture, (k6_partfun1(X1)=esk7_0|k2_relat_1(esk9_0)!=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_139, c_0_140])).
cnf(c_0_147, negated_conjecture, (esk7_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_148, negated_conjecture, (v1_xboole_0(k2_relat_1(esk9_0))|~r1_tarski(esk9_0,k2_zfmisc_1(esk7_0,esk8_0))), inference(spm,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_149, negated_conjecture, (k2_relat_1(esk9_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_78])]), c_0_147])).
cnf(c_0_150, negated_conjecture, (~r1_tarski(esk9_0,k2_zfmisc_1(esk7_0,esk8_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_148]), c_0_149])).
cnf(c_0_151, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_68, c_0_95])).
cnf(c_0_152, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_64]), c_0_72]), c_0_133]), c_0_65])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 153
# Proof object clause steps            : 109
# Proof object formula steps           : 44
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 60
# Proof object simplifying inferences  : 66
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 53
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 50
# Processed clauses                    : 5145
# ...of these trivial                  : 12
# ...subsumed                          : 4065
# ...remaining for further processing  : 1068
# Other redundant clauses eliminated   : 66
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 99
# Backward-rewritten                   : 31
# Generated clauses                    : 21242
# ...of the previous two non-trivial   : 18707
# Contextual simplify-reflections      : 68
# Paramodulations                      : 21176
# Factorizations                       : 0
# Equation resolutions                 : 66
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 872
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 831
# Current number of unprocessed clauses: 13298
# ...number of literals in the above   : 57525
# Current number of archived formulas  : 0
# Current number of archived clauses   : 182
# Clause-clause subsumption calls (NU) : 128674
# Rec. Clause-clause subsumption calls : 76296
# Non-unit clause-clause subsumptions  : 2240
# Unit Clause-clause subsumption calls : 1292
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 281678

# -------------------------------------------------
# User time                : 0.238 s
# System time              : 0.003 s
# Total time               : 0.242 s
# Maximum resident set size: 28688 pages
