# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t72_xboole_1, conjecture, ![X1, X2, X3, X4]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)&(r1_xboole_0(X3,X1)&r1_xboole_0(X4,X2)))=>X3=X2), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t72_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', commutativity_k2_xboole_0)).
fof(t40_xboole_1, axiom, ![X1, X2]:k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t40_xboole_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t39_xboole_1)).
fof(t4_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~(r2_hidden(X3,k3_xboole_0(X1,X2)))))&~((?[X3]:r2_hidden(X3,k3_xboole_0(X1,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t4_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t48_xboole_1)).
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t46_xboole_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t43_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t7_xboole_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t3_xboole_0)).
fof(t4_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t4_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t1_boole)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t66_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', commutativity_k3_xboole_0)).
fof(t51_xboole_1, axiom, ![X1, X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t51_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', t12_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/e987829d331c4697a8622f9579ad91f3.E_input', l32_xboole_1)).
fof(c_0_17, negated_conjecture, ~(![X1, X2, X3, X4]:((k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)&(r1_xboole_0(X3,X1)&r1_xboole_0(X4,X2)))=>X3=X2)), inference(assume_negation,[status(cth)],[t72_xboole_1])).
fof(c_0_18, negated_conjecture, ((k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk3_0,esk4_0)&(r1_xboole_0(esk3_0,esk1_0)&r1_xboole_0(esk4_0,esk2_0)))&esk3_0!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_17])])])).
fof(c_0_19, plain, ![X19, X20]:k2_xboole_0(X19,X20)=k2_xboole_0(X20,X19), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
fof(c_0_20, plain, ![X14, X15]:k4_xboole_0(k2_xboole_0(X14,X15),X15)=k4_xboole_0(X14,X15), inference(variable_rename,[status(thm)],[t40_xboole_1])).
cnf(c_0_21, negated_conjecture, (k2_xboole_0(esk1_0,esk2_0)=k2_xboole_0(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_22, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_23, plain, ![X10, X11]:k2_xboole_0(X10,k4_xboole_0(X11,X10))=k2_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_24, plain, (k4_xboole_0(k2_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, negated_conjecture, (k2_xboole_0(esk2_0,esk1_0)=k2_xboole_0(esk3_0,esk4_0)), inference(rw,[status(thm)],[c_0_21, c_0_22])).
fof(c_0_26, plain, ![X23, X24, X26, X27, X28]:((r1_xboole_0(X23,X24)|r2_hidden(esk5_2(X23,X24),k3_xboole_0(X23,X24)))&(~r2_hidden(X28,k3_xboole_0(X26,X27))|~r1_xboole_0(X26,X27))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xboole_0])])])])])])).
fof(c_0_27, plain, ![X12, X13]:k4_xboole_0(X12,k4_xboole_0(X12,X13))=k3_xboole_0(X12,X13), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_28, plain, ![X35, X36]:k4_xboole_0(X35,k2_xboole_0(X35,X36))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
cnf(c_0_29, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, negated_conjecture, (k4_xboole_0(k2_xboole_0(esk3_0,esk4_0),esk1_0)=k4_xboole_0(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_31, plain, ![X42, X43, X44]:(~r1_tarski(X42,k2_xboole_0(X43,X44))|r1_tarski(k4_xboole_0(X42,X43),X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
fof(c_0_32, plain, ![X17, X18]:r1_tarski(X17,k2_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_33, plain, (~r2_hidden(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_35, plain, ![X29, X30, X32, X33, X34]:(((r2_hidden(esk6_2(X29,X30),X29)|r1_xboole_0(X29,X30))&(r2_hidden(esk6_2(X29,X30),X30)|r1_xboole_0(X29,X30)))&(~r2_hidden(X34,X32)|~r2_hidden(X34,X33)|~r1_xboole_0(X32,X33))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_36, plain, ![X45, X46, X47]:k2_xboole_0(k2_xboole_0(X45,X46),X47)=k2_xboole_0(X45,k2_xboole_0(X46,X47)), inference(variable_rename,[status(thm)],[t4_xboole_1])).
cnf(c_0_37, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (k2_xboole_0(esk1_0,k2_xboole_0(esk3_0,esk4_0))=k2_xboole_0(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_29]), c_0_22]), c_0_25])).
fof(c_0_39, plain, ![X37]:k2_xboole_0(X37,k1_xboole_0)=X37, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_40, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_42, plain, ![X9]:((r1_xboole_0(X9,X9)|X9!=k1_xboole_0)&(X9=k1_xboole_0|~r1_xboole_0(X9,X9))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_43, plain, (~r1_xboole_0(X2,X3)|~r2_hidden(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_44, plain, (r2_hidden(esk6_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_45, plain, ![X21, X22]:k3_xboole_0(X21,X22)=k3_xboole_0(X22,X21), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_46, plain, ![X38, X39]:k2_xboole_0(k3_xboole_0(X38,X39),k4_xboole_0(X38,X39))=X38, inference(variable_rename,[status(thm)],[t51_xboole_1])).
cnf(c_0_47, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_48, negated_conjecture, (k4_xboole_0(esk1_0,k2_xboole_0(esk3_0,esk4_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_49, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_50, plain, ![X40, X41]:(~r1_tarski(X40,X41)|k2_xboole_0(X40,X41)=X41), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
fof(c_0_51, plain, ![X48, X49]:((k4_xboole_0(X48,X49)!=k1_xboole_0|r1_tarski(X48,X49))&(~r1_tarski(X48,X49)|k4_xboole_0(X48,X49)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_52, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_40, c_0_22])).
cnf(c_0_53, negated_conjecture, (r1_tarski(esk2_0,k2_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_41, c_0_25])).
cnf(c_0_54, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_55, plain, (r1_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X2,X3)))|~r1_xboole_0(X2,X3)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_56, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57, plain, (k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_58, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X1,k2_xboole_0(X2,X3)))), inference(spm,[status(thm)],[c_0_41, c_0_47])).
cnf(c_0_59, negated_conjecture, (k2_xboole_0(esk3_0,k2_xboole_0(esk1_0,esk4_0))=k2_xboole_0(esk3_0,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_48]), c_0_49]), c_0_47]), c_0_22])).
cnf(c_0_60, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,esk4_0),esk3_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_63, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_64, negated_conjecture, (r1_xboole_0(esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_65, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_34]), c_0_34])).
cnf(c_0_66, plain, (k2_xboole_0(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k4_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[c_0_57, c_0_34])).
cnf(c_0_67, negated_conjecture, (r1_tarski(k4_xboole_0(X1,esk1_0),esk2_0)|~r1_tarski(X1,k2_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_52, c_0_25])).
cnf(c_0_68, negated_conjecture, (r1_tarski(k2_xboole_0(esk3_0,esk1_0),k2_xboole_0(esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_69, negated_conjecture, (r1_xboole_0(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_70, plain, (k2_xboole_0(X1,X2)=X2|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_71, negated_conjecture, (k2_xboole_0(esk3_0,k4_xboole_0(esk2_0,esk4_0))=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_62]), c_0_22])).
cnf(c_0_72, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk4_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
cnf(c_0_73, plain, (k2_xboole_0(X1,k4_xboole_0(X1,X2))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_22]), c_0_29]), c_0_22])).
cnf(c_0_74, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_75, negated_conjecture, (r1_tarski(k4_xboole_0(esk3_0,esk1_0),esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_24])).
cnf(c_0_76, negated_conjecture, (k4_xboole_0(esk3_0,k4_xboole_0(esk3_0,esk1_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_69])).
cnf(c_0_77, negated_conjecture, (k4_xboole_0(esk2_0,esk4_0)=esk3_0|k4_xboole_0(esk3_0,k4_xboole_0(esk2_0,esk4_0))!=k1_xboole_0), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_78, negated_conjecture, (k4_xboole_0(esk2_0,esk4_0)=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_72]), c_0_49]), c_0_22]), c_0_73])).
cnf(c_0_79, negated_conjecture, (esk3_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_80, negated_conjecture, (k4_xboole_0(k4_xboole_0(esk3_0,esk1_0),esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_81, negated_conjecture, (k4_xboole_0(esk3_0,esk1_0)=esk3_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_76]), c_0_49]), c_0_22]), c_0_73])).
cnf(c_0_82, negated_conjecture, (k4_xboole_0(esk3_0,esk2_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_78]), c_0_78]), c_0_79])).
cnf(c_0_83, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_81]), c_0_82]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 84
# Proof object clause steps            : 49
# Proof object formula steps           : 35
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 25
# Processed clauses                    : 29322
# ...of these trivial                  : 900
# ...subsumed                          : 26400
# ...remaining for further processing  : 2022
# Other redundant clauses eliminated   : 661
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 122
# Backward-rewritten                   : 863
# Generated clauses                    : 334718
# ...of the previous two non-trivial   : 289956
# Contextual simplify-reflections      : 1
# Paramodulations                      : 334057
# Factorizations                       : 0
# Equation resolutions                 : 661
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1011
#    Positive orientable unit clauses  : 249
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 58
#    Non-unit-clauses                  : 699
# Current number of unprocessed clauses: 258900
# ...number of literals in the above   : 519939
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1011
# Clause-clause subsumption calls (NU) : 587787
# Rec. Clause-clause subsumption calls : 368508
# Non-unit clause-clause subsumptions  : 11440
# Unit Clause-clause subsumption calls : 26559
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3147
# BW rewrite match successes           : 399
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4908906

# -------------------------------------------------
# User time                : 2.983 s
# System time              : 0.083 s
# Total time               : 3.066 s
# Maximum resident set size: 203192 pages
