# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t100_xboole_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t12_setfam_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t36_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t48_xboole_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t2_xboole_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', d5_xboole_0)).
fof(t45_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>X2=k2_xboole_0(X1,k4_xboole_0(X2,X1))), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t45_xboole_1)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t47_xboole_1)).
fof(dt_k6_subset_1, axiom, ![X1, X2]:m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1)), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', dt_k6_subset_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', redefinition_k6_subset_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t1_boole)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', redefinition_k7_subset_1)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', d3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', dt_l1_pre_topc)).
fof(t53_pre_topc, conjecture, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)=>k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))&((v2_pre_topc(X1)&k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=>v3_pre_topc(X2,X1))))), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t53_pre_topc)).
fof(t22_pre_topc, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2)), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t22_pre_topc)).
fof(d6_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v4_pre_topc(X2,X1)<=>v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)))), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', d6_pre_topc)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', commutativity_k2_tarski)).
fof(t52_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v4_pre_topc(X2,X1)=>k2_pre_topc(X1,X2)=X2)&((v2_pre_topc(X1)&k2_pre_topc(X1,X2)=X2)=>v4_pre_topc(X2,X1))))), file('data/MPTP/proofs/ea067965bc91434eaa3424ea85606dfd.E_input', t52_pre_topc)).
fof(c_0_20, plain, ![X5, X6]:k4_xboole_0(X5,X6)=k5_xboole_0(X5,k3_xboole_0(X5,X6)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_21, plain, ![X7, X8]:k1_setfam_1(k2_tarski(X7,X8))=k3_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_22, plain, ![X21, X22]:r1_tarski(k4_xboole_0(X21,X22),X21), inference(variable_rename,[status(thm)],[t36_xboole_1])).
cnf(c_0_23, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_25, plain, ![X23, X24]:k4_xboole_0(X23,k4_xboole_0(X23,X24))=k3_xboole_0(X23,X24), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_26, plain, ![X9, X10]:(~r2_hidden(X9,X10)|~r1_tarski(X10,X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_27, plain, ![X11]:r1_tarski(k1_xboole_0,X11), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_28, plain, ![X12, X13, X14, X15, X16, X17, X18, X19]:((((r2_hidden(X15,X12)|~r2_hidden(X15,X14)|X14!=k4_xboole_0(X12,X13))&(~r2_hidden(X15,X13)|~r2_hidden(X15,X14)|X14!=k4_xboole_0(X12,X13)))&(~r2_hidden(X16,X12)|r2_hidden(X16,X13)|r2_hidden(X16,X14)|X14!=k4_xboole_0(X12,X13)))&((~r2_hidden(esk1_3(X17,X18,X19),X19)|(~r2_hidden(esk1_3(X17,X18,X19),X17)|r2_hidden(esk1_3(X17,X18,X19),X18))|X19=k4_xboole_0(X17,X18))&((r2_hidden(esk1_3(X17,X18,X19),X17)|r2_hidden(esk1_3(X17,X18,X19),X19)|X19=k4_xboole_0(X17,X18))&(~r2_hidden(esk1_3(X17,X18,X19),X18)|r2_hidden(esk1_3(X17,X18,X19),X19)|X19=k4_xboole_0(X17,X18))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_29, plain, ![X25, X26]:(~r1_tarski(X25,X26)|X26=k2_xboole_0(X25,k4_xboole_0(X26,X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_xboole_1])])).
fof(c_0_30, plain, ![X34, X35]:k4_xboole_0(X34,k3_xboole_0(X34,X35))=k4_xboole_0(X34,X35), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_31, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_33, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (r2_hidden(esk1_3(X1,X2,X3),X1)|r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_37, plain, ![X30, X31]:m1_subset_1(k6_subset_1(X30,X31),k1_zfmisc_1(X30)), inference(variable_rename,[status(thm)],[dt_k6_subset_1])).
fof(c_0_38, plain, ![X32, X33]:k6_subset_1(X32,X33)=k4_xboole_0(X32,X33), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_39, plain, (X2=k2_xboole_0(X1,k4_xboole_0(X2,X1))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (r1_tarski(k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2))),X1)), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_42, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2))))))=k1_setfam_1(k2_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_24]), c_0_32]), c_0_32])).
cnf(c_0_43, plain, (r2_hidden(esk1_3(X1,X2,X3),X3)|X3=k4_xboole_0(X1,X2)|~r2_hidden(esk1_3(X1,X2,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_44, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45, plain, (X3=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))|r2_hidden(esk1_3(X1,X2,X3),X3)|r2_hidden(esk1_3(X1,X2,X3),X1)), inference(rw,[status(thm)],[c_0_36, c_0_32])).
cnf(c_0_46, plain, (m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48, plain, (X2=k2_xboole_0(X1,k5_xboole_0(X2,k1_setfam_1(k2_tarski(X2,X1))))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_39, c_0_32])).
cnf(c_0_49, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,k1_setfam_1(k2_tarski(X1,X2)))))=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_24]), c_0_32]), c_0_32])).
cnf(c_0_50, plain, (r1_tarski(k1_setfam_1(k2_tarski(X1,X2)),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_51, plain, (X3=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))|r2_hidden(esk1_3(X1,X2,X3),X3)|~r2_hidden(esk1_3(X1,X2,X3),X2)), inference(rw,[status(thm)],[c_0_43, c_0_32])).
cnf(c_0_52, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))=k1_xboole_0|r2_hidden(esk1_3(X1,X2,k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
fof(c_0_53, plain, ![X27]:k2_xboole_0(X27,k1_xboole_0)=X27, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_54, plain, ![X36, X37, X38]:(~m1_subset_1(X37,k1_zfmisc_1(X36))|k7_subset_1(X36,X37,X38)=k4_xboole_0(X37,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
cnf(c_0_55, plain, (m1_subset_1(k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2))),k1_zfmisc_1(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_47]), c_0_32])).
cnf(c_0_56, plain, (k2_xboole_0(k1_setfam_1(k2_tarski(X1,X2)),k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2))))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_50])])).
cnf(c_0_57, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X1)))=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52]), c_0_44])).
cnf(c_0_58, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_59, plain, ![X39]:(~l1_struct_0(X39)|k2_struct_0(X39)=u1_struct_0(X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
fof(c_0_60, plain, ![X40]:(~l1_pre_topc(X40)|l1_struct_0(X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_61, negated_conjecture, ~(![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>((v3_pre_topc(X2,X1)=>k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))&((v2_pre_topc(X1)&k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=>v3_pre_topc(X2,X1)))))), inference(assume_negation,[status(cth)],[t53_pre_topc])).
cnf(c_0_62, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_63, plain, (m1_subset_1(k1_setfam_1(k2_tarski(X1,X2)),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_55, c_0_42])).
cnf(c_0_64, plain, (k1_setfam_1(k2_tarski(X1,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])).
cnf(c_0_65, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_66, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_67, negated_conjecture, (l1_pre_topc(esk2_0)&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))&((((v2_pre_topc(esk2_0)|v3_pre_topc(esk3_0,esk2_0))&(k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)|v3_pre_topc(esk3_0,esk2_0)))&(~v3_pre_topc(esk3_0,esk2_0)|v3_pre_topc(esk3_0,esk2_0)))&(((v2_pre_topc(esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))&(k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)))&(~v3_pre_topc(esk3_0,esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_61])])])])).
cnf(c_0_68, plain, (k7_subset_1(X2,X1,X3)=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X3)))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_62, c_0_32])).
cnf(c_0_69, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
fof(c_0_70, plain, ![X43, X44]:(~l1_struct_0(X43)|(~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X43)))|k7_subset_1(u1_struct_0(X43),k2_struct_0(X43),k7_subset_1(u1_struct_0(X43),k2_struct_0(X43),X44))=X44)), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t22_pre_topc])])])).
cnf(c_0_71, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_72, negated_conjecture, (l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_73, plain, (k1_setfam_1(k2_tarski(X1,k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))))=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_42]), c_0_49])).
cnf(c_0_74, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,X2)))=k7_subset_1(X1,X1,X2)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, plain, (k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2))=X2|~l1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_76, negated_conjecture, (k2_struct_0(esk2_0)=u1_struct_0(esk2_0)), inference(spm,[status(thm)],[c_0_71, c_0_72])).
fof(c_0_77, plain, ![X45, X46]:((~v4_pre_topc(X46,X45)|v3_pre_topc(k7_subset_1(u1_struct_0(X45),k2_struct_0(X45),X46),X45)|~m1_subset_1(X46,k1_zfmisc_1(u1_struct_0(X45)))|~l1_pre_topc(X45))&(~v3_pre_topc(k7_subset_1(u1_struct_0(X45),k2_struct_0(X45),X46),X45)|v4_pre_topc(X46,X45)|~m1_subset_1(X46,k1_zfmisc_1(u1_struct_0(X45)))|~l1_pre_topc(X45))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_pre_topc])])])])).
fof(c_0_78, plain, ![X28, X29]:k2_tarski(X28,X29)=k2_tarski(X29,X28), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_79, plain, (k1_setfam_1(k2_tarski(X1,k7_subset_1(X1,X1,X2)))=k7_subset_1(X1,X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_74]), c_0_74])).
cnf(c_0_80, negated_conjecture, (k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1))=X1|~l1_struct_0(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_81, plain, (v4_pre_topc(X2,X1)|~v3_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_82, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_83, negated_conjecture, (k1_setfam_1(k2_tarski(u1_struct_0(esk2_0),X1))=X1|~l1_struct_0(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_84, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_85, plain, ![X47, X48]:((~v4_pre_topc(X48,X47)|k2_pre_topc(X47,X48)=X48|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X47)))|~l1_pre_topc(X47))&(~v2_pre_topc(X47)|k2_pre_topc(X47,X48)!=X48|v4_pre_topc(X48,X47)|~m1_subset_1(X48,k1_zfmisc_1(u1_struct_0(X47)))|~l1_pre_topc(X47))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t52_pre_topc])])])])).
cnf(c_0_86, plain, (v4_pre_topc(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),X1)|~v3_pre_topc(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(k7_subset_1(u1_struct_0(X1),k2_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_75]), c_0_66])).
cnf(c_0_87, plain, (k5_xboole_0(X1,k1_setfam_1(k2_tarski(X1,k1_setfam_1(k2_tarski(X2,X1)))))=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X2,X1)))), inference(spm,[status(thm)],[c_0_49, c_0_82])).
cnf(c_0_88, negated_conjecture, (k1_setfam_1(k2_tarski(esk3_0,u1_struct_0(esk2_0)))=esk3_0|~l1_struct_0(esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_82])).
cnf(c_0_89, plain, (k2_pre_topc(X2,X1)=X1|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_90, plain, (m1_subset_1(k5_xboole_0(X1,k1_setfam_1(k2_tarski(X2,X1))),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_55, c_0_82])).
cnf(c_0_91, negated_conjecture, (v4_pre_topc(k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),esk2_0)|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_76]), c_0_72])])).
cnf(c_0_92, plain, (m1_subset_1(k7_subset_1(X1,X1,X2),k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_55, c_0_74])).
cnf(c_0_93, plain, (k7_subset_1(X1,X1,k1_setfam_1(k2_tarski(X2,X1)))=k5_xboole_0(X1,k1_setfam_1(k2_tarski(X2,X1)))), inference(rw,[status(thm)],[c_0_87, c_0_74])).
cnf(c_0_94, negated_conjecture, (k1_setfam_1(k2_tarski(esk3_0,u1_struct_0(esk2_0)))=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_66]), c_0_72])])).
cnf(c_0_95, negated_conjecture, (k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)|v3_pre_topc(esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_96, plain, (v3_pre_topc(k7_subset_1(u1_struct_0(X2),k2_struct_0(X2),X1),X2)|~v4_pre_topc(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_97, negated_conjecture, (v2_pre_topc(esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_98, plain, (k2_pre_topc(X1,k5_xboole_0(u1_struct_0(X1),k1_setfam_1(k2_tarski(X2,u1_struct_0(X1)))))=k5_xboole_0(u1_struct_0(X1),k1_setfam_1(k2_tarski(X2,u1_struct_0(X1))))|~v4_pre_topc(k5_xboole_0(u1_struct_0(X1),k1_setfam_1(k2_tarski(X2,u1_struct_0(X1)))),X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_99, negated_conjecture, (v4_pre_topc(k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),esk2_0)|~v3_pre_topc(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_92])])).
cnf(c_0_100, negated_conjecture, (k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0)=k5_xboole_0(u1_struct_0(esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_101, negated_conjecture, (k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0))=k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0)|v3_pre_topc(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_76]), c_0_76])).
cnf(c_0_102, plain, (v3_pre_topc(X1,X2)|~v4_pre_topc(k7_subset_1(u1_struct_0(X2),k2_struct_0(X2),X1),X2)|~l1_pre_topc(X2)|~m1_subset_1(k7_subset_1(u1_struct_0(X2),k2_struct_0(X2),X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_75]), c_0_66])).
cnf(c_0_103, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|k2_pre_topc(X1,X2)!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_104, negated_conjecture, (v2_pre_topc(esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_76]), c_0_76])).
cnf(c_0_105, negated_conjecture, (k2_pre_topc(esk2_0,k5_xboole_0(u1_struct_0(esk2_0),esk3_0))=k5_xboole_0(u1_struct_0(esk2_0),esk3_0)|~v4_pre_topc(k5_xboole_0(u1_struct_0(esk2_0),esk3_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_94]), c_0_72])])).
cnf(c_0_106, negated_conjecture, (v4_pre_topc(k5_xboole_0(u1_struct_0(esk2_0),esk3_0),esk2_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_84])])).
cnf(c_0_107, negated_conjecture, (k2_pre_topc(esk2_0,k5_xboole_0(u1_struct_0(esk2_0),esk3_0))=k5_xboole_0(u1_struct_0(esk2_0),esk3_0)|v3_pre_topc(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_100]), c_0_100])).
cnf(c_0_108, negated_conjecture, (~v3_pre_topc(esk3_0,esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),k2_struct_0(esk2_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_109, negated_conjecture, (v3_pre_topc(X1,esk2_0)|~v4_pre_topc(k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),esk2_0)|~m1_subset_1(k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_76]), c_0_72])])).
cnf(c_0_110, plain, (v4_pre_topc(k5_xboole_0(u1_struct_0(X1),k1_setfam_1(k2_tarski(u1_struct_0(X1),X2))),X1)|k2_pre_topc(X1,k5_xboole_0(u1_struct_0(X1),k1_setfam_1(k2_tarski(u1_struct_0(X1),X2))))!=k5_xboole_0(u1_struct_0(X1),k1_setfam_1(k2_tarski(u1_struct_0(X1),X2)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_103, c_0_55])).
cnf(c_0_111, negated_conjecture, (v2_pre_topc(esk2_0)|k2_pre_topc(esk2_0,k5_xboole_0(u1_struct_0(esk2_0),esk3_0))!=k5_xboole_0(u1_struct_0(esk2_0),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_100]), c_0_100])).
cnf(c_0_112, negated_conjecture, (k2_pre_topc(esk2_0,k5_xboole_0(u1_struct_0(esk2_0),esk3_0))=k5_xboole_0(u1_struct_0(esk2_0),esk3_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_107])).
cnf(c_0_113, negated_conjecture, (k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0))!=k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),esk3_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_76]), c_0_76])).
cnf(c_0_114, negated_conjecture, (v3_pre_topc(X1,esk2_0)|~v4_pre_topc(k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_92])])).
cnf(c_0_115, plain, (v4_pre_topc(k7_subset_1(u1_struct_0(X1),u1_struct_0(X1),X2),X1)|k2_pre_topc(X1,k7_subset_1(u1_struct_0(X1),u1_struct_0(X1),X2))!=k7_subset_1(u1_struct_0(X1),u1_struct_0(X1),X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_74]), c_0_74]), c_0_74])).
cnf(c_0_116, negated_conjecture, (v2_pre_topc(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_117, negated_conjecture, (k2_pre_topc(esk2_0,k5_xboole_0(u1_struct_0(esk2_0),esk3_0))!=k5_xboole_0(u1_struct_0(esk2_0),esk3_0)|~v3_pre_topc(esk3_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_100]), c_0_100])).
cnf(c_0_118, negated_conjecture, (v3_pre_topc(X1,esk2_0)|k2_pre_topc(esk2_0,k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1))!=k7_subset_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_72])]), c_0_116])])).
cnf(c_0_119, negated_conjecture, (~v3_pre_topc(esk3_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_117, c_0_112])])).
cnf(c_0_120, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_100]), c_0_112]), c_0_84])]), c_0_119]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 121
# Proof object clause steps            : 80
# Proof object formula steps           : 41
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 20
# Proof object generating inferences   : 30
# Proof object simplifying inferences  : 68
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 29
# Processed clauses                    : 830
# ...of these trivial                  : 66
# ...subsumed                          : 460
# ...remaining for further processing  : 304
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 15
# Backward-rewritten                   : 102
# Generated clauses                    : 2106
# ...of the previous two non-trivial   : 1578
# Contextual simplify-reflections      : 4
# Paramodulations                      : 2098
# Factorizations                       : 4
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 154
#    Positive orientable unit clauses  : 56
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 83
# Current number of unprocessed clauses: 731
# ...number of literals in the above   : 1470
# Current number of archived formulas  : 0
# Current number of archived clauses   : 150
# Clause-clause subsumption calls (NU) : 2118
# Rec. Clause-clause subsumption calls : 1683
# Non-unit clause-clause subsumptions  : 192
# Unit Clause-clause subsumption calls : 634
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 363
# BW rewrite match successes           : 68
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 27599

# -------------------------------------------------
# User time                : 0.026 s
# System time              : 0.000 s
# Total time               : 0.026 s
# Maximum resident set size: 28640 pages
