# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t69_mcart_1, conjecture, ![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X6))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k5_mcart_1(X1,X2,X3,X5)))))), file('data/MPTP/proofs/eb02873c9f324e3183dd1b89d8d88392.E_input', t69_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/eb02873c9f324e3183dd1b89d8d88392.E_input', d3_mcart_1)).
fof(l44_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&?[X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))&![X5]:(m1_subset_1(X5,X1)=>![X6]:(m1_subset_1(X6,X2)=>![X7]:(m1_subset_1(X7,X3)=>~(X4=k3_mcart_1(X5,X6,X7)))))))))), file('data/MPTP/proofs/eb02873c9f324e3183dd1b89d8d88392.E_input', l44_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/eb02873c9f324e3183dd1b89d8d88392.E_input', d3_zfmisc_1)).
fof(t47_mcart_1, axiom, ![X1, X2, X3]:~((~(X1=k1_xboole_0)&(~(X2=k1_xboole_0)&(~(X3=k1_xboole_0)&?[X4]:(m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))&?[X5, X6, X7]:(X4=k3_mcart_1(X5,X6,X7)&~((k5_mcart_1(X1,X2,X3,X4)=X5&(k6_mcart_1(X1,X2,X3,X4)=X6&k7_mcart_1(X1,X2,X3,X4)=X7))))))))), file('data/MPTP/proofs/eb02873c9f324e3183dd1b89d8d88392.E_input', t47_mcart_1)).
fof(c_0_5, negated_conjecture, ~(![X1, X2, X3, X4, X5]:(m1_subset_1(X5,k3_zfmisc_1(X1,X2,X3))=>(![X6]:(m1_subset_1(X6,X1)=>![X7]:(m1_subset_1(X7,X2)=>![X8]:(m1_subset_1(X8,X3)=>(X5=k3_mcart_1(X6,X7,X8)=>X4=X6))))=>(X1=k1_xboole_0|(X2=k1_xboole_0|(X3=k1_xboole_0|X4=k5_mcart_1(X1,X2,X3,X5))))))), inference(assume_negation,[status(cth)],[t69_mcart_1])).
fof(c_0_6, negated_conjecture, ![X14, X15, X16]:(m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))&((~m1_subset_1(X14,esk1_0)|(~m1_subset_1(X15,esk2_0)|(~m1_subset_1(X16,esk3_0)|(esk5_0!=k3_mcart_1(X14,X15,X16)|esk4_0=X14))))&(esk1_0!=k1_xboole_0&(esk2_0!=k1_xboole_0&(esk3_0!=k1_xboole_0&esk4_0!=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])])).
fof(c_0_7, plain, ![X34, X35, X36]:k3_mcart_1(X34,X35,X36)=k4_tarski(k4_tarski(X34,X35),X36), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_8, plain, ![X17, X18, X19, X20]:((m1_subset_1(esk6_4(X17,X18,X19,X20),X17)|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0)&((m1_subset_1(esk7_4(X17,X18,X19,X20),X18)|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0)&((m1_subset_1(esk8_4(X17,X18,X19,X20),X19)|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0)&(X20=k3_mcart_1(esk6_4(X17,X18,X19,X20),esk7_4(X17,X18,X19,X20),esk8_4(X17,X18,X19,X20))|~m1_subset_1(X20,k3_zfmisc_1(X17,X18,X19))|X19=k1_xboole_0|X18=k1_xboole_0|X17=k1_xboole_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l44_mcart_1])])])])])])])).
fof(c_0_9, plain, ![X31, X32, X33]:k3_zfmisc_1(X31,X32,X33)=k2_zfmisc_1(k2_zfmisc_1(X31,X32),X33), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
cnf(c_0_10, negated_conjecture, (esk4_0=X1|~m1_subset_1(X1,esk1_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X3,esk3_0)|esk5_0!=k3_mcart_1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (X1=k3_mcart_1(esk6_4(X2,X3,X4,X1),esk7_4(X2,X3,X4,X1),esk8_4(X2,X3,X4,X1))|X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|~m1_subset_1(X1,k3_zfmisc_1(X2,X3,X4))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, negated_conjecture, (esk4_0=X1|esk5_0!=k4_tarski(k4_tarski(X1,X2),X3)|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(rw,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (X4=k1_xboole_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k4_tarski(k4_tarski(esk6_4(X2,X3,X4,X1),esk7_4(X2,X3,X4,X1)),esk8_4(X2,X3,X4,X1))|~m1_subset_1(X1,k2_zfmisc_1(k2_zfmisc_1(X2,X3),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_11])).
cnf(c_0_16, plain, (m1_subset_1(esk8_4(X1,X2,X3,X4),X3)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_17, negated_conjecture, (esk6_4(X1,X2,X3,esk5_0)=esk4_0|X1=k1_xboole_0|X2=k1_xboole_0|X3=k1_xboole_0|~m1_subset_1(esk8_4(X1,X2,X3,esk5_0),esk3_0)|~m1_subset_1(esk7_4(X1,X2,X3,esk5_0),esk2_0)|~m1_subset_1(esk6_4(X1,X2,X3,esk5_0),esk1_0)|~m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15])])).
cnf(c_0_18, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk8_4(X1,X2,X3,X4),X3)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_16, c_0_13])).
cnf(c_0_19, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_20, plain, (m1_subset_1(esk7_4(X1,X2,X3,X4),X2)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_21, plain, ![X24, X25, X26, X27, X28, X29, X30]:((k5_mcart_1(X24,X25,X26,X27)=X28|X27!=k3_mcart_1(X28,X29,X30)|~m1_subset_1(X27,k3_zfmisc_1(X24,X25,X26))|X26=k1_xboole_0|X25=k1_xboole_0|X24=k1_xboole_0)&((k6_mcart_1(X24,X25,X26,X27)=X29|X27!=k3_mcart_1(X28,X29,X30)|~m1_subset_1(X27,k3_zfmisc_1(X24,X25,X26))|X26=k1_xboole_0|X25=k1_xboole_0|X24=k1_xboole_0)&(k7_mcart_1(X24,X25,X26,X27)=X30|X27!=k3_mcart_1(X28,X29,X30)|~m1_subset_1(X27,k3_zfmisc_1(X24,X25,X26))|X26=k1_xboole_0|X25=k1_xboole_0|X24=k1_xboole_0))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t47_mcart_1])])])])])).
cnf(c_0_22, negated_conjecture, (esk6_4(X1,X2,esk3_0,esk5_0)=esk4_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk7_4(X1,X2,esk3_0,esk5_0),esk2_0)|~m1_subset_1(esk6_4(X1,X2,esk3_0,esk5_0),esk1_0)|~m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),esk3_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19])).
cnf(c_0_23, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk7_4(X1,X2,X3,X4),X2)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_20, c_0_13])).
cnf(c_0_24, negated_conjecture, (esk2_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_25, plain, (m1_subset_1(esk6_4(X1,X2,X3,X4),X1)|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk5_0,k3_zfmisc_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_27, plain, (k5_mcart_1(X1,X2,X3,X4)=X5|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|X4!=k3_mcart_1(X5,X6,X7)|~m1_subset_1(X4,k3_zfmisc_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, negated_conjecture, (esk6_4(X1,esk2_0,esk3_0,esk5_0)=esk4_0|X1=k1_xboole_0|~m1_subset_1(esk6_4(X1,esk2_0,esk3_0,esk5_0),esk1_0)|~m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(X1,esk2_0),esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_19])).
cnf(c_0_29, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|m1_subset_1(esk6_4(X1,X2,X3,X4),X1)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[c_0_25, c_0_13])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0),esk3_0))), inference(rw,[status(thm)],[c_0_26, c_0_13])).
cnf(c_0_31, negated_conjecture, (esk1_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_32, plain, (X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|k5_mcart_1(X1,X2,X3,X4)=X5|X4!=k4_tarski(k4_tarski(X5,X6),X7)|~m1_subset_1(X4,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_13]), c_0_11])).
cnf(c_0_33, negated_conjecture, (esk6_4(esk1_0,esk2_0,esk3_0,esk5_0)=esk4_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])]), c_0_31]), c_0_19]), c_0_24])).
cnf(c_0_34, plain, (k5_mcart_1(X1,X2,X3,k4_tarski(k4_tarski(X4,X5),X6))=X4|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(k4_tarski(k4_tarski(X4,X5),X6),k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_35, negated_conjecture, (k4_tarski(k4_tarski(esk4_0,esk7_4(esk1_0,esk2_0,esk3_0,esk5_0)),esk8_4(esk1_0,esk2_0,esk3_0,esk5_0))=esk5_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_15, c_0_33]), c_0_30])]), c_0_19]), c_0_24]), c_0_31])).
cnf(c_0_36, negated_conjecture, (esk4_0!=k5_mcart_1(esk1_0,esk2_0,esk3_0,esk5_0)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_37, negated_conjecture, (k5_mcart_1(X1,X2,X3,esk5_0)=esk4_0|X3=k1_xboole_0|X2=k1_xboole_0|X1=k1_xboole_0|~m1_subset_1(esk5_0,k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_30])]), c_0_19]), c_0_24]), c_0_31]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 28
# Proof object formula steps           : 11
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 29
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 13
# Processed clauses                    : 37
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 37
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 25
# ...of the previous two non-trivial   : 23
# Contextual simplify-reflections      : 0
# Paramodulations                      : 21
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 13
# Current number of unprocessed clauses: 12
# ...number of literals in the above   : 88
# Current number of archived formulas  : 0
# Current number of archived clauses   : 15
# Clause-clause subsumption calls (NU) : 19
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 10
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1998

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.003 s
# Total time               : 0.009 s
# Maximum resident set size: 28684 pages
