# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', d1_tarski)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', redefinition_k6_domain_1)).
fof(d1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>(v1_zfmisc_1(X1)<=>?[X2]:(m1_subset_1(X2,X1)&X1=k6_domain_1(X1,X2)))), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', d1_tex_2)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t70_enumset1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', d3_tarski)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', d1_xboole_0)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t7_ordinal1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t2_xboole_1)).
fof(t49_zfmisc_1, axiom, ![X1, X2]:~(k2_xboole_0(k1_tarski(X1),X2)=k1_xboole_0), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t49_zfmisc_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t72_enumset1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t1_boole)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t2_tarski)).
fof(t2_tex_2, conjecture, ![X1]:((~(v1_xboole_0(X1))&v1_zfmisc_1(X1))=>![X2]:(~(v1_xboole_0(k3_xboole_0(X1,X2)))=>r1_tarski(X1,X2))), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t2_tex_2)).
fof(l3_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,k1_tarski(X2))<=>(X1=k1_xboole_0|X1=k1_tarski(X2))), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', l3_zfmisc_1)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t12_setfam_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t100_xboole_1)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t17_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t3_boole)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t2_boole)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', t48_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', idempotence_k3_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', l32_xboole_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/eb527d6960e14edb8a45601e8dc3a2bd.E_input', fc1_xboole_0)).
fof(c_0_25, plain, ![X5, X6, X7, X8, X9, X10]:(((~r2_hidden(X7,X6)|X7=X5|X6!=k1_tarski(X5))&(X8!=X5|r2_hidden(X8,X6)|X6!=k1_tarski(X5)))&((~r2_hidden(esk1_2(X9,X10),X10)|esk1_2(X9,X10)!=X9|X10=k1_tarski(X9))&(r2_hidden(esk1_2(X9,X10),X10)|esk1_2(X9,X10)=X9|X10=k1_tarski(X9)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_26, plain, ![X12, X13]:(v1_xboole_0(X12)|~m1_subset_1(X13,X12)|k6_domain_1(X12,X13)=k1_tarski(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_27, plain, ![X14, X16]:(((m1_subset_1(esk2_1(X14),X14)|~v1_zfmisc_1(X14)|v1_xboole_0(X14))&(X14=k6_domain_1(X14,esk2_1(X14))|~v1_zfmisc_1(X14)|v1_xboole_0(X14)))&(~m1_subset_1(X16,X14)|X14!=k6_domain_1(X14,X16)|v1_zfmisc_1(X14)|v1_xboole_0(X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_tex_2])])])])])])).
fof(c_0_28, plain, ![X17, X18]:k3_tarski(k2_tarski(X17,X18))=k2_xboole_0(X17,X18), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_29, plain, ![X19, X20]:k1_enumset1(X19,X19,X20)=k2_tarski(X19,X20), inference(variable_rename,[status(thm)],[t70_enumset1])).
cnf(c_0_30, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_31, plain, ![X21, X22, X23, X24, X25]:((~r1_tarski(X21,X22)|(~r2_hidden(X23,X21)|r2_hidden(X23,X22)))&((r2_hidden(esk3_2(X24,X25),X24)|r1_tarski(X24,X25))&(~r2_hidden(esk3_2(X24,X25),X25)|r1_tarski(X24,X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_32, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (m1_subset_1(esk2_1(X1),X1)|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_34, plain, ![X27, X28, X29]:((~v1_xboole_0(X27)|~r2_hidden(X28,X27))&(r2_hidden(esk4_1(X29),X29)|v1_xboole_0(X29))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_35, plain, ![X31, X32]:(~r2_hidden(X31,X32)|~r1_tarski(X32,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_36, plain, ![X33]:r1_tarski(k1_xboole_0,X33), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_37, plain, ![X34, X35]:k2_xboole_0(k1_tarski(X34),X35)!=k1_xboole_0, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t49_zfmisc_1])])).
cnf(c_0_38, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_39, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_40, plain, ![X36, X37, X38]:k2_enumset1(X36,X36,X37,X38)=k1_enumset1(X36,X37,X38), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_41, plain, ![X39, X40, X41, X42]:k3_enumset1(X39,X39,X40,X41,X42)=k2_enumset1(X39,X40,X41,X42), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_42, plain, ![X43]:k2_xboole_0(X43,k1_xboole_0)=X43, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_43, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_30])).
cnf(c_0_44, plain, (r2_hidden(esk3_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_45, plain, (X1=k6_domain_1(X1,esk2_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_46, plain, (k6_domain_1(X1,esk2_1(X1))=k1_tarski(esk2_1(X1))|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_47, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_48, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_49, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_50, plain, ![X44, X45]:((~r2_hidden(esk5_2(X44,X45),X44)|~r2_hidden(esk5_2(X44,X45),X45)|X44=X45)&(r2_hidden(esk5_2(X44,X45),X44)|r2_hidden(esk5_2(X44,X45),X45)|X44=X45)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
cnf(c_0_51, plain, (k2_xboole_0(k1_tarski(X1),X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_52, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_53, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_54, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_55, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_56, plain, (esk3_2(k1_tarski(X1),X2)=X1|r1_tarski(k1_tarski(X1),X2)), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_57, plain, (k1_tarski(esk2_1(X1))=X1|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_58, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_47, c_0_44])).
fof(c_0_59, negated_conjecture, ~(![X1]:((~(v1_xboole_0(X1))&v1_zfmisc_1(X1))=>![X2]:(~(v1_xboole_0(k3_xboole_0(X1,X2)))=>r1_tarski(X1,X2)))), inference(assume_negation,[status(cth)],[t2_tex_2])).
cnf(c_0_60, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_61, plain, (r2_hidden(esk5_2(X1,X2),X1)|r2_hidden(esk5_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_62, plain, (k3_tarski(k3_enumset1(k1_tarski(X1),k1_tarski(X1),k1_tarski(X1),k1_tarski(X1),X2))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_54])).
cnf(c_0_63, plain, (k3_tarski(k3_enumset1(X1,X1,X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_52]), c_0_53]), c_0_54])).
cnf(c_0_64, plain, (esk3_2(X1,X2)=esk2_1(X1)|r1_tarski(X1,X2)|~v1_zfmisc_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_58])).
fof(c_0_65, negated_conjecture, ((~v1_xboole_0(esk6_0)&v1_zfmisc_1(esk6_0))&(~v1_xboole_0(k3_xboole_0(esk6_0,esk7_0))&~r1_tarski(esk6_0,esk7_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_59])])])])).
cnf(c_0_66, plain, (k1_xboole_0=X1|r2_hidden(esk5_2(k1_xboole_0,X1),X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_67, plain, (k1_tarski(X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_62, c_0_63])).
fof(c_0_68, plain, ![X53, X54]:((~r1_tarski(X53,k1_tarski(X54))|(X53=k1_xboole_0|X53=k1_tarski(X54)))&((X53!=k1_xboole_0|r1_tarski(X53,k1_tarski(X54)))&(X53!=k1_tarski(X54)|r1_tarski(X53,k1_tarski(X54))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_zfmisc_1])])])).
cnf(c_0_69, plain, (r1_tarski(X1,X2)|~r2_hidden(esk3_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_70, plain, (r1_tarski(X1,X2)|r2_hidden(esk2_1(X1),X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_44, c_0_64])).
cnf(c_0_71, negated_conjecture, (v1_zfmisc_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_72, plain, (esk5_2(k1_xboole_0,k1_tarski(X1))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_66]), c_0_67])).
fof(c_0_73, plain, ![X49, X50]:k1_setfam_1(k2_tarski(X49,X50))=k3_xboole_0(X49,X50), inference(variable_rename,[status(thm)],[t12_setfam_1])).
cnf(c_0_74, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|~r1_tarski(X1,k1_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_75, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_69, c_0_56])).
cnf(c_0_76, negated_conjecture, (r1_tarski(esk6_0,X1)|r2_hidden(esk2_1(esk6_0),esk6_0)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_77, plain, (esk2_1(X1)=esk5_2(k1_xboole_0,X1)|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_72, c_0_57])).
fof(c_0_78, plain, ![X51, X52]:k4_xboole_0(X51,X52)=k5_xboole_0(X51,k3_xboole_0(X51,X52)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_79, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
fof(c_0_80, plain, ![X57, X58]:r1_tarski(k3_xboole_0(X57,X58),X57), inference(variable_rename,[status(thm)],[t17_xboole_1])).
cnf(c_0_81, plain, (k1_tarski(X1)=k1_tarski(X2)|~r2_hidden(X1,k1_tarski(X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_67])).
cnf(c_0_82, negated_conjecture, (r2_hidden(esk2_1(esk6_0),esk6_0)|~r2_hidden(X1,esk6_0)), inference(spm,[status(thm)],[c_0_48, c_0_76])).
cnf(c_0_83, plain, (r2_hidden(esk4_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_84, negated_conjecture, (~v1_xboole_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_85, plain, (k1_tarski(esk5_2(k1_xboole_0,X1))=X1|v1_xboole_0(X1)|~v1_zfmisc_1(X1)), inference(spm,[status(thm)],[c_0_57, c_0_77])).
fof(c_0_86, plain, ![X55]:k4_xboole_0(X55,k1_xboole_0)=X55, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_87, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_88, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_79, c_0_39])).
fof(c_0_89, plain, ![X56]:k3_xboole_0(X56,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_90, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_91, plain, (k1_tarski(X1)=X2|~v1_zfmisc_1(X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_57]), c_0_47])).
cnf(c_0_92, negated_conjecture, (r2_hidden(esk2_1(esk6_0),esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])).
cnf(c_0_93, plain, (X1=esk5_2(k1_xboole_0,X2)|~v1_zfmisc_1(X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_85]), c_0_47])).
fof(c_0_94, plain, ![X59, X60]:k4_xboole_0(X59,k4_xboole_0(X59,X60))=k3_xboole_0(X59,X60), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_95, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_96, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k1_setfam_1(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_97, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_98, plain, ![X61]:k3_xboole_0(X61,X61)=X61, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
fof(c_0_99, plain, ![X62, X63]:((k4_xboole_0(X62,X63)!=k1_xboole_0|r1_tarski(X62,X63))&(~r1_tarski(X62,X63)|k4_xboole_0(X62,X63)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_100, plain, (r1_tarski(k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_88]), c_0_53]), c_0_54])).
cnf(c_0_101, negated_conjecture, (k1_tarski(esk2_1(esk6_0))=esk6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_71])])).
cnf(c_0_102, negated_conjecture, (esk2_1(esk6_0)=esk5_2(k1_xboole_0,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_92]), c_0_71])])).
cnf(c_0_103, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_104, plain, (k5_xboole_0(X1,k1_setfam_1(k3_enumset1(X1,X1,X1,X1,k1_xboole_0)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_96]), c_0_53]), c_0_54])).
cnf(c_0_105, plain, (k1_setfam_1(k3_enumset1(X1,X1,X1,X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_88]), c_0_53]), c_0_54])).
cnf(c_0_106, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_107, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_108, plain, (k1_setfam_1(k3_enumset1(k1_tarski(X1),k1_tarski(X1),k1_tarski(X1),k1_tarski(X1),X2))=k1_tarski(X1)|k1_setfam_1(k3_enumset1(k1_tarski(X1),k1_tarski(X1),k1_tarski(X1),k1_tarski(X1),X2))=k1_xboole_0), inference(spm,[status(thm)],[c_0_74, c_0_100])).
cnf(c_0_109, negated_conjecture, (k1_tarski(esk5_2(k1_xboole_0,esk6_0))=esk6_0), inference(rw,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_110, plain, (k5_xboole_0(X1,k1_setfam_1(k3_enumset1(X1,X1,X1,X1,k5_xboole_0(X1,k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2))))))=k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_88]), c_0_96]), c_0_96]), c_0_53]), c_0_53]), c_0_53]), c_0_54]), c_0_54]), c_0_54])).
cnf(c_0_111, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(rw,[status(thm)],[c_0_104, c_0_105])).
cnf(c_0_112, plain, (k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_88]), c_0_53]), c_0_54])).
cnf(c_0_113, negated_conjecture, (~v1_xboole_0(k3_xboole_0(esk6_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_114, plain, (r1_tarski(X1,X2)|k5_xboole_0(X1,k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_96]), c_0_53]), c_0_54])).
cnf(c_0_115, negated_conjecture, (k1_setfam_1(k3_enumset1(esk6_0,esk6_0,esk6_0,esk6_0,X1))=k1_xboole_0|k1_setfam_1(k3_enumset1(esk6_0,esk6_0,esk6_0,esk6_0,X1))=esk6_0), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_116, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_105]), c_0_111]), c_0_112])).
cnf(c_0_117, negated_conjecture, (~v1_xboole_0(k1_setfam_1(k3_enumset1(esk6_0,esk6_0,esk6_0,esk6_0,esk7_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_88]), c_0_53]), c_0_54])).
cnf(c_0_118, negated_conjecture, (k1_setfam_1(k3_enumset1(esk6_0,esk6_0,esk6_0,esk6_0,X1))=k1_xboole_0|r1_tarski(esk6_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_116])])).
cnf(c_0_119, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_120, negated_conjecture, (~r1_tarski(esk6_0,esk7_0)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_121, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_119])]), c_0_120]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 122
# Proof object clause steps            : 72
# Proof object formula steps           : 50
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 25
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 56
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 25
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 34
# Processed clauses                    : 653
# ...of these trivial                  : 15
# ...subsumed                          : 399
# ...remaining for further processing  : 239
# Other redundant clauses eliminated   : 17
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 8
# Generated clauses                    : 2236
# ...of the previous two non-trivial   : 2043
# Contextual simplify-reflections      : 18
# Paramodulations                      : 2214
# Factorizations                       : 6
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 193
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 155
# Current number of unprocessed clauses: 1406
# ...number of literals in the above   : 5483
# Current number of archived formulas  : 0
# Current number of archived clauses   : 48
# Clause-clause subsumption calls (NU) : 1764
# Rec. Clause-clause subsumption calls : 1388
# Non-unit clause-clause subsumptions  : 228
# Unit Clause-clause subsumption calls : 95
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 26
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 25323

# -------------------------------------------------
# User time                : 0.024 s
# System time              : 0.005 s
# Total time               : 0.029 s
# Maximum resident set size: 28688 pages
