# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t84_mcart_1, conjecture, ![X1, X2, X3, X4, X5, X6, X7, X8]:(r2_hidden(k4_mcart_1(X1,X2,X3,X4),k4_zfmisc_1(X5,X6,X7,X8))<=>(r2_hidden(X1,X5)&(r2_hidden(X2,X6)&(r2_hidden(X3,X7)&r2_hidden(X4,X8))))), file('data/MPTP/proofs/ec66e24077ab45e69fba7c206ad09914.E_input', t84_mcart_1)).
fof(t31_mcart_1, axiom, ![X1, X2, X3, X4]:k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4), file('data/MPTP/proofs/ec66e24077ab45e69fba7c206ad09914.E_input', t31_mcart_1)).
fof(t53_mcart_1, axiom, ![X1, X2, X3, X4]:k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4), file('data/MPTP/proofs/ec66e24077ab45e69fba7c206ad09914.E_input', t53_mcart_1)).
fof(t73_mcart_1, axiom, ![X1, X2, X3, X4, X5, X6]:(r2_hidden(k3_mcart_1(X1,X2,X3),k3_zfmisc_1(X4,X5,X6))<=>(r2_hidden(X1,X4)&(r2_hidden(X2,X5)&r2_hidden(X3,X6)))), file('data/MPTP/proofs/ec66e24077ab45e69fba7c206ad09914.E_input', t73_mcart_1)).
fof(d3_mcart_1, axiom, ![X1, X2, X3]:k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3), file('data/MPTP/proofs/ec66e24077ab45e69fba7c206ad09914.E_input', d3_mcart_1)).
fof(d3_zfmisc_1, axiom, ![X1, X2, X3]:k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3), file('data/MPTP/proofs/ec66e24077ab45e69fba7c206ad09914.E_input', d3_zfmisc_1)).
fof(l54_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('data/MPTP/proofs/ec66e24077ab45e69fba7c206ad09914.E_input', l54_zfmisc_1)).
fof(c_0_7, negated_conjecture, ~(![X1, X2, X3, X4, X5, X6, X7, X8]:(r2_hidden(k4_mcart_1(X1,X2,X3,X4),k4_zfmisc_1(X5,X6,X7,X8))<=>(r2_hidden(X1,X5)&(r2_hidden(X2,X6)&(r2_hidden(X3,X7)&r2_hidden(X4,X8)))))), inference(assume_negation,[status(cth)],[t84_mcart_1])).
fof(c_0_8, negated_conjecture, ((~r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))|(~r2_hidden(esk1_0,esk5_0)|(~r2_hidden(esk2_0,esk6_0)|(~r2_hidden(esk3_0,esk7_0)|~r2_hidden(esk4_0,esk8_0)))))&((r2_hidden(esk1_0,esk5_0)|r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0)))&((r2_hidden(esk2_0,esk6_0)|r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0)))&((r2_hidden(esk3_0,esk7_0)|r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0)))&(r2_hidden(esk4_0,esk8_0)|r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_7])])])])).
fof(c_0_9, plain, ![X21, X22, X23, X24]:k4_mcart_1(X21,X22,X23,X24)=k4_tarski(k4_tarski(k4_tarski(X21,X22),X23),X24), inference(variable_rename,[status(thm)],[t31_mcart_1])).
fof(c_0_10, plain, ![X17, X18, X19, X20]:k4_zfmisc_1(X17,X18,X19,X20)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X17,X18),X19),X20), inference(variable_rename,[status(thm)],[t53_mcart_1])).
fof(c_0_11, plain, ![X25, X26, X27, X28, X29, X30]:(((r2_hidden(X25,X28)|~r2_hidden(k3_mcart_1(X25,X26,X27),k3_zfmisc_1(X28,X29,X30)))&((r2_hidden(X26,X29)|~r2_hidden(k3_mcart_1(X25,X26,X27),k3_zfmisc_1(X28,X29,X30)))&(r2_hidden(X27,X30)|~r2_hidden(k3_mcart_1(X25,X26,X27),k3_zfmisc_1(X28,X29,X30)))))&(~r2_hidden(X25,X28)|(~r2_hidden(X26,X29)|~r2_hidden(X27,X30))|r2_hidden(k3_mcart_1(X25,X26,X27),k3_zfmisc_1(X28,X29,X30)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t73_mcart_1])])])).
fof(c_0_12, plain, ![X48, X49, X50]:k3_mcart_1(X48,X49,X50)=k4_tarski(k4_tarski(X48,X49),X50), inference(variable_rename,[status(thm)],[d3_mcart_1])).
fof(c_0_13, plain, ![X51, X52, X53]:k3_zfmisc_1(X51,X52,X53)=k2_zfmisc_1(k2_zfmisc_1(X51,X52),X53), inference(variable_rename,[status(thm)],[d3_zfmisc_1])).
cnf(c_0_14, negated_conjecture, (~r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))|~r2_hidden(esk1_0,esk5_0)|~r2_hidden(esk2_0,esk6_0)|~r2_hidden(esk3_0,esk7_0)|~r2_hidden(esk4_0,esk8_0)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_15, plain, (k4_mcart_1(X1,X2,X3,X4)=k4_tarski(k4_tarski(k4_tarski(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_16, plain, (k4_zfmisc_1(X1,X2,X3,X4)=k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3),X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_17, plain, ![X54, X55, X56, X57]:(((r2_hidden(X54,X56)|~r2_hidden(k4_tarski(X54,X55),k2_zfmisc_1(X56,X57)))&(r2_hidden(X55,X57)|~r2_hidden(k4_tarski(X54,X55),k2_zfmisc_1(X56,X57))))&(~r2_hidden(X54,X56)|~r2_hidden(X55,X57)|r2_hidden(k4_tarski(X54,X55),k2_zfmisc_1(X56,X57)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l54_zfmisc_1])])])).
cnf(c_0_18, plain, (r2_hidden(X1,X2)|~r2_hidden(k3_mcart_1(X1,X3,X4),k3_zfmisc_1(X2,X5,X6))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (k3_mcart_1(X1,X2,X3)=k4_tarski(k4_tarski(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (k3_zfmisc_1(X1,X2,X3)=k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, negated_conjecture, (r2_hidden(esk1_0,esk5_0)|r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_22, negated_conjecture, (r2_hidden(esk2_0,esk6_0)|r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_23, negated_conjecture, (r2_hidden(esk3_0,esk7_0)|r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_24, plain, (r2_hidden(X1,X2)|~r2_hidden(k3_mcart_1(X3,X1,X4),k3_zfmisc_1(X5,X2,X6))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_25, negated_conjecture, (~r2_hidden(esk1_0,esk5_0)|~r2_hidden(esk2_0,esk6_0)|~r2_hidden(esk3_0,esk7_0)|~r2_hidden(esk4_0,esk8_0)|~r2_hidden(k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_26, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(k4_tarski(X1,X3),X4),k2_zfmisc_1(k2_zfmisc_1(X2,X5),X6))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_20])).
cnf(c_0_28, negated_conjecture, (r2_hidden(esk1_0,esk5_0)|r2_hidden(k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_15]), c_0_16])).
cnf(c_0_29, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, negated_conjecture, (r2_hidden(esk2_0,esk6_0)|r2_hidden(k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_15]), c_0_16])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk3_0,esk7_0)|r2_hidden(k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_23, c_0_15]), c_0_16])).
cnf(c_0_32, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(k4_tarski(X3,X1),X4),k2_zfmisc_1(k2_zfmisc_1(X5,X2),X6))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_19]), c_0_20])).
cnf(c_0_33, negated_conjecture, (r2_hidden(esk4_0,esk8_0)|r2_hidden(k4_mcart_1(esk1_0,esk2_0,esk3_0,esk4_0),k4_zfmisc_1(esk5_0,esk6_0,esk7_0,esk8_0))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_34, negated_conjecture, (~r2_hidden(k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0))|~r2_hidden(esk1_0,esk5_0)|~r2_hidden(esk2_0,esk6_0)|~r2_hidden(esk3_0,esk7_0)), inference(csr,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_35, negated_conjecture, (r2_hidden(esk1_0,esk5_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_28]), c_0_29])).
cnf(c_0_36, negated_conjecture, (r2_hidden(esk2_0,esk6_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_30]), c_0_26])).
cnf(c_0_37, negated_conjecture, (r2_hidden(esk3_0,esk7_0)), inference(csr,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_38, negated_conjecture, (r2_hidden(esk4_0,esk8_0)|r2_hidden(k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_15]), c_0_16])).
cnf(c_0_39, negated_conjecture, (~r2_hidden(k4_tarski(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),esk4_0),k2_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0),esk8_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35])]), c_0_36]), c_0_37])])).
cnf(c_0_40, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk4_0,esk8_0)), inference(csr,[status(thm)],[c_0_38, c_0_26])).
cnf(c_0_42, negated_conjecture, (~r2_hidden(k4_tarski(k4_tarski(esk1_0,esk2_0),esk3_0),k2_zfmisc_1(k2_zfmisc_1(esk5_0,esk6_0),esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])])).
cnf(c_0_43, negated_conjecture, (~r2_hidden(k4_tarski(esk1_0,esk2_0),k2_zfmisc_1(esk5_0,esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_40]), c_0_37])])).
cnf(c_0_44, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_40]), c_0_36]), c_0_35])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 45
# Proof object clause steps            : 30
# Proof object formula steps           : 15
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 31
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 20
# Processed clauses                    : 44
# ...of these trivial                  : 0
# ...subsumed                          : 3
# ...remaining for further processing  : 41
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 22
# ...of the previous two non-trivial   : 16
# Contextual simplify-reflections      : 5
# Paramodulations                      : 18
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 14
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 7
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 23
# Current number of archived formulas  : 0
# Current number of archived clauses   : 27
# Clause-clause subsumption calls (NU) : 56
# Rec. Clause-clause subsumption calls : 53
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2049

# -------------------------------------------------
# User time                : 0.010 s
# System time              : 0.000 s
# Total time               : 0.010 s
# Maximum resident set size: 28632 pages
