# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PS_S3S
# and selection function SelectNewComplexAHPExceptUniqMaxHorn.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t3_xboole_0)).
fof(t55_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_xboole_0(k2_tarski(X1,X2),X3)&r2_hidden(X1,X3))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t55_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t70_enumset1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t75_enumset1)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t66_xboole_1)).
fof(t166_relat_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>(r2_hidden(X1,k10_relat_1(X3,X2))<=>?[X4]:(r2_hidden(X4,k2_relat_1(X3))&(r2_hidden(k4_tarski(X1,X4),X3)&r2_hidden(X4,X2))))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t166_relat_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t7_xboole_0)).
fof(t168_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>k10_relat_1(X2,X1)=k10_relat_1(X2,k3_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t168_relat_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', d7_xboole_0)).
fof(t173_relat_1, conjecture, ![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', t173_relat_1)).
fof(d14_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2, X3]:(X3=k10_relat_1(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(k4_tarski(X4,X5),X1)&r2_hidden(X5,X2))))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', d14_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', d5_relat_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/ecf789397fc24045adfb34b6a0b4376e.E_input', symmetry_r1_xboole_0)).
fof(c_0_17, plain, ![X8, X9, X11, X12, X13]:(((r2_hidden(esk1_2(X8,X9),X8)|r1_xboole_0(X8,X9))&(r2_hidden(esk1_2(X8,X9),X9)|r1_xboole_0(X8,X9)))&(~r2_hidden(X13,X11)|~r2_hidden(X13,X12)|~r1_xboole_0(X11,X12))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_18, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_19, plain, (r2_hidden(esk1_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_20, plain, ![X14, X15, X16]:(~r1_xboole_0(k2_tarski(X14,X15),X16)|~r2_hidden(X14,X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_zfmisc_1])])).
fof(c_0_21, plain, ![X17, X18]:k1_enumset1(X17,X17,X18)=k2_tarski(X17,X18), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_22, plain, ![X19, X20, X21]:k2_enumset1(X19,X19,X20,X21)=k1_enumset1(X19,X20,X21), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_23, plain, ![X22, X23, X24, X25]:k3_enumset1(X22,X22,X23,X24,X25)=k2_enumset1(X22,X23,X24,X25), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_24, plain, ![X26, X27, X28, X29, X30]:k4_enumset1(X26,X26,X27,X28,X29,X30)=k3_enumset1(X26,X27,X28,X29,X30), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_25, plain, ![X31, X32, X33, X34, X35, X36]:k5_enumset1(X31,X31,X32,X33,X34,X35,X36)=k4_enumset1(X31,X32,X33,X34,X35,X36), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_26, plain, ![X37, X38, X39, X40, X41, X42, X43]:k6_enumset1(X37,X37,X38,X39,X40,X41,X42,X43)=k5_enumset1(X37,X38,X39,X40,X41,X42,X43), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_27, plain, (r1_xboole_0(X1,X2)|~r2_hidden(esk1_2(X1,X2),X3)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_28, plain, ![X44]:((r1_xboole_0(X44,X44)|X44!=k1_xboole_0)&(X44=k1_xboole_0|~r1_xboole_0(X44,X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
cnf(c_0_29, plain, (~r1_xboole_0(k2_tarski(X1,X2),X3)|~r2_hidden(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_32, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_33, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_34, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_35, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_36, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_27, c_0_19])).
cnf(c_0_37, plain, (r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (~r2_hidden(X1,X3)|~r1_xboole_0(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_30]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_35])).
cnf(c_0_39, plain, (r1_xboole_0(X1,X2)|X2!=k1_xboole_0), inference(spm,[status(thm)],[c_0_36, c_0_37])).
fof(c_0_40, plain, ![X45, X46, X47, X49]:(((r2_hidden(esk2_3(X45,X46,X47),k2_relat_1(X47))|~r2_hidden(X45,k10_relat_1(X47,X46))|~v1_relat_1(X47))&((r2_hidden(k4_tarski(X45,esk2_3(X45,X46,X47)),X47)|~r2_hidden(X45,k10_relat_1(X47,X46))|~v1_relat_1(X47))&(r2_hidden(esk2_3(X45,X46,X47),X46)|~r2_hidden(X45,k10_relat_1(X47,X46))|~v1_relat_1(X47))))&(~r2_hidden(X49,k2_relat_1(X47))|(~r2_hidden(k4_tarski(X45,X49),X47)|~r2_hidden(X49,X46))|r2_hidden(X45,k10_relat_1(X47,X46))|~v1_relat_1(X47))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t166_relat_1])])])])])).
cnf(c_0_41, plain, (X1!=k1_xboole_0|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_42, plain, (r2_hidden(esk2_3(X1,X2,X3),X2)|~r2_hidden(X1,k10_relat_1(X3,X2))|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_43, plain, ![X56]:(X56=k1_xboole_0|r2_hidden(esk5_1(X56),X56)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_44, plain, ![X50, X51]:(~v1_relat_1(X51)|k10_relat_1(X51,X50)=k10_relat_1(X51,k3_xboole_0(k2_relat_1(X51),X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t168_relat_1])])).
fof(c_0_45, plain, ![X52, X53]:((~r1_xboole_0(X52,X53)|k3_xboole_0(X52,X53)=k1_xboole_0)&(k3_xboole_0(X52,X53)!=k1_xboole_0|r1_xboole_0(X52,X53))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_46, negated_conjecture, ~(![X1, X2]:(v1_relat_1(X2)=>(k10_relat_1(X2,X1)=k1_xboole_0<=>r1_xboole_0(k2_relat_1(X2),X1)))), inference(assume_negation,[status(cth)],[t173_relat_1])).
cnf(c_0_47, plain, (X1!=k1_xboole_0|~v1_relat_1(X2)|~r2_hidden(X3,k10_relat_1(X2,X1))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_48, plain, (X1=k1_xboole_0|r2_hidden(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (k10_relat_1(X1,X2)=k10_relat_1(X1,k3_xboole_0(k2_relat_1(X1),X2))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_50, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_51, negated_conjecture, (v1_relat_1(esk4_0)&((k10_relat_1(esk4_0,esk3_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk4_0),esk3_0))&(k10_relat_1(esk4_0,esk3_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk4_0),esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_46])])])).
cnf(c_0_52, plain, (k10_relat_1(X1,X2)=k1_xboole_0|X2!=k1_xboole_0|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_53, plain, (k10_relat_1(X1,k1_xboole_0)=k10_relat_1(X1,X2)|~v1_relat_1(X1)|~r1_xboole_0(k2_relat_1(X1),X2)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_54, negated_conjecture, (k10_relat_1(esk4_0,esk3_0)=k1_xboole_0|r1_xboole_0(k2_relat_1(esk4_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_55, negated_conjecture, (v1_relat_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_56, plain, (k10_relat_1(X1,k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(er,[status(thm)],[c_0_52])).
fof(c_0_57, plain, ![X58, X59, X60, X61, X63, X64, X65, X66, X68]:((((r2_hidden(k4_tarski(X61,esk6_4(X58,X59,X60,X61)),X58)|~r2_hidden(X61,X60)|X60!=k10_relat_1(X58,X59)|~v1_relat_1(X58))&(r2_hidden(esk6_4(X58,X59,X60,X61),X59)|~r2_hidden(X61,X60)|X60!=k10_relat_1(X58,X59)|~v1_relat_1(X58)))&(~r2_hidden(k4_tarski(X63,X64),X58)|~r2_hidden(X64,X59)|r2_hidden(X63,X60)|X60!=k10_relat_1(X58,X59)|~v1_relat_1(X58)))&((~r2_hidden(esk7_3(X58,X65,X66),X66)|(~r2_hidden(k4_tarski(esk7_3(X58,X65,X66),X68),X58)|~r2_hidden(X68,X65))|X66=k10_relat_1(X58,X65)|~v1_relat_1(X58))&((r2_hidden(k4_tarski(esk7_3(X58,X65,X66),esk8_3(X58,X65,X66)),X58)|r2_hidden(esk7_3(X58,X65,X66),X66)|X66=k10_relat_1(X58,X65)|~v1_relat_1(X58))&(r2_hidden(esk8_3(X58,X65,X66),X65)|r2_hidden(esk7_3(X58,X65,X66),X66)|X66=k10_relat_1(X58,X65)|~v1_relat_1(X58))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_relat_1])])])])])])).
cnf(c_0_58, negated_conjecture, (k10_relat_1(esk4_0,esk3_0)=k10_relat_1(esk4_0,k1_xboole_0)|k10_relat_1(esk4_0,esk3_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55])])).
cnf(c_0_59, negated_conjecture, (k10_relat_1(esk4_0,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_56, c_0_55])).
cnf(c_0_60, plain, (r2_hidden(X1,X5)|~r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(X2,X4)|X5!=k10_relat_1(X3,X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_61, negated_conjecture, (k10_relat_1(esk4_0,esk3_0)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_62, negated_conjecture, (X1!=k1_xboole_0|~r2_hidden(k4_tarski(X2,X3),esk4_0)|~r2_hidden(X3,esk3_0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_55])]), c_0_41])).
fof(c_0_63, plain, ![X70, X71, X72, X74, X75, X76, X77, X79]:(((~r2_hidden(X72,X71)|r2_hidden(k4_tarski(esk9_3(X70,X71,X72),X72),X70)|X71!=k2_relat_1(X70))&(~r2_hidden(k4_tarski(X75,X74),X70)|r2_hidden(X74,X71)|X71!=k2_relat_1(X70)))&((~r2_hidden(esk10_2(X76,X77),X77)|~r2_hidden(k4_tarski(X79,esk10_2(X76,X77)),X76)|X77=k2_relat_1(X76))&(r2_hidden(esk10_2(X76,X77),X77)|r2_hidden(k4_tarski(esk11_2(X76,X77),esk10_2(X76,X77)),X76)|X77=k2_relat_1(X76)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_64, negated_conjecture, (~r2_hidden(k4_tarski(X1,X2),esk4_0)|~r2_hidden(X2,esk3_0)), inference(er,[status(thm)],[c_0_62])).
cnf(c_0_65, plain, (r2_hidden(k4_tarski(esk9_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_66, negated_conjecture, (X1!=k2_relat_1(esk4_0)|~r2_hidden(X2,esk3_0)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_67, negated_conjecture, (~r2_hidden(X1,k2_relat_1(esk4_0))|~r2_hidden(X1,esk3_0)), inference(er,[status(thm)],[c_0_66])).
fof(c_0_68, plain, ![X81, X82]:(~r1_xboole_0(X81,X82)|r1_xboole_0(X82,X81)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
cnf(c_0_69, negated_conjecture, (r1_xboole_0(X1,k2_relat_1(esk4_0))|~r2_hidden(esk1_2(X1,k2_relat_1(esk4_0)),esk3_0)), inference(spm,[status(thm)],[c_0_67, c_0_19])).
cnf(c_0_70, plain, (r2_hidden(esk1_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_71, negated_conjecture, (k10_relat_1(esk4_0,esk3_0)!=k1_xboole_0|~r1_xboole_0(k2_relat_1(esk4_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_72, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_73, negated_conjecture, (r1_xboole_0(esk3_0,k2_relat_1(esk4_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_74, negated_conjecture, (~r1_xboole_0(k2_relat_1(esk4_0),esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_61])])).
cnf(c_0_75, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_74]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 76
# Proof object clause steps            : 41
# Proof object formula steps           : 35
# Proof object conjectures             : 17
# Proof object clause conjectures      : 14
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 6
# Initial clauses in saturation        : 28
# Processed clauses                    : 323
# ...of these trivial                  : 5
# ...subsumed                          : 177
# ...remaining for further processing  : 141
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 10
# Generated clauses                    : 565
# ...of the previous two non-trivial   : 495
# Contextual simplify-reflections      : 2
# Paramodulations                      : 550
# Factorizations                       : 1
# Equation resolutions                 : 14
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 101
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 88
# Current number of unprocessed clauses: 221
# ...number of literals in the above   : 848
# Current number of archived formulas  : 0
# Current number of archived clauses   : 46
# Clause-clause subsumption calls (NU) : 2275
# Rec. Clause-clause subsumption calls : 1086
# Non-unit clause-clause subsumptions  : 136
# Unit Clause-clause subsumption calls : 132
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7821

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28600 pages
