# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_B07_F1_AE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_zfmisc_1, conjecture, ![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4)))), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t28_zfmisc_1)).
fof(t70_xboole_1, axiom, ![X1, X2, X3]:(~((~(r1_xboole_0(X1,k2_xboole_0(X2,X3)))&(r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3))))&~((~((r1_xboole_0(X1,X2)&r1_xboole_0(X1,X3)))&r1_xboole_0(X1,k2_xboole_0(X2,X3))))), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t70_xboole_1)).
fof(t51_xboole_1, axiom, ![X1, X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t51_xboole_1)).
fof(l45_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_tarski(X2,X3))<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', l45_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t70_enumset1)).
fof(t89_enumset1, axiom, ![X1, X2, X3]:k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t89_enumset1)).
fof(t81_xboole_1, axiom, ![X1, X2, X3]:(r1_xboole_0(X1,k4_xboole_0(X2,X3))=>r1_xboole_0(X2,k4_xboole_0(X1,X3))), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t81_xboole_1)).
fof(t66_xboole_1, axiom, ![X1]:(~((~(r1_xboole_0(X1,X1))&X1=k1_xboole_0))&~((~(X1=k1_xboole_0)&r1_xboole_0(X1,X1)))), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t66_xboole_1)).
fof(t25_zfmisc_1, axiom, ![X1, X2, X3]:~((r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))&(~(X1=X2)&~(X1=X3)))), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t25_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', d10_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', l32_xboole_1)).
fof(t9_zfmisc_1, axiom, ![X1, X2, X3]:(k1_tarski(X1)=k2_tarski(X2,X3)=>X2=X3), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t9_zfmisc_1)).
fof(t12_zfmisc_1, axiom, ![X1, X2]:r1_tarski(k1_tarski(X1),k2_tarski(X1,X2)), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t12_zfmisc_1)).
fof(t18_zfmisc_1, axiom, ![X1, X2]:(k3_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2), file('data/MPTP/proofs/ed2900750a4345bb8940683014739864.E_input', t18_zfmisc_1)).
fof(c_0_14, negated_conjecture, ~(![X1, X2, X3, X4]:~((r1_tarski(k2_tarski(X1,X2),k2_tarski(X3,X4))&(~(X1=X3)&~(X1=X4))))), inference(assume_negation,[status(cth)],[t28_zfmisc_1])).
fof(c_0_15, plain, ![X17, X18, X19]:((r1_xboole_0(X17,k2_xboole_0(X18,X19))|(~r1_xboole_0(X17,X18)|~r1_xboole_0(X17,X19)))&((r1_xboole_0(X17,X18)|~r1_xboole_0(X17,k2_xboole_0(X18,X19)))&(r1_xboole_0(X17,X19)|~r1_xboole_0(X17,k2_xboole_0(X18,X19))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t70_xboole_1])])])])).
fof(c_0_16, plain, ![X20, X21]:k2_xboole_0(k3_xboole_0(X20,X21),k4_xboole_0(X20,X21))=X20, inference(variable_rename,[status(thm)],[t51_xboole_1])).
fof(c_0_17, plain, ![X9, X10, X11]:((~r1_tarski(X9,k2_tarski(X10,X11))|(X9=k1_xboole_0|(X9=k1_tarski(X10)|(X9=k1_tarski(X11)|X9=k2_tarski(X10,X11)))))&((X9!=k1_xboole_0|r1_tarski(X9,k2_tarski(X10,X11)))&((X9!=k1_tarski(X10)|r1_tarski(X9,k2_tarski(X10,X11)))&((X9!=k1_tarski(X11)|r1_tarski(X9,k2_tarski(X10,X11)))&(X9!=k2_tarski(X10,X11)|r1_tarski(X9,k2_tarski(X10,X11))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l45_zfmisc_1])])])])).
fof(c_0_18, plain, ![X12, X13]:k1_enumset1(X12,X12,X13)=k2_tarski(X12,X13), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_19, plain, ![X14, X15, X16]:k5_enumset1(X14,X14,X14,X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t89_enumset1])).
fof(c_0_20, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))&(esk1_0!=esk3_0&esk1_0!=esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_21, plain, ![X25, X26, X27]:(~r1_xboole_0(X25,k4_xboole_0(X26,X27))|r1_xboole_0(X26,k4_xboole_0(X25,X27))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t81_xboole_1])])).
cnf(c_0_22, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_24, plain, ![X28]:((r1_xboole_0(X28,X28)|X28!=k1_xboole_0)&(X28=k1_xboole_0|~r1_xboole_0(X28,X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t66_xboole_1])])])).
fof(c_0_25, plain, ![X22, X23, X24]:(~r1_tarski(k1_tarski(X22),k2_tarski(X23,X24))|(X22=X23|X22=X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_zfmisc_1])])])).
cnf(c_0_26, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|X1=k1_tarski(X3)|X1=k2_tarski(X2,X3)|~r1_tarski(X1,k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_28, plain, (k5_enumset1(X1,X1,X1,X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_29, negated_conjecture, (r1_tarski(k2_tarski(esk1_0,esk2_0),k2_tarski(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (r1_xboole_0(X2,k4_xboole_0(X1,X3))|~r1_xboole_0(X1,k4_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_32, plain, (r1_xboole_0(X1,X1)|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_33, plain, (X1=X2|X1=X3|~r1_tarski(k1_tarski(X1),k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_34, plain, (X1=k1_xboole_0|X1=k1_tarski(X3)|X1=k1_tarski(X2)|X1=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)|~r1_tarski(X1,k5_enumset1(X2,X2,X2,X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_27]), c_0_28]), c_0_28])).
cnf(c_0_35, negated_conjecture, (r1_tarski(k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_29, c_0_27]), c_0_27]), c_0_28]), c_0_28])).
fof(c_0_36, plain, ![X29, X30]:(((r1_tarski(X29,X30)|X29!=X30)&(r1_tarski(X30,X29)|X29!=X30))&(~r1_tarski(X29,X30)|~r1_tarski(X30,X29)|X29=X30)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_37, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_38, plain, (r1_xboole_0(k1_xboole_0,k1_xboole_0)), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (r1_tarski(X1,k2_tarski(X3,X2))|X1!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_40, plain, (X1=X3|X1=X2|~r1_tarski(k1_tarski(X1),k5_enumset1(X2,X2,X2,X2,X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_27]), c_0_28])).
cnf(c_0_41, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_43, plain, ![X31, X32]:((k4_xboole_0(X31,X32)!=k1_xboole_0|r1_tarski(X31,X32))&(~r1_tarski(X31,X32)|k4_xboole_0(X31,X32)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_44, plain, (r1_xboole_0(k1_xboole_0,k4_xboole_0(k1_xboole_0,X1))), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (r1_tarski(X1,k5_enumset1(X3,X3,X3,X3,X3,X3,X2))|X1!=k1_tarski(X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_27]), c_0_28])).
cnf(c_0_46, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|X1=esk1_0|X1=esk2_0|~r1_tarski(k1_tarski(X1),k1_tarski(esk4_0))), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_48, negated_conjecture, (esk1_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_49, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, plain, (X1=k1_xboole_0|~r1_xboole_0(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_51, plain, (r1_xboole_0(k4_xboole_0(k1_xboole_0,X1),k4_xboole_0(k1_xboole_0,X2))), inference(spm,[status(thm)],[c_0_37, c_0_44])).
fof(c_0_52, plain, ![X33, X34, X35]:(k1_tarski(X33)!=k2_tarski(X34,X35)|X34=X35), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_zfmisc_1])])).
cnf(c_0_53, plain, (r1_tarski(k1_tarski(X1),k5_enumset1(X2,X2,X2,X2,X2,X2,X1))), inference(er,[status(thm)],[c_0_45])).
cnf(c_0_54, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk4_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_55, plain, (r1_xboole_0(X1,X2)|~r1_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_56, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r1_xboole_0(X2,k1_xboole_0)|~r1_tarski(X1,X3)), inference(spm,[status(thm)],[c_0_30, c_0_49])).
cnf(c_0_57, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_58, plain, (X2=X3|k1_tarski(X1)!=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_59, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_60, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk4_0|r1_tarski(k1_tarski(esk2_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_62, plain, (r1_xboole_0(X1,k3_xboole_0(X2,X3))|~r1_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_55, c_0_23])).
cnf(c_0_63, plain, (r1_xboole_0(X1,k1_xboole_0)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_38])])).
cnf(c_0_64, plain, (X2=X3|k1_tarski(X1)!=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_27]), c_0_28])).
fof(c_0_65, plain, ![X38, X39]:r1_tarski(k1_tarski(X38),k2_tarski(X38,X39)), inference(variable_rename,[status(thm)],[t12_zfmisc_1])).
fof(c_0_66, plain, ![X36, X37]:(k3_xboole_0(k1_tarski(X36),k1_tarski(X37))!=k1_tarski(X36)|X36=X37), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_zfmisc_1])])).
cnf(c_0_67, plain, (k5_enumset1(X1,X1,X1,X1,X1,X1,X2)=k1_tarski(X2)|~r1_tarski(k5_enumset1(X1,X1,X1,X1,X1,X1,X2),k1_tarski(X2))), inference(spm,[status(thm)],[c_0_59, c_0_53])).
cnf(c_0_68, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk3_0|esk2_0=esk4_0), inference(spm,[status(thm)],[c_0_40, c_0_60])).
cnf(c_0_69, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_61, c_0_57])).
cnf(c_0_70, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(k3_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_50, c_0_62])).
cnf(c_0_71, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_63, c_0_47])).
cnf(c_0_72, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk1_0|k1_tarski(X1)!=k1_tarski(esk4_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_41]), c_0_64])).
cnf(c_0_73, plain, (r1_tarski(k1_tarski(X1),k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_74, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|k1_tarski(esk4_0)!=k1_tarski(esk3_0)), inference(ef,[status(thm)],[c_0_41])).
cnf(c_0_75, plain, (X1=X2|k3_xboole_0(k1_tarski(X1),k1_tarski(X2))!=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_76, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k1_tarski(esk2_0)=k1_xboole_0|esk2_0=esk4_0|esk2_0=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69])])).
cnf(c_0_77, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_78, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk1_0), inference(er,[status(thm)],[c_0_72])).
cnf(c_0_79, plain, (r1_tarski(X1,k2_tarski(X2,X3))|X1!=k2_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_80, plain, (r1_tarski(k1_tarski(X1),k5_enumset1(X1,X1,X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_27]), c_0_28])).
cnf(c_0_81, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|X1=esk1_0|X1=esk2_0|k1_tarski(esk4_0)!=k1_tarski(esk3_0)|~r1_tarski(k1_tarski(X1),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_40, c_0_74])).
cnf(c_0_82, negated_conjecture, (esk1_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_83, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|esk2_0=esk3_0|esk2_0=esk4_0|esk2_0=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])])).
cnf(c_0_84, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk1_0|X1=esk1_0|X1=esk2_0|~r1_tarski(k1_tarski(X1),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_40, c_0_78])).
cnf(c_0_85, plain, (r1_tarski(X1,k5_enumset1(X2,X2,X2,X2,X2,X2,X3))|X1!=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_27]), c_0_27]), c_0_28]), c_0_28])).
cnf(c_0_86, plain, (k5_enumset1(X1,X1,X1,X1,X1,X1,X2)=k1_tarski(X1)|~r1_tarski(k5_enumset1(X1,X1,X1,X1,X1,X1,X2),k1_tarski(X1))), inference(spm,[status(thm)],[c_0_59, c_0_80])).
cnf(c_0_87, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk3_0|k1_tarski(esk4_0)!=k1_tarski(esk3_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_80]), c_0_82])).
cnf(c_0_88, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|esk2_0=esk4_0|esk2_0=esk3_0|X1!=esk3_0), inference(ef,[status(thm)],[c_0_83])).
cnf(c_0_89, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk3_0|esk2_0=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_80]), c_0_82])).
cnf(c_0_90, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk4_0|esk2_0=esk1_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_53]), c_0_48])).
cnf(c_0_91, plain, (r1_tarski(X1,k2_tarski(X2,X3))|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_92, plain, (X1=X2|X1=X3|k1_tarski(X1)!=k5_enumset1(X2,X2,X2,X2,X2,X2,X3)), inference(spm,[status(thm)],[c_0_40, c_0_85])).
cnf(c_0_93, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|k1_xboole_0=k1_tarski(esk1_0)|esk2_0=esk3_0|k1_tarski(esk4_0)!=k1_tarski(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_69])])).
cnf(c_0_94, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_tarski(esk3_0)|esk2_0=esk3_0|esk2_0=esk4_0), inference(er,[status(thm)],[c_0_88])).
cnf(c_0_95, negated_conjecture, (esk2_0=esk3_0|esk2_0=esk1_0|k1_tarski(X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_89])).
cnf(c_0_96, negated_conjecture, (esk2_0=esk4_0|esk2_0=esk1_0|k1_tarski(X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_64, c_0_90])).
cnf(c_0_97, negated_conjecture, (k5_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk1_0,esk2_0)=k1_xboole_0|esk2_0=esk1_0|r1_tarski(k1_tarski(esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(spm,[status(thm)],[c_0_80, c_0_78])).
cnf(c_0_98, plain, (r1_tarski(X1,k5_enumset1(X2,X2,X2,X2,X2,X2,X3))|X1!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_27]), c_0_28])).
cnf(c_0_99, negated_conjecture, (k1_xboole_0=k1_tarski(esk1_0)|esk2_0=esk3_0|X1=esk2_0|X1=esk1_0|k1_tarski(esk4_0)!=k1_tarski(esk3_0)|k1_tarski(X1)!=k1_tarski(esk3_0)), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_100, negated_conjecture, (esk2_0=esk4_0|esk2_0=esk3_0|X1=esk2_0|X1=esk1_0|k1_tarski(X1)!=k1_tarski(esk3_0)), inference(spm,[status(thm)],[c_0_92, c_0_94])).
cnf(c_0_101, negated_conjecture, (esk2_0=esk1_0|esk2_0=esk3_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_89]), c_0_69])]), c_0_95])).
cnf(c_0_102, negated_conjecture, (esk2_0=esk1_0|esk2_0=esk4_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_90]), c_0_69])]), c_0_96])).
cnf(c_0_103, negated_conjecture, (esk2_0=esk1_0|r1_tarski(k1_tarski(esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_97]), c_0_69])]), c_0_98])).
cnf(c_0_104, negated_conjecture, (k1_xboole_0=k1_tarski(esk1_0)|esk2_0=esk3_0|k1_tarski(esk4_0)!=k1_tarski(esk3_0)), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_99]), c_0_82])).
cnf(c_0_105, negated_conjecture, (esk2_0=esk3_0|esk2_0=esk4_0), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_100]), c_0_82])).
cnf(c_0_106, negated_conjecture, (esk2_0=esk1_0|esk4_0=esk3_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_48])).
cnf(c_0_107, negated_conjecture, (r1_tarski(k1_tarski(esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk4_0))|k1_tarski(esk4_0)!=k1_tarski(esk3_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_82]), c_0_98])).
cnf(c_0_108, negated_conjecture, (esk4_0=esk3_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_82]), c_0_48])).
cnf(c_0_109, negated_conjecture, (r1_tarski(k1_tarski(esk1_0),k5_enumset1(esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_108]), c_0_108])])).
cnf(c_0_110, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_109]), c_0_82]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 111
# Proof object clause steps            : 82
# Proof object formula steps           : 29
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 49
# Proof object simplifying inferences  : 56
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 24
# Processed clauses                    : 2532
# ...of these trivial                  : 35
# ...subsumed                          : 1547
# ...remaining for further processing  : 950
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 376
# Backward-rewritten                   : 369
# Generated clauses                    : 31880
# ...of the previous two non-trivial   : 28525
# Contextual simplify-reflections      : 59
# Paramodulations                      : 31714
# Factorizations                       : 108
# Equation resolutions                 : 58
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 180
#    Positive orientable unit clauses  : 15
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 161
# Current number of unprocessed clauses: 24820
# ...number of literals in the above   : 149949
# Current number of archived formulas  : 0
# Current number of archived clauses   : 770
# Clause-clause subsumption calls (NU) : 44209
# Rec. Clause-clause subsumption calls : 16089
# Non-unit clause-clause subsumptions  : 1752
# Unit Clause-clause subsumption calls : 359
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1096418

# -------------------------------------------------
# User time                : 0.412 s
# System time              : 0.003 s
# Total time               : 0.415 s
# Maximum resident set size: 26208 pages
