# No SInE strategy applied
# Auto-Mode selected heuristic H_____042_B03_F1_AE_Q4_SP_S2S
# and selection function SelectNewComplexAHP.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t28_yellow_6, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(r1_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X3)))), file('data/MPTP/proofs/ed47f2ab7b614b2fb780490428b31fb2.E_input', t28_yellow_6)).
fof(d11_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r1_waybel_0(X1,X2,X3)<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&![X5]:(m1_subset_1(X5,u1_struct_0(X2))=>(r1_orders_2(X2,X4,X5)=>r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/ed47f2ab7b614b2fb780490428b31fb2.E_input', d11_waybel_0)).
fof(dt_l1_waybel_0, axiom, ![X1]:(l1_struct_0(X1)=>![X2]:(l1_waybel_0(X2,X1)=>l1_orders_2(X2))), file('data/MPTP/proofs/ed47f2ab7b614b2fb780490428b31fb2.E_input', dt_l1_waybel_0)).
fof(d5_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>(v7_waybel_0(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>?[X4]:(m1_subset_1(X4,u1_struct_0(X1))&(r1_orders_2(X1,X2,X4)&r1_orders_2(X1,X3,X4))))))), file('data/MPTP/proofs/ed47f2ab7b614b2fb780490428b31fb2.E_input', d5_yellow_6)).
fof(d12_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&l1_waybel_0(X2,X1))=>![X3]:(r2_waybel_0(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X2))=>?[X5]:(m1_subset_1(X5,u1_struct_0(X2))&(r1_orders_2(X2,X4,X5)&r2_hidden(k2_waybel_0(X1,X2,X5),X3))))))), file('data/MPTP/proofs/ed47f2ab7b614b2fb780490428b31fb2.E_input', d12_waybel_0)).
fof(c_0_5, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(r1_waybel_0(X1,X2,X3)=>r2_waybel_0(X1,X2,X3))))), inference(assume_negation,[status(cth)],[t28_yellow_6])).
fof(c_0_6, plain, ![X9, X10, X11, X13, X14, X15]:(((m1_subset_1(esk4_3(X9,X10,X11),u1_struct_0(X10))|~r1_waybel_0(X9,X10,X11)|(v2_struct_0(X10)|~l1_waybel_0(X10,X9))|(v2_struct_0(X9)|~l1_struct_0(X9)))&(~m1_subset_1(X13,u1_struct_0(X10))|(~r1_orders_2(X10,esk4_3(X9,X10,X11),X13)|r2_hidden(k2_waybel_0(X9,X10,X13),X11))|~r1_waybel_0(X9,X10,X11)|(v2_struct_0(X10)|~l1_waybel_0(X10,X9))|(v2_struct_0(X9)|~l1_struct_0(X9))))&((m1_subset_1(esk5_4(X9,X10,X14,X15),u1_struct_0(X10))|~m1_subset_1(X15,u1_struct_0(X10))|r1_waybel_0(X9,X10,X14)|(v2_struct_0(X10)|~l1_waybel_0(X10,X9))|(v2_struct_0(X9)|~l1_struct_0(X9)))&((r1_orders_2(X10,X15,esk5_4(X9,X10,X14,X15))|~m1_subset_1(X15,u1_struct_0(X10))|r1_waybel_0(X9,X10,X14)|(v2_struct_0(X10)|~l1_waybel_0(X10,X9))|(v2_struct_0(X9)|~l1_struct_0(X9)))&(~r2_hidden(k2_waybel_0(X9,X10,esk5_4(X9,X10,X14,X15)),X14)|~m1_subset_1(X15,u1_struct_0(X10))|r1_waybel_0(X9,X10,X14)|(v2_struct_0(X10)|~l1_waybel_0(X10,X9))|(v2_struct_0(X9)|~l1_struct_0(X9)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d11_waybel_0])])])])])])])).
fof(c_0_7, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_struct_0(esk1_0))&((~v2_struct_0(esk2_0)&(v4_orders_2(esk2_0)&(v7_waybel_0(esk2_0)&l1_waybel_0(esk2_0,esk1_0))))&(r1_waybel_0(esk1_0,esk2_0,esk3_0)&~r2_waybel_0(esk1_0,esk2_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_5])])])])).
fof(c_0_8, plain, ![X17, X18]:(~l1_struct_0(X17)|(~l1_waybel_0(X18,X17)|l1_orders_2(X18))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_waybel_0])])])).
fof(c_0_9, plain, ![X19, X20, X21, X25]:(((m1_subset_1(esk6_3(X19,X20,X21),u1_struct_0(X19))|~m1_subset_1(X21,u1_struct_0(X19))|~m1_subset_1(X20,u1_struct_0(X19))|~v7_waybel_0(X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&((r1_orders_2(X19,X20,esk6_3(X19,X20,X21))|~m1_subset_1(X21,u1_struct_0(X19))|~m1_subset_1(X20,u1_struct_0(X19))|~v7_waybel_0(X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&(r1_orders_2(X19,X21,esk6_3(X19,X20,X21))|~m1_subset_1(X21,u1_struct_0(X19))|~m1_subset_1(X20,u1_struct_0(X19))|~v7_waybel_0(X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))))&((m1_subset_1(esk7_1(X19),u1_struct_0(X19))|v7_waybel_0(X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&((m1_subset_1(esk8_1(X19),u1_struct_0(X19))|v7_waybel_0(X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))&(~m1_subset_1(X25,u1_struct_0(X19))|(~r1_orders_2(X19,esk7_1(X19),X25)|~r1_orders_2(X19,esk8_1(X19),X25))|v7_waybel_0(X19)|(v2_struct_0(X19)|~l1_orders_2(X19)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_yellow_6])])])])])])).
cnf(c_0_10, plain, (m1_subset_1(esk4_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_0(X1,X2,X3)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, negated_conjecture, (r1_waybel_0(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, negated_conjecture, (l1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_14, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_15, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, plain, (l1_orders_2(X2)|~l1_struct_0(X1)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_17, plain, ![X26, X27, X28, X29, X31, X33]:(((m1_subset_1(esk9_4(X26,X27,X28,X29),u1_struct_0(X27))|~m1_subset_1(X29,u1_struct_0(X27))|~r2_waybel_0(X26,X27,X28)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26)))&((r1_orders_2(X27,X29,esk9_4(X26,X27,X28,X29))|~m1_subset_1(X29,u1_struct_0(X27))|~r2_waybel_0(X26,X27,X28)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26)))&(r2_hidden(k2_waybel_0(X26,X27,esk9_4(X26,X27,X28,X29)),X28)|~m1_subset_1(X29,u1_struct_0(X27))|~r2_waybel_0(X26,X27,X28)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26)))))&((m1_subset_1(esk10_3(X26,X27,X31),u1_struct_0(X27))|r2_waybel_0(X26,X27,X31)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26)))&(~m1_subset_1(X33,u1_struct_0(X27))|(~r1_orders_2(X27,esk10_3(X26,X27,X31),X33)|~r2_hidden(k2_waybel_0(X26,X27,X33),X31))|r2_waybel_0(X26,X27,X31)|(v2_struct_0(X27)|~l1_waybel_0(X27,X26))|(v2_struct_0(X26)|~l1_struct_0(X26))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d12_waybel_0])])])])])])])).
cnf(c_0_18, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v7_waybel_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, negated_conjecture, (m1_subset_1(esk4_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])]), c_0_14]), c_0_15])).
cnf(c_0_20, negated_conjecture, (l1_orders_2(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_12]), c_0_13])])).
cnf(c_0_21, negated_conjecture, (v7_waybel_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, plain, (m1_subset_1(esk10_3(X1,X2,X3),u1_struct_0(X2))|r2_waybel_0(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~l1_waybel_0(X2,X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (m1_subset_1(esk6_3(esk2_0,X1,esk4_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])]), c_0_14])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk10_3(esk1_0,esk2_0,X1),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_12]), c_0_13])]), c_0_14]), c_0_15])).
cnf(c_0_25, plain, (r2_hidden(k2_waybel_0(X3,X2,X1),X4)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,esk4_3(X3,X2,X4),X1)|~r1_waybel_0(X3,X2,X4)|~l1_waybel_0(X2,X3)|~l1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_26, negated_conjecture, (m1_subset_1(esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk2_0))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, negated_conjecture, (r2_hidden(k2_waybel_0(X1,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X2),esk4_3(esk1_0,esk2_0,esk3_0))),X3)|r2_waybel_0(esk1_0,esk2_0,X2)|v2_struct_0(X1)|~r1_orders_2(esk2_0,esk4_3(X1,esk2_0,X3),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X2),esk4_3(esk1_0,esk2_0,esk3_0)))|~r1_waybel_0(X1,esk2_0,X3)|~l1_waybel_0(esk2_0,X1)|~l1_struct_0(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_14])).
cnf(c_0_28, plain, (r1_orders_2(X1,X2,esk6_3(X1,X3,X2))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v7_waybel_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_29, plain, (r2_waybel_0(X3,X2,X4)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,esk10_3(X3,X2,X4),X1)|~r2_hidden(k2_waybel_0(X3,X2,X1),X4)|~l1_waybel_0(X2,X3)|~l1_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, negated_conjecture, (r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),X2)|r2_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk4_3(esk1_0,esk2_0,X2),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))|~r1_waybel_0(esk1_0,esk2_0,X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_12]), c_0_13])]), c_0_15])).
cnf(c_0_31, negated_conjecture, (r1_orders_2(esk2_0,X1,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X2),X1))|r2_waybel_0(esk1_0,esk2_0,X2)|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_24]), c_0_20]), c_0_21])]), c_0_14])).
cnf(c_0_32, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,X1)|r2_waybel_0(X2,esk2_0,X3)|v2_struct_0(X2)|~r2_hidden(k2_waybel_0(X2,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),X3)|~r1_orders_2(esk2_0,esk10_3(X2,esk2_0,X3),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))|~l1_waybel_0(esk2_0,X2)|~l1_struct_0(X2)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_26]), c_0_14])).
cnf(c_0_33, negated_conjecture, (r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),esk3_0)|r2_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk4_3(esk1_0,esk2_0,esk3_0),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_30, c_0_11])).
cnf(c_0_34, negated_conjecture, (r1_orders_2(esk2_0,esk4_3(esk1_0,esk2_0,esk3_0),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_31, c_0_19])).
cnf(c_0_35, plain, (r1_orders_2(X1,X2,esk6_3(X1,X2,X3))|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v7_waybel_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_36, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,X1)|r2_waybel_0(esk1_0,esk2_0,X2)|~r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),X2)|~r1_orders_2(esk2_0,esk10_3(esk1_0,esk2_0,X2),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_12]), c_0_13])]), c_0_15])).
cnf(c_0_37, negated_conjecture, (r2_hidden(k2_waybel_0(esk1_0,esk2_0,esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0))),esk3_0)|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38, negated_conjecture, (~r2_waybel_0(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_39, negated_conjecture, (r1_orders_2(esk2_0,X1,esk6_3(esk2_0,X1,esk4_3(esk1_0,esk2_0,esk3_0)))|~m1_subset_1(X1,u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_19]), c_0_20]), c_0_21])]), c_0_14])).
cnf(c_0_40, negated_conjecture, (r2_waybel_0(esk1_0,esk2_0,X1)|~r1_orders_2(esk2_0,esk10_3(esk1_0,esk2_0,esk3_0),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_41, negated_conjecture, (r1_orders_2(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk6_3(esk2_0,esk10_3(esk1_0,esk2_0,X1),esk4_3(esk1_0,esk2_0,esk3_0)))|r2_waybel_0(esk1_0,esk2_0,X1)), inference(spm,[status(thm)],[c_0_39, c_0_24])).
cnf(c_0_42, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_38]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 43
# Proof object clause steps            : 32
# Proof object formula steps           : 11
# Proof object conjectures             : 27
# Proof object clause conjectures      : 24
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 1034
# ...of these trivial                  : 1
# ...subsumed                          : 0
# ...remaining for further processing  : 1033
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 57601
# ...of the previous two non-trivial   : 57076
# Contextual simplify-reflections      : 0
# Paramodulations                      : 57601
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1031
#    Positive orientable unit clauses  : 44
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 983
# Current number of unprocessed clauses: 56067
# ...number of literals in the above   : 217005
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 64906
# Rec. Clause-clause subsumption calls : 12183
# Non-unit clause-clause subsumptions  : 2
# Unit Clause-clause subsumption calls : 1322
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 701
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3574786

# -------------------------------------------------
# User time                : 0.575 s
# System time              : 0.020 s
# Total time               : 0.595 s
# Maximum resident set size: 52428 pages
