# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t20_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', t20_tex_2)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', t69_enumset1)).
fof(t7_tex_2, axiom, ![X1]:((~(v1_xboole_0(X1))&~(v1_zfmisc_1(X1)))=>![X2]:(m1_subset_1(X2,X1)=>v1_subset_1(k6_domain_1(X1,X2),X1))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', t7_tex_2)).
fof(cc1_zfmisc_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_zfmisc_1(X1)), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', cc1_zfmisc_1)).
fof(dt_k1_tex_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v1_pre_topc(k1_tex_2(X1,X2))&m1_pre_topc(k1_tex_2(X1,X2),X1)))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', dt_k1_tex_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', dt_l1_pre_topc)).
fof(d3_tex_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_tex_2(X2,X1)<=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>v1_subset_1(X3,u1_struct_0(X1))))))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', d3_tex_2)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', dt_m1_pre_topc)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', d7_subset_1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', redefinition_k6_domain_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', fc2_struct_0)).
fof(fc2_tex_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v7_struct_0(k1_tex_2(X1,X2))&v1_pre_topc(k1_tex_2(X1,X2))))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', fc2_tex_2)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', d1_xboole_0)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', d2_subset_1)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', dt_k6_domain_1)).
fof(t8_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~((v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))&v7_struct_0(X1))))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', t8_tex_2)).
fof(t6_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>~((v1_subset_1(k6_domain_1(X1,X2),X1)&v1_zfmisc_1(X1))))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', t6_tex_2)).
fof(l3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(r2_hidden(X3,X2)=>r2_hidden(X3,X1))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', l3_subset_1)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', t1_tsep_1)).
fof(dt_k2_subset_1, axiom, ![X1]:m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1)), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', dt_k2_subset_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', d4_subset_1)).
fof(fc14_subset_1, axiom, ![X1]:~(v1_subset_1(k2_subset_1(X1),X1)), file('data/MPTP/proofs/ed94e21abb974af8bdeaa1851943de6f.E_input', fc14_subset_1)).
fof(c_0_23, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1)))))), inference(assume_negation,[status(cth)],[t20_tex_2])).
fof(c_0_24, plain, ![X6, X7, X8, X9, X10, X11]:(((~r2_hidden(X8,X7)|X8=X6|X7!=k1_tarski(X6))&(X9!=X6|r2_hidden(X9,X7)|X7!=k1_tarski(X6)))&((~r2_hidden(esk3_2(X10,X11),X11)|esk3_2(X10,X11)!=X10|X11=k1_tarski(X10))&(r2_hidden(esk3_2(X10,X11),X11)|esk3_2(X10,X11)=X10|X11=k1_tarski(X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_25, plain, ![X13]:k2_tarski(X13,X13)=k1_tarski(X13), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_26, plain, ![X14, X15]:(v1_xboole_0(X14)|v1_zfmisc_1(X14)|(~m1_subset_1(X15,X14)|v1_subset_1(k6_domain_1(X14,X15),X14))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_tex_2])])])])).
fof(c_0_27, plain, ![X16]:(~v1_xboole_0(X16)|v1_zfmisc_1(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_zfmisc_1])])).
fof(c_0_28, plain, ![X17, X18]:((~v2_struct_0(k1_tex_2(X17,X18))|(v2_struct_0(X17)|~l1_pre_topc(X17)|~m1_subset_1(X18,u1_struct_0(X17))))&((v1_pre_topc(k1_tex_2(X17,X18))|(v2_struct_0(X17)|~l1_pre_topc(X17)|~m1_subset_1(X18,u1_struct_0(X17))))&(m1_pre_topc(k1_tex_2(X17,X18),X17)|(v2_struct_0(X17)|~l1_pre_topc(X17)|~m1_subset_1(X18,u1_struct_0(X17)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tex_2])])])])).
fof(c_0_29, negated_conjecture, ((~v2_struct_0(esk1_0)&l1_pre_topc(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&((~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0)))&(v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_23])])])])).
cnf(c_0_30, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_32, plain, ![X19]:(~l1_pre_topc(X19)|l1_struct_0(X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_33, plain, (v1_xboole_0(X1)|v1_zfmisc_1(X1)|v1_subset_1(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (v1_zfmisc_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_35, plain, ![X20, X21, X22]:((~v1_tex_2(X21,X20)|(~m1_subset_1(X22,k1_zfmisc_1(u1_struct_0(X20)))|(X22!=u1_struct_0(X21)|v1_subset_1(X22,u1_struct_0(X20))))|~m1_pre_topc(X21,X20)|~l1_pre_topc(X20))&((m1_subset_1(esk4_2(X20,X21),k1_zfmisc_1(u1_struct_0(X20)))|v1_tex_2(X21,X20)|~m1_pre_topc(X21,X20)|~l1_pre_topc(X20))&((esk4_2(X20,X21)=u1_struct_0(X21)|v1_tex_2(X21,X20)|~m1_pre_topc(X21,X20)|~l1_pre_topc(X20))&(~v1_subset_1(esk4_2(X20,X21),u1_struct_0(X20))|v1_tex_2(X21,X20)|~m1_pre_topc(X21,X20)|~l1_pre_topc(X20))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tex_2])])])])])).
fof(c_0_36, plain, ![X24, X25]:(~l1_pre_topc(X24)|(~m1_pre_topc(X25,X24)|l1_pre_topc(X25))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_37, plain, (m1_pre_topc(k1_tex_2(X1,X2),X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_40, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_41, plain, ![X26, X27]:((~v1_subset_1(X27,X26)|X27!=X26|~m1_subset_1(X27,k1_zfmisc_1(X26)))&(X27=X26|v1_subset_1(X27,X26)|~m1_subset_1(X27,k1_zfmisc_1(X26)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_42, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_30, c_0_31])).
fof(c_0_43, plain, ![X28, X29]:(v1_xboole_0(X28)|~m1_subset_1(X29,X28)|k6_domain_1(X28,X29)=k1_tarski(X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
fof(c_0_44, plain, ![X30]:(v2_struct_0(X30)|~l1_struct_0(X30)|~v1_xboole_0(u1_struct_0(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_45, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_46, negated_conjecture, (~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|~v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_47, plain, (v1_zfmisc_1(X1)|v1_subset_1(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_48, plain, (esk4_2(X1,X2)=u1_struct_0(X2)|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_49, plain, ![X31, X32]:((~v2_struct_0(k1_tex_2(X31,X32))|(v2_struct_0(X31)|~l1_pre_topc(X31)|~m1_subset_1(X32,u1_struct_0(X31))))&((v7_struct_0(k1_tex_2(X31,X32))|(v2_struct_0(X31)|~l1_pre_topc(X31)|~m1_subset_1(X32,u1_struct_0(X31))))&(v1_pre_topc(k1_tex_2(X31,X32))|(v2_struct_0(X31)|~l1_pre_topc(X31)|~m1_subset_1(X32,u1_struct_0(X31)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_tex_2])])])])).
cnf(c_0_50, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_51, negated_conjecture, (m1_pre_topc(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])]), c_0_40])).
cnf(c_0_52, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_54, plain, (v1_tex_2(X2,X1)|~v1_subset_1(esk4_2(X1,X2),u1_struct_0(X1))|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_55, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_56, plain, ![X33, X34, X35]:((~v1_xboole_0(X33)|~r2_hidden(X34,X33))&(r2_hidden(esk5_1(X35),X35)|v1_xboole_0(X35))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_57, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X1)), inference(er,[status(thm)],[c_0_42])).
fof(c_0_58, plain, ![X37, X38]:(((~m1_subset_1(X38,X37)|r2_hidden(X38,X37)|v1_xboole_0(X37))&(~r2_hidden(X38,X37)|m1_subset_1(X38,X37)|v1_xboole_0(X37)))&((~m1_subset_1(X38,X37)|v1_xboole_0(X38)|~v1_xboole_0(X37))&(~v1_xboole_0(X38)|m1_subset_1(X38,X37)|~v1_xboole_0(X37)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_59, plain, ![X39, X40]:(v1_xboole_0(X39)|~m1_subset_1(X40,X39)|m1_subset_1(k6_domain_1(X39,X40),k1_zfmisc_1(X39))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_60, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_61, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_62, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_45, c_0_39])).
fof(c_0_63, plain, ![X41, X42]:(v2_struct_0(X41)|~l1_struct_0(X41)|(~m1_subset_1(X42,u1_struct_0(X41))|(~v1_subset_1(k6_domain_1(u1_struct_0(X41),X42),u1_struct_0(X41))|~v7_struct_0(X41)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_tex_2])])])])).
cnf(c_0_64, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk1_0))|~v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_38])])).
cnf(c_0_65, negated_conjecture, (u1_struct_0(X1)=esk4_2(esk1_0,X1)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_48, c_0_39])).
cnf(c_0_66, plain, (v7_struct_0(k1_tex_2(X1,X2))|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_67, negated_conjecture, (l1_pre_topc(k1_tex_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_39])])).
cnf(c_0_68, plain, (esk4_2(X1,X2)=u1_struct_0(X1)|v1_tex_2(X2,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_69, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_55, c_0_31])).
cnf(c_0_70, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_71, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_72, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_73, plain, ![X43, X44]:(v1_xboole_0(X43)|(~m1_subset_1(X44,X43)|(~v1_subset_1(k6_domain_1(X43,X44),X43)|~v1_zfmisc_1(X43)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_tex_2])])])])).
cnf(c_0_74, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_75, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_60, c_0_31])).
cnf(c_0_76, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_61]), c_0_62])])).
cnf(c_0_77, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))|~v7_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_78, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,esk2_0))=esk4_2(esk1_0,k1_tex_2(esk1_0,esk2_0))|v1_zfmisc_1(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_51])])).
cnf(c_0_79, negated_conjecture, (v7_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_38]), c_0_39])]), c_0_40])).
cnf(c_0_80, negated_conjecture, (l1_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_45, c_0_67])).
cnf(c_0_81, negated_conjecture, (esk4_2(esk1_0,X1)=u1_struct_0(esk1_0)|v1_tex_2(X1,esk1_0)|~m1_pre_topc(X1,esk1_0)), inference(spm,[status(thm)],[c_0_68, c_0_39])).
cnf(c_0_82, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_69])).
cnf(c_0_83, plain, (r2_hidden(esk5_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_84, plain, (~v1_xboole_0(k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_85, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_72, c_0_70])).
cnf(c_0_86, plain, (v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~v1_subset_1(k6_domain_1(X1,X2),X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_87, plain, (k6_domain_1(X1,X2)=X1|v1_xboole_0(X1)|v1_subset_1(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_52, c_0_74])).
fof(c_0_88, plain, ![X45, X46, X47]:(~m1_subset_1(X46,k1_zfmisc_1(X45))|(~r2_hidden(X47,X46)|r2_hidden(X47,X45))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l3_subset_1])])])).
fof(c_0_89, plain, ![X48, X49]:(~l1_pre_topc(X48)|(~m1_pre_topc(X49,X48)|m1_subset_1(u1_struct_0(X49),k1_zfmisc_1(u1_struct_0(X48))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
cnf(c_0_90, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_38]), c_0_76])).
cnf(c_0_91, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk1_0))|v2_struct_0(k1_tex_2(esk1_0,esk2_0))|~v1_subset_1(k6_domain_1(esk4_2(esk1_0,k1_tex_2(esk1_0,esk2_0)),X1),esk4_2(esk1_0,k1_tex_2(esk1_0,esk2_0)))|~m1_subset_1(X1,esk4_2(esk1_0,k1_tex_2(esk1_0,esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79]), c_0_80])])).
cnf(c_0_92, negated_conjecture, (esk4_2(esk1_0,k1_tex_2(esk1_0,esk2_0))=u1_struct_0(esk1_0)|v1_zfmisc_1(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_81]), c_0_51])])).
fof(c_0_93, plain, ![X50]:m1_subset_1(k2_subset_1(X50),k1_zfmisc_1(X50)), inference(variable_rename,[status(thm)],[dt_k2_subset_1])).
fof(c_0_94, plain, ![X51]:k2_subset_1(X51)=X51, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_95, plain, ![X52]:~v1_subset_1(k2_subset_1(X52),X52), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc14_subset_1])])).
cnf(c_0_96, plain, (esk5_1(k2_tarski(X1,X1))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84])).
cnf(c_0_97, plain, (k2_tarski(X1,X1)=k6_domain_1(X2,X1)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_85]), c_0_70])).
cnf(c_0_98, plain, (k6_domain_1(X1,X2)=X1|v1_xboole_0(X1)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_99, plain, (r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_100, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_101, negated_conjecture, (m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_90]), c_0_38])]), c_0_76])).
cnf(c_0_102, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk1_0))|v2_struct_0(k1_tex_2(esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_47])).
cnf(c_0_103, plain, (m1_subset_1(k2_subset_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_104, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_105, plain, (~v1_subset_1(k2_subset_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_106, plain, (esk5_1(k6_domain_1(X1,X2))=X2|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_107, plain, (k6_domain_1(X1,X2)=X1|~v1_zfmisc_1(X1)|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_85]), c_0_70])).
cnf(c_0_108, plain, (r2_hidden(X1,u1_struct_0(X2))|~m1_pre_topc(X3,X2)|~r2_hidden(X1,u1_struct_0(X3))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_109, negated_conjecture, (~v1_zfmisc_1(u1_struct_0(esk1_0))|~v1_subset_1(k2_tarski(esk2_0,esk2_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_90]), c_0_38])]), c_0_76])).
cnf(c_0_110, negated_conjecture, (u1_struct_0(esk1_0)=k2_tarski(esk2_0,esk2_0)|v1_subset_1(k2_tarski(esk2_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_52, c_0_101])).
cnf(c_0_111, plain, (v2_struct_0(X1)|~v2_struct_0(k1_tex_2(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_112, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk1_0))|v2_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_102, c_0_38])).
cnf(c_0_113, plain, (v1_subset_1(X3,u1_struct_0(X2))|~v1_tex_2(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|X3!=u1_struct_0(X1)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_114, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_115, plain, (~v1_subset_1(X1,X1)), inference(rw,[status(thm)],[c_0_105, c_0_104])).
cnf(c_0_116, negated_conjecture, (v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)|v1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_117, plain, (esk5_1(X1)=X2|~v1_zfmisc_1(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_118, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk5_1(u1_struct_0(X1)),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_108, c_0_83])).
cnf(c_0_119, negated_conjecture, (u1_struct_0(esk1_0)=k2_tarski(esk2_0,esk2_0)|~v1_zfmisc_1(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_120, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_38]), c_0_39])]), c_0_40])).
cnf(c_0_121, plain, (v1_zfmisc_1(X1)|v1_subset_1(k2_tarski(X2,X2),X1)|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_97]), c_0_85])).
cnf(c_0_122, plain, (v2_struct_0(X1)|~l1_struct_0(k1_tex_2(X1,X2))|~v1_xboole_0(u1_struct_0(k1_tex_2(X1,X2)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_111, c_0_61])).
cnf(c_0_123, plain, (u1_struct_0(X1)!=u1_struct_0(X2)|~m1_pre_topc(X2,X1)|~v1_tex_2(X2,X1)|~l1_pre_topc(X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_115])).
cnf(c_0_124, negated_conjecture, (v1_subset_1(k2_tarski(esk2_0,esk2_0),u1_struct_0(esk1_0))|v1_tex_2(k1_tex_2(esk1_0,esk2_0),esk1_0)), inference(rw,[status(thm)],[c_0_116, c_0_90])).
cnf(c_0_125, plain, (k6_domain_1(u1_struct_0(X1),X2)=u1_struct_0(X1)|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_87]), c_0_61])).
cnf(c_0_126, plain, (esk5_1(u1_struct_0(X1))=esk5_1(u1_struct_0(X2))|v1_xboole_0(u1_struct_0(X2))|~m1_pre_topc(X2,X1)|~v1_zfmisc_1(u1_struct_0(X1))|~l1_pre_topc(X1)), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_127, negated_conjecture, (u1_struct_0(esk1_0)=k2_tarski(esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_128, plain, (v1_zfmisc_1(k2_tarski(X1,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_121]), c_0_71])])).
cnf(c_0_129, negated_conjecture, (~v1_xboole_0(u1_struct_0(k1_tex_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_38]), c_0_80]), c_0_39])]), c_0_40])).
cnf(c_0_130, negated_conjecture, (v1_subset_1(k2_tarski(esk2_0,esk2_0),u1_struct_0(esk1_0))|u1_struct_0(k1_tex_2(esk1_0,esk2_0))!=u1_struct_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_51]), c_0_39])])).
cnf(c_0_131, plain, (k6_domain_1(u1_struct_0(X1),X2)=u1_struct_0(X1)|v2_struct_0(X1)|~v7_struct_0(X1)|~l1_struct_0(X1)|~r2_hidden(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_125, c_0_85])).
cnf(c_0_132, negated_conjecture, (esk5_1(u1_struct_0(k1_tex_2(esk1_0,esk2_0)))=esk2_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_51]), c_0_127]), c_0_96]), c_0_127]), c_0_128]), c_0_39])]), c_0_129])).
cnf(c_0_133, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,esk2_0))!=u1_struct_0(esk1_0)|~v1_zfmisc_1(u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_109, c_0_130])).
cnf(c_0_134, plain, (k2_tarski(X1,X1)=u1_struct_0(X2)|v2_struct_0(X2)|~v7_struct_0(X2)|~l1_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_97, c_0_131])).
cnf(c_0_135, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(k1_tex_2(esk1_0,esk2_0)))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_132]), c_0_129])).
cnf(c_0_136, negated_conjecture, (u1_struct_0(k1_tex_2(esk1_0,esk2_0))!=k2_tarski(esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_120])]), c_0_127])).
cnf(c_0_137, negated_conjecture, (v2_struct_0(k1_tex_2(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135]), c_0_79]), c_0_80])]), c_0_136])).
cnf(c_0_138, negated_conjecture, (m1_subset_1(esk2_0,k2_tarski(esk2_0,esk2_0))), inference(rw,[status(thm)],[c_0_38, c_0_127])).
cnf(c_0_139, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_137]), c_0_127]), c_0_138]), c_0_39])]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 140
# Proof object clause steps            : 93
# Proof object formula steps           : 47
# Proof object conjectures             : 39
# Proof object clause conjectures      : 36
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 48
# Proof object simplifying inferences  : 79
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 40
# Processed clauses                    : 3624
# ...of these trivial                  : 24
# ...subsumed                          : 2292
# ...remaining for further processing  : 1308
# Other redundant clauses eliminated   : 141
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 167
# Backward-rewritten                   : 193
# Generated clauses                    : 42073
# ...of the previous two non-trivial   : 39504
# Contextual simplify-reflections      : 138
# Paramodulations                      : 41875
# Factorizations                       : 8
# Equation resolutions                 : 187
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 908
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 19
#    Non-unit-clauses                  : 862
# Current number of unprocessed clauses: 35278
# ...number of literals in the above   : 205345
# Current number of archived formulas  : 0
# Current number of archived clauses   : 399
# Clause-clause subsumption calls (NU) : 215410
# Rec. Clause-clause subsumption calls : 67431
# Non-unit clause-clause subsumptions  : 1958
# Unit Clause-clause subsumption calls : 2300
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 20
# BW rewrite match successes           : 16
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 842060

# -------------------------------------------------
# User time                : 0.617 s
# System time              : 0.010 s
# Total time               : 0.627 s
# Maximum resident set size: 31868 pages
