# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_TT_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d21_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(X3=k15_lattice3(X1,X2)<=>(r4_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r4_lattice3(X1,X4,X2)=>r1_lattices(X1,X3,X4)))))))), file('data/MPTP/proofs/edd2da3eaf7e4893aa1496dcac9aebc0.E_input', d21_lattice3)).
fof(d16_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r3_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X2,X4)))))), file('data/MPTP/proofs/edd2da3eaf7e4893aa1496dcac9aebc0.E_input', d16_lattice3)).
fof(fraenkel_a_2_1_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&l3_lattices(X2))=>(r2_hidden(X1,a_2_1_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r3_lattice3(X2,X4,X3))))), file('data/MPTP/proofs/edd2da3eaf7e4893aa1496dcac9aebc0.E_input', fraenkel_a_2_1_lattice3)).
fof(d17_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(r4_lattice3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)=>r1_lattices(X1,X4,X2)))))), file('data/MPTP/proofs/edd2da3eaf7e4893aa1496dcac9aebc0.E_input', d17_lattice3)).
fof(dt_k15_lattice3, axiom, ![X1, X2]:((~(v2_struct_0(X1))&l3_lattices(X1))=>m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))), file('data/MPTP/proofs/edd2da3eaf7e4893aa1496dcac9aebc0.E_input', dt_k15_lattice3)).
fof(fraenkel_a_2_2_lattice3, axiom, ![X1, X2, X3]:((~(v2_struct_0(X2))&(v10_lattices(X2)&(v4_lattice3(X2)&l3_lattices(X2))))=>(r2_hidden(X1,a_2_2_lattice3(X2,X3))<=>?[X4]:(m1_subset_1(X4,u1_struct_0(X2))&(X1=X4&r4_lattice3(X2,X4,X3))))), file('data/MPTP/proofs/edd2da3eaf7e4893aa1496dcac9aebc0.E_input', fraenkel_a_2_2_lattice3)).
fof(t37_lattice3, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2))), file('data/MPTP/proofs/edd2da3eaf7e4893aa1496dcac9aebc0.E_input', t37_lattice3)).
fof(d22_lattice3, axiom, ![X1]:((~(v2_struct_0(X1))&l3_lattices(X1))=>![X2]:k16_lattice3(X1,X2)=k15_lattice3(X1,a_2_1_lattice3(X1,X2))), file('data/MPTP/proofs/edd2da3eaf7e4893aa1496dcac9aebc0.E_input', d22_lattice3)).
fof(c_0_8, plain, ![X5, X6, X7, X8]:(((r4_lattice3(X5,X7,X6)|X7!=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5)))&(~m1_subset_1(X8,u1_struct_0(X5))|(~r4_lattice3(X5,X8,X6)|r1_lattices(X5,X7,X8))|X7!=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5))))&((m1_subset_1(esk1_3(X5,X6,X7),u1_struct_0(X5))|~r4_lattice3(X5,X7,X6)|X7=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5)))&((r4_lattice3(X5,esk1_3(X5,X6,X7),X6)|~r4_lattice3(X5,X7,X6)|X7=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5)))&(~r1_lattices(X5,X7,esk1_3(X5,X6,X7))|~r4_lattice3(X5,X7,X6)|X7=k15_lattice3(X5,X6)|~m1_subset_1(X7,u1_struct_0(X5))|(v2_struct_0(X5)|(~v10_lattices(X5)|(~v4_lattice3(X5)|~l3_lattices(X5))))|(v2_struct_0(X5)|~l3_lattices(X5)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d21_lattice3])])])])])])).
fof(c_0_9, plain, ![X10, X11, X12, X13, X14]:((~r3_lattice3(X10,X11,X12)|(~m1_subset_1(X13,u1_struct_0(X10))|(~r2_hidden(X13,X12)|r1_lattices(X10,X11,X13)))|~m1_subset_1(X11,u1_struct_0(X10))|(v2_struct_0(X10)|~l3_lattices(X10)))&((m1_subset_1(esk2_3(X10,X11,X14),u1_struct_0(X10))|r3_lattice3(X10,X11,X14)|~m1_subset_1(X11,u1_struct_0(X10))|(v2_struct_0(X10)|~l3_lattices(X10)))&((r2_hidden(esk2_3(X10,X11,X14),X14)|r3_lattice3(X10,X11,X14)|~m1_subset_1(X11,u1_struct_0(X10))|(v2_struct_0(X10)|~l3_lattices(X10)))&(~r1_lattices(X10,X11,esk2_3(X10,X11,X14))|r3_lattice3(X10,X11,X14)|~m1_subset_1(X11,u1_struct_0(X10))|(v2_struct_0(X10)|~l3_lattices(X10)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_lattice3])])])])])])])).
fof(c_0_10, plain, ![X16, X17, X18, X20, X21]:(((m1_subset_1(esk3_3(X16,X17,X18),u1_struct_0(X17))|~r2_hidden(X16,a_2_1_lattice3(X17,X18))|(v2_struct_0(X17)|~l3_lattices(X17)))&((X16=esk3_3(X16,X17,X18)|~r2_hidden(X16,a_2_1_lattice3(X17,X18))|(v2_struct_0(X17)|~l3_lattices(X17)))&(r3_lattice3(X17,esk3_3(X16,X17,X18),X18)|~r2_hidden(X16,a_2_1_lattice3(X17,X18))|(v2_struct_0(X17)|~l3_lattices(X17)))))&(~m1_subset_1(X21,u1_struct_0(X17))|(X16!=X21|~r3_lattice3(X17,X21,X20))|r2_hidden(X16,a_2_1_lattice3(X17,X20))|(v2_struct_0(X17)|~l3_lattices(X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_1_lattice3])])])])])])])).
fof(c_0_11, plain, ![X22, X23, X24, X25, X26]:((~r4_lattice3(X22,X23,X24)|(~m1_subset_1(X25,u1_struct_0(X22))|(~r2_hidden(X25,X24)|r1_lattices(X22,X25,X23)))|~m1_subset_1(X23,u1_struct_0(X22))|(v2_struct_0(X22)|~l3_lattices(X22)))&((m1_subset_1(esk4_3(X22,X23,X26),u1_struct_0(X22))|r4_lattice3(X22,X23,X26)|~m1_subset_1(X23,u1_struct_0(X22))|(v2_struct_0(X22)|~l3_lattices(X22)))&((r2_hidden(esk4_3(X22,X23,X26),X26)|r4_lattice3(X22,X23,X26)|~m1_subset_1(X23,u1_struct_0(X22))|(v2_struct_0(X22)|~l3_lattices(X22)))&(~r1_lattices(X22,esk4_3(X22,X23,X26),X23)|r4_lattice3(X22,X23,X26)|~m1_subset_1(X23,u1_struct_0(X22))|(v2_struct_0(X22)|~l3_lattices(X22)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_lattice3])])])])])])])).
cnf(c_0_12, plain, (r4_lattice3(X1,esk1_3(X1,X2,X3),X2)|X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (r1_lattices(X1,X2,X4)|v2_struct_0(X1)|~r3_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (r3_lattice3(X1,esk3_3(X2,X1,X3),X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_lattice3(X1,X3))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (m1_subset_1(esk3_3(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_lattice3(X2,X3))|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (X1=esk3_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_lattice3(X2,X3))|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (r2_hidden(esk4_3(X1,X2,X3),X3)|r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (r1_lattices(X2,X4,X1)|v2_struct_0(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r4_lattice3(X2,X1,X3)|X4!=k15_lattice3(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
fof(c_0_20, plain, ![X28, X29]:(v2_struct_0(X28)|~l3_lattices(X28)|m1_subset_1(k15_lattice3(X28,X29),u1_struct_0(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k15_lattice3])])])).
fof(c_0_21, plain, ![X30, X31, X32, X34, X35]:(((m1_subset_1(esk5_3(X30,X31,X32),u1_struct_0(X31))|~r2_hidden(X30,a_2_2_lattice3(X31,X32))|(v2_struct_0(X31)|(~v10_lattices(X31)|(~v4_lattice3(X31)|~l3_lattices(X31)))))&((X30=esk5_3(X30,X31,X32)|~r2_hidden(X30,a_2_2_lattice3(X31,X32))|(v2_struct_0(X31)|(~v10_lattices(X31)|(~v4_lattice3(X31)|~l3_lattices(X31)))))&(r4_lattice3(X31,esk5_3(X30,X31,X32),X32)|~r2_hidden(X30,a_2_2_lattice3(X31,X32))|(v2_struct_0(X31)|(~v10_lattices(X31)|(~v4_lattice3(X31)|~l3_lattices(X31)))))))&(~m1_subset_1(X35,u1_struct_0(X31))|(X30!=X35|~r4_lattice3(X31,X35,X34))|r2_hidden(X30,a_2_2_lattice3(X31,X34))|(v2_struct_0(X31)|(~v10_lattices(X31)|(~v4_lattice3(X31)|~l3_lattices(X31)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_2_2_lattice3])])])])])])])).
cnf(c_0_22, plain, (X2=k15_lattice3(X1,X3)|v2_struct_0(X1)|v2_struct_0(X1)|~r1_lattices(X1,X2,esk1_3(X1,X3,X2))|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_23, plain, (r1_lattices(X1,X4,X2)|v2_struct_0(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_24, plain, (X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|r4_lattice3(X1,esk1_3(X1,X2,X3),X2)|~l3_lattices(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_12])).
cnf(c_0_25, plain, (X3=k15_lattice3(X1,X2)|v2_struct_0(X1)|m1_subset_1(esk1_3(X1,X2,X3),u1_struct_0(X1))|~l3_lattices(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_13])).
cnf(c_0_26, plain, (r1_lattices(X1,esk3_3(X2,X1,X3),X4)|v2_struct_0(X1)|~r2_hidden(X2,a_2_1_lattice3(X1,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X4,u1_struct_0(X1))|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_14, c_0_15]), c_0_16])).
cnf(c_0_27, plain, (esk3_3(esk4_3(X1,X2,a_2_1_lattice3(X3,X4)),X3,X4)=esk4_3(X1,X2,a_2_1_lattice3(X3,X4))|r4_lattice3(X1,X2,a_2_1_lattice3(X3,X4))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X3)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_28, plain, (v2_struct_0(X2)|r1_lattices(X2,X4,X1)|X4!=k15_lattice3(X2,X3)|~l3_lattices(X2)|~v10_lattices(X2)|~v4_lattice3(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(cn,[status(thm)],[c_0_19])).
cnf(c_0_29, plain, (v2_struct_0(X1)|m1_subset_1(k15_lattice3(X1,X2),u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_30, plain, (X1=esk5_3(X1,X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,a_2_2_lattice3(X2,X3))|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_31, plain, (r2_hidden(esk2_3(X1,X2,X3),X3)|r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_32, plain, (X2=k15_lattice3(X1,X3)|v2_struct_0(X1)|~l3_lattices(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~r4_lattice3(X1,X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_lattices(X1,X2,esk1_3(X1,X3,X2))), inference(cn,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (X1=k15_lattice3(X2,X3)|r1_lattices(X2,X4,esk1_3(X2,X3,X1))|v2_struct_0(X2)|~r2_hidden(X4,X3)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_34, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,esk4_3(X1,X2,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_35, plain, (r1_lattices(X1,esk4_3(X2,X3,a_2_1_lattice3(X1,X4)),X5)|r4_lattice3(X2,X3,a_2_1_lattice3(X1,X4))|v2_struct_0(X2)|v2_struct_0(X1)|~r2_hidden(X5,X4)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X2))|~l3_lattices(X1)|~l3_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_18])).
cnf(c_0_36, plain, (r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattices(X1,X2,esk2_3(X1,X2,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_37, plain, (r1_lattices(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_28]), c_0_29])).
cnf(c_0_38, plain, (r4_lattice3(X1,esk5_3(X2,X1,X3),X3)|v2_struct_0(X1)|~r2_hidden(X2,a_2_2_lattice3(X1,X3))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_39, plain, (esk5_3(esk2_3(X1,X2,a_2_2_lattice3(X3,X4)),X3,X4)=esk2_3(X1,X2,a_2_2_lattice3(X3,X4))|r3_lattice3(X1,X2,a_2_2_lattice3(X3,X4))|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_lattice3(X3)|~v10_lattices(X3)|~l3_lattices(X3)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40, plain, (X1=k15_lattice3(X2,X3)|v2_struct_0(X2)|~r2_hidden(X1,X3)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41, plain, (r4_lattice3(X1,X2,a_2_1_lattice3(X1,X3))|v2_struct_0(X1)|~r2_hidden(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_42, plain, (r2_hidden(X3,a_2_1_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r3_lattice3(X2,X1,X4)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_43, plain, (r3_lattice3(X1,k15_lattice3(X1,X2),X3)|v2_struct_0(X1)|~r4_lattice3(X1,esk2_3(X1,k15_lattice3(X1,X2),X3),X2)|~m1_subset_1(esk2_3(X1,k15_lattice3(X1,X2),X3),u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_29])).
cnf(c_0_44, plain, (r3_lattice3(X1,X2,a_2_2_lattice3(X3,X4))|r4_lattice3(X3,esk2_3(X1,X2,a_2_2_lattice3(X3,X4)),X4)|v2_struct_0(X1)|v2_struct_0(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_lattice3(X3)|~v10_lattices(X3)|~l3_lattices(X3)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_31])).
cnf(c_0_45, plain, (X1=k15_lattice3(X2,a_2_1_lattice3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X1,a_2_1_lattice3(X2,X3))|~r2_hidden(X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, plain, (r2_hidden(X1,a_2_1_lattice3(X2,X3))|v2_struct_0(X2)|~r3_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~l3_lattices(X2)), inference(er,[status(thm)],[c_0_42])).
cnf(c_0_47, plain, (r3_lattice3(X1,k15_lattice3(X1,X2),a_2_2_lattice3(X1,X2))|v2_struct_0(X1)|~m1_subset_1(esk2_3(X1,k15_lattice3(X1,X2),a_2_2_lattice3(X1,X2)),u1_struct_0(X1))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_29])).
cnf(c_0_48, plain, (m1_subset_1(esk2_3(X1,X2,X3),u1_struct_0(X1))|r3_lattice3(X1,X2,X3)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_49, plain, (r4_lattice3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X1)|X2!=k15_lattice3(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~v10_lattices(X1)|~v4_lattice3(X1)|~l3_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_50, plain, (X1=k15_lattice3(X2,a_2_1_lattice3(X2,X3))|v2_struct_0(X2)|~r2_hidden(X1,X3)|~r3_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_51, plain, (r3_lattice3(X1,k15_lattice3(X1,X2),a_2_2_lattice3(X1,X2))|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48]), c_0_29])).
cnf(c_0_52, plain, (r2_hidden(X3,a_2_2_lattice3(X2,X4))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=X1|~r4_lattice3(X2,X1,X4)|~v10_lattices(X2)|~v4_lattice3(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_53, plain, (v2_struct_0(X1)|r4_lattice3(X1,X2,X3)|X2!=k15_lattice3(X1,X3)|~l3_lattices(X1)|~v10_lattices(X1)|~v4_lattice3(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(cn,[status(thm)],[c_0_49])).
fof(c_0_54, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&(v4_lattice3(X1)&l3_lattices(X1))))=>![X2]:k15_lattice3(X1,X2)=k16_lattice3(X1,a_2_2_lattice3(X1,X2)))), inference(assume_negation,[status(cth)],[t37_lattice3])).
fof(c_0_55, plain, ![X38, X39]:(v2_struct_0(X38)|~l3_lattices(X38)|k16_lattice3(X38,X39)=k15_lattice3(X38,a_2_1_lattice3(X38,X39))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d22_lattice3])])])])).
cnf(c_0_56, plain, (k15_lattice3(X1,a_2_1_lattice3(X1,a_2_2_lattice3(X1,X2)))=k15_lattice3(X1,X2)|v2_struct_0(X1)|~r2_hidden(k15_lattice3(X1,X2),a_2_2_lattice3(X1,X2))|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_29])).
cnf(c_0_57, plain, (r2_hidden(X1,a_2_2_lattice3(X2,X3))|v2_struct_0(X2)|~r4_lattice3(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_lattice3(X2)|~v10_lattices(X2)|~l3_lattices(X2)), inference(er,[status(thm)],[c_0_52])).
cnf(c_0_58, plain, (r4_lattice3(X1,k15_lattice3(X1,X2),X2)|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_53]), c_0_29])).
fof(c_0_59, negated_conjecture, ((~v2_struct_0(esk6_0)&(v10_lattices(esk6_0)&(v4_lattice3(esk6_0)&l3_lattices(esk6_0))))&k15_lattice3(esk6_0,esk7_0)!=k16_lattice3(esk6_0,a_2_2_lattice3(esk6_0,esk7_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_54])])])])).
cnf(c_0_60, plain, (v2_struct_0(X1)|k16_lattice3(X1,X2)=k15_lattice3(X1,a_2_1_lattice3(X1,X2))|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_61, plain, (k15_lattice3(X1,a_2_1_lattice3(X1,a_2_2_lattice3(X1,X2)))=k15_lattice3(X1,X2)|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_29]), c_0_58])).
cnf(c_0_62, negated_conjecture, (k15_lattice3(esk6_0,esk7_0)!=k16_lattice3(esk6_0,a_2_2_lattice3(esk6_0,esk7_0))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_63, plain, (k16_lattice3(X1,a_2_2_lattice3(X1,X2))=k15_lattice3(X1,X2)|v2_struct_0(X1)|~v4_lattice3(X1)|~v10_lattices(X1)|~l3_lattices(X1)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_64, negated_conjecture, (v4_lattice3(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_65, negated_conjecture, (v10_lattices(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_66, negated_conjecture, (l3_lattices(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_67, negated_conjecture, (~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_68, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64]), c_0_65]), c_0_66])]), c_0_67]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 69
# Proof object clause steps            : 52
# Proof object formula steps           : 17
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 26
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 26
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 28
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 28
# Processed clauses                    : 135
# ...of these trivial                  : 0
# ...subsumed                          : 21
# ...remaining for further processing  : 114
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 163
# ...of the previous two non-trivial   : 145
# Contextual simplify-reflections      : 33
# Paramodulations                      : 159
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 78
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 73
# Current number of unprocessed clauses: 64
# ...number of literals in the above   : 600
# Current number of archived formulas  : 0
# Current number of archived clauses   : 32
# Clause-clause subsumption calls (NU) : 3075
# Rec. Clause-clause subsumption calls : 85
# Non-unit clause-clause subsumptions  : 58
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9048

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.003 s
# Total time               : 0.016 s
# Maximum resident set size: 28684 pages
