# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t37_yellow19, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m2_yellow_6(X3,X1,X2)&v3_yellow_6(X3,X1))))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', t37_yellow19)).
fof(d19_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>(v3_yellow_6(X2,X1)<=>~(k10_yellow_6(X1,X2)=k1_xboole_0)))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', d19_yellow_6)).
fof(t32_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>~((r3_waybel_9(X1,X2,X3)&![X4]:(m2_yellow_6(X4,X1,X2)=>~(r2_hidden(X3,k10_yellow_6(X1,X4))))))))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', t32_waybel_9)).
fof(dt_m2_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_struct_0(X1))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>(~(v2_struct_0(X3))&(v4_orders_2(X3)&(v7_waybel_0(X3)&l1_waybel_0(X3,X1)))))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', dt_m2_yellow_6)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', dt_l1_pre_topc)).
fof(d18_yellow_6, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=k10_yellow_6(X1,X2)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r2_hidden(X4,X3)<=>![X5]:(m1_connsp_2(X5,X1,X4)=>r1_waybel_0(X1,X2,X5)))))))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', d18_yellow_6)).
fof(t35_yellow19, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&r3_waybel_9(X1,X2,X3))))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', t35_yellow19)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', t4_subset_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', t7_ordinal1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', t3_subset)).
fof(dt_k10_yellow_6, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&(~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1)))))=>m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', dt_k10_yellow_6)).
fof(t29_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k10_yellow_6(X1,X2))=>r3_waybel_9(X1,X2,X3))))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', t29_waybel_9)).
fof(t31_waybel_9, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>![X3]:(m2_yellow_6(X3,X1,X2)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r3_waybel_9(X1,X3,X4)=>r3_waybel_9(X1,X2,X4)))))), file('data/MPTP/proofs/ee5f08ad0d5b48be89301d598b522e37.E_input', t31_waybel_9)).
fof(c_0_13, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>(v1_compts_1(X1)<=>![X2]:((~(v2_struct_0(X2))&(v4_orders_2(X2)&(v7_waybel_0(X2)&l1_waybel_0(X2,X1))))=>?[X3]:(m2_yellow_6(X3,X1,X2)&v3_yellow_6(X3,X1)))))), inference(assume_negation,[status(cth)],[t37_yellow19])).
fof(c_0_14, plain, ![X24, X25]:((~v3_yellow_6(X25,X24)|k10_yellow_6(X24,X25)!=k1_xboole_0|(v2_struct_0(X25)|(~v4_orders_2(X25)|(~v7_waybel_0(X25)|~l1_waybel_0(X25,X24))))|(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))))&(k10_yellow_6(X24,X25)=k1_xboole_0|v3_yellow_6(X25,X24)|(v2_struct_0(X25)|(~v4_orders_2(X25)|(~v7_waybel_0(X25)|~l1_waybel_0(X25,X24))))|(v2_struct_0(X24)|(~v2_pre_topc(X24)|~l1_pre_topc(X24))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d19_yellow_6])])])])])).
fof(c_0_15, negated_conjecture, ![X8, X9]:((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((((~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0))&((v4_orders_2(esk2_0)|~v1_compts_1(esk1_0))&((v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0))&(l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)))))&(~m2_yellow_6(X8,esk1_0,esk2_0)|~v3_yellow_6(X8,esk1_0)|~v1_compts_1(esk1_0)))&((m2_yellow_6(esk3_1(X9),esk1_0,X9)|(v2_struct_0(X9)|(~v4_orders_2(X9)|(~v7_waybel_0(X9)|~l1_waybel_0(X9,esk1_0))))|v1_compts_1(esk1_0))&(v3_yellow_6(esk3_1(X9),esk1_0)|(v2_struct_0(X9)|(~v4_orders_2(X9)|(~v7_waybel_0(X9)|~l1_waybel_0(X9,esk1_0))))|v1_compts_1(esk1_0))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_13])])])])])])).
cnf(c_0_16, plain, (k10_yellow_6(X1,X2)=k1_xboole_0|v3_yellow_6(X2,X1)|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_17, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (~m2_yellow_6(X1,esk1_0,esk2_0)|~v3_yellow_6(X1,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, negated_conjecture, (k10_yellow_6(esk1_0,X1)=k1_xboole_0|v3_yellow_6(X1,esk1_0)|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17]), c_0_18])]), c_0_19])).
fof(c_0_22, plain, ![X45, X46, X47]:((m2_yellow_6(esk9_3(X45,X46,X47),X45,X46)|~r3_waybel_9(X45,X46,X47)|~m1_subset_1(X47,u1_struct_0(X45))|(v2_struct_0(X46)|(~v4_orders_2(X46)|(~v7_waybel_0(X46)|~l1_waybel_0(X46,X45))))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))))&(r2_hidden(X47,k10_yellow_6(X45,esk9_3(X45,X46,X47)))|~r3_waybel_9(X45,X46,X47)|~m1_subset_1(X47,u1_struct_0(X45))|(v2_struct_0(X46)|(~v4_orders_2(X46)|(~v7_waybel_0(X46)|~l1_waybel_0(X46,X45))))|(v2_struct_0(X45)|(~v2_pre_topc(X45)|~l1_pre_topc(X45))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t32_waybel_9])])])])])])).
fof(c_0_23, plain, ![X12, X13, X14]:((~v2_struct_0(X14)|~m2_yellow_6(X14,X12,X13)|(v2_struct_0(X12)|~l1_struct_0(X12)|(v2_struct_0(X13)|(~v4_orders_2(X13)|(~v7_waybel_0(X13)|~l1_waybel_0(X13,X12))))))&((v4_orders_2(X14)|~m2_yellow_6(X14,X12,X13)|(v2_struct_0(X12)|~l1_struct_0(X12)|(v2_struct_0(X13)|(~v4_orders_2(X13)|(~v7_waybel_0(X13)|~l1_waybel_0(X13,X12))))))&((v7_waybel_0(X14)|~m2_yellow_6(X14,X12,X13)|(v2_struct_0(X12)|~l1_struct_0(X12)|(v2_struct_0(X13)|(~v4_orders_2(X13)|(~v7_waybel_0(X13)|~l1_waybel_0(X13,X12))))))&(l1_waybel_0(X14,X12)|~m2_yellow_6(X14,X12,X13)|(v2_struct_0(X12)|~l1_struct_0(X12)|(v2_struct_0(X13)|(~v4_orders_2(X13)|(~v7_waybel_0(X13)|~l1_waybel_0(X13,X12))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_yellow_6])])])])])).
fof(c_0_24, plain, ![X11]:(~l1_pre_topc(X11)|l1_struct_0(X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_25, negated_conjecture, (k10_yellow_6(esk1_0,X1)=k1_xboole_0|v2_struct_0(X1)|~m2_yellow_6(X1,esk1_0,esk2_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v1_compts_1(esk1_0)), inference(spm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, plain, (m2_yellow_6(esk9_3(X1,X2,X3),X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, negated_conjecture, (v4_orders_2(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_28, negated_conjecture, (v7_waybel_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_29, negated_conjecture, (l1_waybel_0(esk2_0,esk1_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_30, negated_conjecture, (~v2_struct_0(esk2_0)|~v1_compts_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_31, plain, (l1_waybel_0(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_32, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
fof(c_0_33, plain, ![X15, X16, X17, X18, X19, X23]:(((~r2_hidden(X18,X17)|(~m1_connsp_2(X19,X15,X18)|r1_waybel_0(X15,X16,X19))|~m1_subset_1(X18,u1_struct_0(X15))|X17!=k10_yellow_6(X15,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))))&((m1_connsp_2(esk4_4(X15,X16,X17,X18),X15,X18)|r2_hidden(X18,X17)|~m1_subset_1(X18,u1_struct_0(X15))|X17!=k10_yellow_6(X15,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))))&(~r1_waybel_0(X15,X16,esk4_4(X15,X16,X17,X18))|r2_hidden(X18,X17)|~m1_subset_1(X18,u1_struct_0(X15))|X17!=k10_yellow_6(X15,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))))))&((m1_subset_1(esk5_3(X15,X16,X17),u1_struct_0(X15))|X17=k10_yellow_6(X15,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))))&(((m1_connsp_2(esk6_3(X15,X16,X17),X15,esk5_3(X15,X16,X17))|~r2_hidden(esk5_3(X15,X16,X17),X17)|X17=k10_yellow_6(X15,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))))&(~r1_waybel_0(X15,X16,esk6_3(X15,X16,X17))|~r2_hidden(esk5_3(X15,X16,X17),X17)|X17=k10_yellow_6(X15,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15)))))&(r2_hidden(esk5_3(X15,X16,X17),X17)|(~m1_connsp_2(X23,X15,esk5_3(X15,X16,X17))|r1_waybel_0(X15,X16,X23))|X17=k10_yellow_6(X15,X16)|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X15)))|(v2_struct_0(X16)|(~v4_orders_2(X16)|(~v7_waybel_0(X16)|~l1_waybel_0(X16,X15))))|(v2_struct_0(X15)|(~v2_pre_topc(X15)|~l1_pre_topc(X15))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_yellow_6])])])])])])).
fof(c_0_34, plain, ![X40, X41, X44]:(((m1_subset_1(esk7_2(X40,X41),u1_struct_0(X40))|(v2_struct_0(X41)|(~v4_orders_2(X41)|(~v7_waybel_0(X41)|~l1_waybel_0(X41,X40))))|~v1_compts_1(X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&(r3_waybel_9(X40,X41,esk7_2(X40,X41))|(v2_struct_0(X41)|(~v4_orders_2(X41)|(~v7_waybel_0(X41)|~l1_waybel_0(X41,X40))))|~v1_compts_1(X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))))&(((~v2_struct_0(esk8_1(X40))|v1_compts_1(X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&((v4_orders_2(esk8_1(X40))|v1_compts_1(X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&((v7_waybel_0(esk8_1(X40))|v1_compts_1(X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40))))&(l1_waybel_0(esk8_1(X40),X40)|v1_compts_1(X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))))))&(~m1_subset_1(X44,u1_struct_0(X40))|~r3_waybel_9(X40,esk8_1(X40),X44)|v1_compts_1(X40)|(v2_struct_0(X40)|(~v2_pre_topc(X40)|~l1_pre_topc(X40)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_yellow19])])])])])])).
cnf(c_0_35, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|v2_struct_0(esk9_3(esk1_0,esk2_0,X1))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_waybel_0(esk9_3(esk1_0,esk2_0,X1),esk1_0)|~v7_waybel_0(esk9_3(esk1_0,esk2_0,X1))|~v4_orders_2(esk9_3(esk1_0,esk2_0,X1))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26]), c_0_18]), c_0_17])]), c_0_19]), c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_36, plain, (l1_waybel_0(esk9_3(X1,X2,X3),X1)|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_26]), c_0_32])).
cnf(c_0_37, plain, (v7_waybel_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_38, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|X3=k10_yellow_6(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_39, plain, (l1_waybel_0(esk8_1(X1),X1)|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_40, plain, (v4_orders_2(esk8_1(X1))|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (v7_waybel_0(esk8_1(X1))|v1_compts_1(X1)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, plain, (v1_compts_1(X1)|v2_struct_0(X1)|~v2_struct_0(esk8_1(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_43, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|v2_struct_0(esk9_3(esk1_0,esk2_0,X1))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v7_waybel_0(esk9_3(esk1_0,esk2_0,X1))|~v4_orders_2(esk9_3(esk1_0,esk2_0,X1))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_18]), c_0_17])]), c_0_19]), c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_44, plain, (v7_waybel_0(esk9_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_26]), c_0_32])).
cnf(c_0_45, plain, (v4_orders_2(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_46, plain, (X1=k10_yellow_6(X2,esk8_1(X2))|m1_subset_1(esk5_3(X2,esk8_1(X2),X1),u1_struct_0(X2))|v1_compts_1(X2)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40]), c_0_41]), c_0_42])).
fof(c_0_47, plain, ![X26]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X26)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_48, plain, (v2_struct_0(X2)|v2_struct_0(X3)|~v2_struct_0(X1)|~m2_yellow_6(X1,X2,X3)|~l1_struct_0(X2)|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_49, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|v2_struct_0(esk9_3(esk1_0,esk2_0,X1))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v4_orders_2(esk9_3(esk1_0,esk2_0,X1))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_18]), c_0_17])]), c_0_19]), c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_50, plain, (v4_orders_2(esk9_3(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_26]), c_0_32])).
fof(c_0_51, plain, ![X36, X37]:(~r2_hidden(X36,X37)|~r1_tarski(X37,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
fof(c_0_52, plain, ![X38, X39]:((~m1_subset_1(X38,k1_zfmisc_1(X39))|r1_tarski(X38,X39))&(~r1_tarski(X38,X39)|m1_subset_1(X38,k1_zfmisc_1(X39)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_53, plain, (r2_hidden(X1,k10_yellow_6(X2,esk9_3(X2,X3,X1)))|v2_struct_0(X3)|v2_struct_0(X2)|~r3_waybel_9(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v4_orders_2(X3)|~v7_waybel_0(X3)|~l1_waybel_0(X3,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_54, negated_conjecture, (X1=k10_yellow_6(esk1_0,esk8_1(esk1_0))|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),X1),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_55, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~r3_waybel_9(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_waybel_0(X2,X1)|~v7_waybel_0(X2)|~v4_orders_2(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_struct_0(esk9_3(X1,X2,X3))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_26]), c_0_32])).
cnf(c_0_57, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|v2_struct_0(esk9_3(esk1_0,esk2_0,X1))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_18]), c_0_17])]), c_0_19]), c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_58, plain, (r3_waybel_9(X1,X2,esk7_2(X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v1_compts_1(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_59, plain, (m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v1_compts_1(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_60, plain, (m1_connsp_2(esk4_4(X1,X2,X3,X4),X1,X4)|r2_hidden(X4,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_61, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63, negated_conjecture, (r2_hidden(X1,k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1)))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_29]), c_0_18]), c_0_17])]), c_0_19]), c_0_27]), c_0_28]), c_0_30])).
cnf(c_0_64, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_65, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_57]), c_0_18]), c_0_17])]), c_0_19]), c_0_27]), c_0_28]), c_0_29]), c_0_30])).
cnf(c_0_66, negated_conjecture, (r3_waybel_9(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_29]), c_0_18]), c_0_17])]), c_0_19]), c_0_27]), c_0_28]), c_0_30])).
cnf(c_0_67, negated_conjecture, (m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_29]), c_0_18]), c_0_17])]), c_0_19]), c_0_27]), c_0_28]), c_0_30])).
cnf(c_0_68, plain, (m1_connsp_2(esk4_4(X1,esk8_1(X1),X2,X3),X1,X3)|r2_hidden(X3,X2)|v1_compts_1(X1)|v2_struct_0(X1)|X2!=k10_yellow_6(X1,esk8_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_39]), c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_69, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_70, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(X1,k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1)))|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_71, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_64])).
cnf(c_0_72, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r3_waybel_9(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0))|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_66, c_0_64])).
cnf(c_0_73, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),u1_struct_0(esk1_0))|m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_67, c_0_64])).
fof(c_0_74, plain, ![X27, X28]:(v2_struct_0(X27)|(~v2_pre_topc(X27)|~l1_pre_topc(X27))|(v2_struct_0(X28)|(~v4_orders_2(X28)|(~v7_waybel_0(X28)|~l1_waybel_0(X28,X27))))|m1_subset_1(k10_yellow_6(X27,X28),k1_zfmisc_1(u1_struct_0(X27)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_yellow_6])])])).
cnf(c_0_75, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|r1_waybel_0(X1,X2,X4)|X3=k10_yellow_6(X1,X2)|v2_struct_0(X2)|v2_struct_0(X1)|~m1_connsp_2(X4,X1,esk5_3(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_76, negated_conjecture, (m1_connsp_2(esk4_4(esk1_0,esk8_1(esk1_0),X1,X2),esk1_0,X2)|r2_hidden(X2,X1)|v1_compts_1(esk1_0)|X1!=k10_yellow_6(esk1_0,esk8_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_77, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),u1_struct_0(esk1_0))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1)),k1_zfmisc_1(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_78, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0)))=k1_xboole_0|k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73])).
cnf(c_0_79, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_subset_1(k10_yellow_6(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_80, negated_conjecture, (X1=k10_yellow_6(esk1_0,X2)|r1_waybel_0(esk1_0,X2,esk4_4(esk1_0,esk8_1(esk1_0),X3,esk5_3(esk1_0,X2,X1)))|r2_hidden(esk5_3(esk1_0,X2,X1),X3)|r2_hidden(esk5_3(esk1_0,X2,X1),X1)|v1_compts_1(esk1_0)|v2_struct_0(X2)|X3!=k10_yellow_6(esk1_0,esk8_1(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_81, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_55])]), c_0_73]), c_0_72])).
cnf(c_0_82, plain, (m1_subset_1(k10_yellow_6(X1,esk8_1(X1)),k1_zfmisc_1(u1_struct_0(X1)))|v1_compts_1(X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_39]), c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_83, plain, (r2_hidden(X4,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_waybel_0(X1,X2,esk4_4(X1,X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(X1))|X3!=k10_yellow_6(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_84, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),X1,esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0)))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|X1!=k10_yellow_6(esk1_0,esk8_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81]), c_0_55])])).
cnf(c_0_85, negated_conjecture, (m1_subset_1(k10_yellow_6(esk1_0,esk8_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_17]), c_0_18])]), c_0_19])).
cnf(c_0_86, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|X1!=k10_yellow_6(esk1_0,esk8_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_18]), c_0_17])]), c_0_19]), c_0_81])).
cnf(c_0_87, negated_conjecture, (r2_hidden(X1,k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1)))|m1_subset_1(k10_yellow_6(esk1_0,esk8_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_63, c_0_85])).
cnf(c_0_88, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|X1!=k10_yellow_6(esk1_0,esk8_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_39]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_89, negated_conjecture, (m1_subset_1(k10_yellow_6(esk1_0,esk8_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1)),k1_zfmisc_1(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_69, c_0_87])).
cnf(c_0_90, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|m1_subset_1(k10_yellow_6(esk1_0,esk8_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_85])).
cnf(c_0_91, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|X1!=k10_yellow_6(esk1_0,esk8_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_41]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_92, negated_conjecture, (m1_subset_1(k10_yellow_6(esk1_0,esk8_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_90]), c_0_55])])).
cnf(c_0_93, negated_conjecture, (r3_waybel_9(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0))|m1_subset_1(k10_yellow_6(esk1_0,esk8_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(spm,[status(thm)],[c_0_66, c_0_85])).
cnf(c_0_94, negated_conjecture, (m1_subset_1(k10_yellow_6(esk1_0,esk8_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_67, c_0_85])).
cnf(c_0_95, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|X1!=k10_yellow_6(esk1_0,esk8_1(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_40]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_96, negated_conjecture, (m1_subset_1(k10_yellow_6(esk1_0,esk8_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94])).
fof(c_0_97, plain, ![X33, X34, X35]:(v2_struct_0(X33)|(~v2_pre_topc(X33)|~l1_pre_topc(X33))|(v2_struct_0(X34)|(~v4_orders_2(X34)|(~v7_waybel_0(X34)|~l1_waybel_0(X34,X33)))|(~m1_subset_1(X35,u1_struct_0(X33))|(~r2_hidden(X35,k10_yellow_6(X33,X34))|r3_waybel_9(X33,X34,X35))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t29_waybel_9])])])])).
cnf(c_0_98, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k10_yellow_6(esk1_0,esk8_1(esk1_0)))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))), inference(spm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_99, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~r2_hidden(X3,k10_yellow_6(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_100, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k10_yellow_6(esk1_0,esk8_1(esk1_0)))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_98]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_101, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r3_waybel_9(esk1_0,esk8_1(esk1_0),esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_18]), c_0_17])]), c_0_19]), c_0_81])).
cnf(c_0_102, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r3_waybel_9(esk1_0,esk8_1(esk1_0),esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_39]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_103, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r3_waybel_9(esk1_0,esk8_1(esk1_0),esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_41]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_104, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r3_waybel_9(esk1_0,esk8_1(esk1_0),esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_40]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_105, plain, (v1_compts_1(X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~r3_waybel_9(X2,esk8_1(X2),X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_106, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r3_waybel_9(esk1_0,esk8_1(esk1_0),esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_104]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_107, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|v1_compts_1(esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_18]), c_0_17])]), c_0_19]), c_0_81])).
cnf(c_0_108, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r3_waybel_9(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)), inference(spm,[status(thm)],[c_0_66, c_0_107])).
cnf(c_0_109, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_67, c_0_107])).
cnf(c_0_110, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_65, c_0_107])).
cnf(c_0_111, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r3_waybel_9(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_108]), c_0_55])])).
cnf(c_0_112, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_109]), c_0_55])])).
cnf(c_0_113, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|r2_hidden(X1,k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1)))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_63, c_0_107])).
cnf(c_0_114, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0)))=k1_xboole_0|k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112])).
cnf(c_0_115, negated_conjecture, (m2_yellow_6(esk3_1(X1),esk1_0,X1)|v2_struct_0(X1)|v1_compts_1(esk1_0)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_116, negated_conjecture, (l1_struct_0(esk1_0)), inference(spm,[status(thm)],[c_0_32, c_0_18])).
cnf(c_0_117, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k1_xboole_0),k1_xboole_0)|r2_hidden(esk7_2(esk1_0,esk2_0),k1_xboole_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_114]), c_0_112]), c_0_111])).
cnf(c_0_118, negated_conjecture, (l1_waybel_0(esk3_1(X1),esk1_0)|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_115]), c_0_19]), c_0_116])])).
cnf(c_0_119, negated_conjecture, (v4_orders_2(esk3_1(X1))|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_115]), c_0_19]), c_0_116])])).
cnf(c_0_120, negated_conjecture, (v7_waybel_0(esk3_1(X1))|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_115]), c_0_19]), c_0_116])])).
cnf(c_0_121, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)|~v2_struct_0(esk3_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_115]), c_0_19]), c_0_116])])).
cnf(c_0_122, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0|r2_hidden(esk7_2(esk1_0,esk2_0),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_117]), c_0_55])])).
cnf(c_0_123, negated_conjecture, (m1_subset_1(k10_yellow_6(esk1_0,esk3_1(X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_118]), c_0_18]), c_0_17])]), c_0_19]), c_0_119]), c_0_120]), c_0_121])).
cnf(c_0_124, negated_conjecture, (k10_yellow_6(esk1_0,esk8_1(esk1_0))=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_122]), c_0_55])])).
cnf(c_0_125, negated_conjecture, (k10_yellow_6(esk1_0,esk3_1(X1))=k10_yellow_6(esk1_0,esk8_1(esk1_0))|m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_54, c_0_123])).
cnf(c_0_126, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v3_yellow_6(X1,X2)|k10_yellow_6(X2,X1)!=k1_xboole_0|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_127, negated_conjecture, (v3_yellow_6(esk3_1(X1),esk1_0)|v2_struct_0(X1)|v1_compts_1(esk1_0)|~v4_orders_2(X1)|~v7_waybel_0(X1)|~l1_waybel_0(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_128, negated_conjecture, (r3_waybel_9(esk1_0,esk8_1(esk1_0),X1)|v2_struct_0(esk8_1(esk1_0))|~r2_hidden(X1,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_124]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_129, negated_conjecture, (k10_yellow_6(esk1_0,esk3_1(X1))=k10_yellow_6(esk1_0,esk8_1(esk1_0))|r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),X2,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1)))))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))),k10_yellow_6(esk1_0,esk3_1(X1)))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))),X2)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X1)|X2!=k10_yellow_6(esk1_0,esk8_1(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v7_waybel_0(X1)|~v4_orders_2(esk8_1(esk1_0))|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_125]), c_0_123])).
cnf(c_0_130, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(X1)|k10_yellow_6(esk1_0,esk3_1(X1))!=k1_xboole_0|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_127]), c_0_18]), c_0_17])]), c_0_19]), c_0_119]), c_0_120]), c_0_118]), c_0_121])).
cnf(c_0_131, negated_conjecture, (r3_waybel_9(esk1_0,esk8_1(esk1_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~r2_hidden(X1,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_39]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_132, negated_conjecture, (r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),X1,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2)))))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2))),k10_yellow_6(esk1_0,esk3_1(X2)))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2))),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X2)|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v7_waybel_0(X2)|~v4_orders_2(esk8_1(esk1_0))|~v4_orders_2(X2)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_124]), c_0_124]), c_0_130])).
cnf(c_0_133, negated_conjecture, (r3_waybel_9(esk1_0,esk8_1(esk1_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~r2_hidden(X1,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_41]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_134, negated_conjecture, (r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),X1,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2)))))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2))),k10_yellow_6(esk1_0,esk3_1(X2)))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2))),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X2)|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v7_waybel_0(X2)|~v4_orders_2(esk8_1(esk1_0))|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_39]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_135, negated_conjecture, (r3_waybel_9(esk1_0,esk8_1(esk1_0),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~r2_hidden(X1,k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_40]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_136, negated_conjecture, (m1_subset_1(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))),u1_struct_0(esk1_0))|v1_compts_1(esk1_0)|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_124]), c_0_130])).
cnf(c_0_137, negated_conjecture, (r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),X1,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2)))))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2))),k10_yellow_6(esk1_0,esk3_1(X2)))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2))),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X2)|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~v4_orders_2(esk8_1(esk1_0))|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_41]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_138, negated_conjecture, (r3_waybel_9(esk1_0,esk8_1(esk1_0),esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))))|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X1)|~r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))),k1_xboole_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_139, negated_conjecture, (r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),X1,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2)))))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2))),k10_yellow_6(esk1_0,esk3_1(X2)))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X2))),X1)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X2)|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~l1_waybel_0(X2,esk1_0)|~v7_waybel_0(X2)|~v4_orders_2(X2)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_40]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_140, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X1)|~r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))),k1_xboole_0)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_138]), c_0_18]), c_0_17])]), c_0_19]), c_0_136])).
fof(c_0_141, plain, ![X29, X30, X31, X32]:(v2_struct_0(X29)|(~v2_pre_topc(X29)|~l1_pre_topc(X29))|(v2_struct_0(X30)|(~v4_orders_2(X30)|(~v7_waybel_0(X30)|~l1_waybel_0(X30,X29)))|(~m2_yellow_6(X31,X29,X30)|(~m1_subset_1(X32,u1_struct_0(X29))|(~r3_waybel_9(X29,X31,X32)|r3_waybel_9(X29,X30,X32)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t31_waybel_9])])])])).
cnf(c_0_142, negated_conjecture, (r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),k1_xboole_0,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1)))))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))),k10_yellow_6(esk1_0,esk3_1(X1)))|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_55]), c_0_140])).
cnf(c_0_143, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r3_waybel_9(X1,X2,X4)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_orders_2(X2)|~v7_waybel_0(X2)|~l1_waybel_0(X2,X1)|~m2_yellow_6(X3,X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r3_waybel_9(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_144, negated_conjecture, (r3_waybel_9(esk1_0,esk3_1(X1),esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1))))|r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),k1_xboole_0,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(X1)))))|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|v2_struct_0(X1)|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_142]), c_0_18]), c_0_17])]), c_0_19]), c_0_119]), c_0_120]), c_0_118]), c_0_136]), c_0_121])).
cnf(c_0_145, negated_conjecture, (r3_waybel_9(esk1_0,X1,X2)|v1_compts_1(esk1_0)|v2_struct_0(X1)|~r3_waybel_9(esk1_0,esk3_1(X1),X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~l1_waybel_0(X1,esk1_0)|~v7_waybel_0(X1)|~v4_orders_2(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_115]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_146, negated_conjecture, (r3_waybel_9(esk1_0,esk3_1(esk8_1(esk1_0)),esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(esk8_1(esk1_0)))))|r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),k1_xboole_0,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(esk8_1(esk1_0))))))|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_39]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_147, negated_conjecture, (r3_waybel_9(esk1_0,esk8_1(esk1_0),esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(esk8_1(esk1_0)))))|r1_waybel_0(esk1_0,esk8_1(esk1_0),esk4_4(esk1_0,esk8_1(esk1_0),k1_xboole_0,esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(esk8_1(esk1_0))))))|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146]), c_0_136])).
cnf(c_0_148, negated_conjecture, (r3_waybel_9(esk1_0,esk8_1(esk1_0),esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(esk8_1(esk1_0)))))|r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(esk8_1(esk1_0)))),k1_xboole_0)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_147]), c_0_124]), c_0_55]), c_0_18]), c_0_17])]), c_0_19]), c_0_136])).
cnf(c_0_149, negated_conjecture, (r2_hidden(esk5_3(esk1_0,esk8_1(esk1_0),k10_yellow_6(esk1_0,esk3_1(esk8_1(esk1_0)))),k1_xboole_0)|v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_148]), c_0_18]), c_0_17])]), c_0_19]), c_0_136])).
cnf(c_0_150, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~l1_waybel_0(esk8_1(esk1_0),esk1_0)|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_149]), c_0_55])])).
cnf(c_0_151, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~v7_waybel_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_39]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_152, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))|~v4_orders_2(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_41]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_153, negated_conjecture, (v1_compts_1(esk1_0)|v2_struct_0(esk8_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_40]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_154, negated_conjecture, (v1_compts_1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_153]), c_0_18]), c_0_17])]), c_0_19])).
cnf(c_0_155, negated_conjecture, (r2_hidden(X1,k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1)))|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_154])])).
cnf(c_0_156, negated_conjecture, (k10_yellow_6(esk1_0,esk9_3(esk1_0,esk2_0,X1))=k1_xboole_0|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_154])])).
cnf(c_0_157, negated_conjecture, (r2_hidden(X1,k1_xboole_0)|~r3_waybel_9(esk1_0,esk2_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_155, c_0_156])).
cnf(c_0_158, negated_conjecture, (r3_waybel_9(esk1_0,esk2_0,esk7_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_154])])).
cnf(c_0_159, negated_conjecture, (m1_subset_1(esk7_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_154])])).
cnf(c_0_160, negated_conjecture, (r2_hidden(esk7_2(esk1_0,esk2_0),k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158]), c_0_159])])).
cnf(c_0_161, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_160]), c_0_55])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 162
# Proof object clause steps            : 135
# Proof object formula steps           : 27
# Proof object conjectures             : 104
# Proof object clause conjectures      : 101
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 36
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 93
# Proof object simplifying inferences  : 285
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 40
# Processed clauses                    : 311
# ...of these trivial                  : 8
# ...subsumed                          : 19
# ...remaining for further processing  : 284
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 78
# Backward-rewritten                   : 81
# Generated clauses                    : 342
# ...of the previous two non-trivial   : 343
# Contextual simplify-reflections      : 194
# Paramodulations                      : 340
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 85
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 70
# Current number of unprocessed clauses: 74
# ...number of literals in the above   : 966
# Current number of archived formulas  : 0
# Current number of archived clauses   : 199
# Clause-clause subsumption calls (NU) : 18210
# Rec. Clause-clause subsumption calls : 910
# Non-unit clause-clause subsumptions  : 286
# Unit Clause-clause subsumption calls : 94
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 26047

# -------------------------------------------------
# User time                : 0.039 s
# System time              : 0.000 s
# Total time               : 0.039 s
# Maximum resident set size: 28680 pages
