# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', d7_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t48_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t3_boole)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t65_xboole_1)).
fof(t5_xboole_0, axiom, ![X1, X2, X3]:~((r1_xboole_0(X1,X2)&(r2_hidden(X3,k2_xboole_0(X1,X2))&(~((r2_hidden(X3,X1)&~(r2_hidden(X3,X2))))&~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))))), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t5_xboole_0)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', idempotence_k2_xboole_0)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t41_xboole_1)).
fof(t73_xboole_1, conjecture, ![X1, X2, X3]:((r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,X3))=>r1_tarski(X1,X2)), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t73_xboole_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', d3_tarski)).
fof(t17_xboole_1, axiom, ![X1, X2]:r1_tarski(k3_xboole_0(X1,X2),X1), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t17_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', commutativity_k3_xboole_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', commutativity_k2_xboole_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t7_xboole_1)).
fof(t12_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k2_xboole_0(X1,X2)=X2), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t12_xboole_1)).
fof(t4_xboole_1, axiom, ![X1, X2, X3]:k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/ee7bdd032c454b518867a6918d4cd465.E_input', t4_xboole_1)).
fof(c_0_15, plain, ![X10, X11]:((~r1_xboole_0(X10,X11)|k3_xboole_0(X10,X11)=k1_xboole_0)&(k3_xboole_0(X10,X11)!=k1_xboole_0|r1_xboole_0(X10,X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
fof(c_0_16, plain, ![X8, X9]:k4_xboole_0(X8,k4_xboole_0(X8,X9))=k3_xboole_0(X8,X9), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_17, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, ![X12]:k4_xboole_0(X12,k1_xboole_0)=X12, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_20, plain, ![X34]:r1_xboole_0(X34,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_21, plain, ![X48, X49, X50]:(((r2_hidden(X50,X49)|r2_hidden(X50,X48)|~r2_hidden(X50,k2_xboole_0(X48,X49))|~r1_xboole_0(X48,X49))&(~r2_hidden(X50,X48)|r2_hidden(X50,X48)|~r2_hidden(X50,k2_xboole_0(X48,X49))|~r1_xboole_0(X48,X49)))&((r2_hidden(X50,X49)|~r2_hidden(X50,X49)|~r2_hidden(X50,k2_xboole_0(X48,X49))|~r1_xboole_0(X48,X49))&(~r2_hidden(X50,X48)|~r2_hidden(X50,X49)|~r2_hidden(X50,k2_xboole_0(X48,X49))|~r1_xboole_0(X48,X49)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_xboole_0])])])])).
fof(c_0_22, plain, ![X17]:k2_xboole_0(X17,X17)=X17, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
fof(c_0_23, plain, ![X29, X30, X31]:k4_xboole_0(k4_xboole_0(X29,X30),X31)=k4_xboole_0(X29,k2_xboole_0(X30,X31)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_24, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_25, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_27, negated_conjecture, ~(![X1, X2, X3]:((r1_tarski(X1,k2_xboole_0(X2,X3))&r1_xboole_0(X1,X3))=>r1_tarski(X1,X2))), inference(assume_negation,[status(cth)],[t73_xboole_1])).
cnf(c_0_28, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X2,X3))|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_30, plain, ![X42, X43, X44, X45, X46]:((~r1_tarski(X42,X43)|(~r2_hidden(X44,X42)|r2_hidden(X44,X43)))&((r2_hidden(esk5_2(X45,X46),X45)|r1_tarski(X45,X46))&(~r2_hidden(esk5_2(X45,X46),X46)|r1_tarski(X45,X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_31, plain, ![X35, X36]:r1_tarski(k3_xboole_0(X35,X36),X35), inference(variable_rename,[status(thm)],[t17_xboole_1])).
fof(c_0_32, plain, ![X18, X19]:k3_xboole_0(X18,X19)=k3_xboole_0(X19,X18), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_33, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
fof(c_0_35, negated_conjecture, ((r1_tarski(esk1_0,k2_xboole_0(esk2_0,esk3_0))&r1_xboole_0(esk1_0,esk3_0))&~r1_tarski(esk1_0,esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])).
fof(c_0_36, plain, ![X37, X38]:k2_xboole_0(X37,X38)=k2_xboole_0(X38,X37), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_37, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(X2,X2)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_38, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_39, plain, ![X15, X16]:r1_tarski(X15,k2_xboole_0(X15,X16)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_40, plain, (r1_tarski(k3_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_42, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k4_xboole_0(k1_xboole_0,X2)), inference(spm,[status(thm)],[c_0_33, c_0_34])).
fof(c_0_43, plain, ![X13, X14]:(~r1_tarski(X13,X14)|k2_xboole_0(X13,X14)=X14), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_xboole_1])])).
cnf(c_0_44, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_45, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_46, plain, (r1_tarski(X1,X2)|~r1_xboole_0(X1,X1)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_47, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_48, plain, ![X39, X40, X41]:k2_xboole_0(k2_xboole_0(X39,X40),X41)=k2_xboole_0(X39,k2_xboole_0(X40,X41)), inference(variable_rename,[status(thm)],[t4_xboole_1])).
cnf(c_0_49, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X1)), inference(rw,[status(thm)],[c_0_40, c_0_18])).
cnf(c_0_50, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_18]), c_0_18])).
cnf(c_0_51, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_29]), c_0_34])).
cnf(c_0_52, plain, (k2_xboole_0(X1,X2)=X2|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, negated_conjecture, (r1_tarski(esk1_0,k2_xboole_0(esk3_0,esk2_0))), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_54, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_46, c_0_26])).
cnf(c_0_55, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_47, c_0_45])).
cnf(c_0_56, plain, (k2_xboole_0(k2_xboole_0(X1,X2),X3)=k2_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_58, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(rw,[status(thm)],[c_0_42, c_0_51])).
cnf(c_0_59, negated_conjecture, (k2_xboole_0(esk1_0,k2_xboole_0(esk3_0,esk2_0))=k2_xboole_0(esk3_0,esk2_0)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_60, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_52, c_0_54])).
cnf(c_0_61, plain, (r1_tarski(X1,k2_xboole_0(X2,k2_xboole_0(X3,X1)))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_62, plain, (k2_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_49]), c_0_45])).
cnf(c_0_63, plain, (r1_tarski(k4_xboole_0(X1,k2_xboole_0(X2,k4_xboole_0(X1,k2_xboole_0(X2,X3)))),X3)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_33]), c_0_33])).
cnf(c_0_64, negated_conjecture, (k4_xboole_0(esk1_0,k2_xboole_0(esk3_0,esk2_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_65, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(spm,[status(thm)],[c_0_45, c_0_60])).
cnf(c_0_66, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k2_xboole_0(X3,X1))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_67, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,esk3_0),esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_65])).
cnf(c_0_68, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),k2_xboole_0(X3,X2))), inference(spm,[status(thm)],[c_0_66, c_0_50])).
cnf(c_0_69, negated_conjecture, (k2_xboole_0(esk2_0,k4_xboole_0(esk1_0,esk3_0))=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_67]), c_0_45])).
cnf(c_0_70, negated_conjecture, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,k4_xboole_0(esk1_0,esk3_0))),esk2_0)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_71, negated_conjecture, (r1_xboole_0(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_72, negated_conjecture, (~r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_73, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_24]), c_0_25]), c_0_71])]), c_0_72]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 74
# Proof object clause steps            : 43
# Proof object formula steps           : 31
# Proof object conjectures             : 13
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 17
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 21
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 32
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 29
# Processed clauses                    : 667
# ...of these trivial                  : 170
# ...subsumed                          : 227
# ...remaining for further processing  : 270
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 22
# Backward-rewritten                   : 34
# Generated clauses                    : 7449
# ...of the previous two non-trivial   : 3767
# Contextual simplify-reflections      : 1
# Paramodulations                      : 7435
# Factorizations                       : 2
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 209
#    Positive orientable unit clauses  : 138
#    Positive unorientable unit clauses: 4
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 56
# Current number of unprocessed clauses: 3063
# ...number of literals in the above   : 4439
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 1359
# Rec. Clause-clause subsumption calls : 577
# Non-unit clause-clause subsumptions  : 94
# Unit Clause-clause subsumption calls : 160
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 931
# BW rewrite match successes           : 200
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 69113

# -------------------------------------------------
# User time                : 0.035 s
# System time              : 0.007 s
# Total time               : 0.042 s
# Maximum resident set size: 28680 pages
