# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t60_funct_1, conjecture, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X1)&(k1_relat_1(X1)=k2_relat_1(X2)&(k2_relat_1(X1)=k1_relat_1(X2)&![X3, X4]:((r2_hidden(X3,k1_relat_1(X1))&r2_hidden(X4,k1_relat_1(X2)))=>(k1_funct_1(X1,X3)=X4<=>k1_funct_1(X2,X4)=X3)))))=>X2=k2_funct_1(X1)))), file('data/MPTP/proofs/ef0027bb2b4b44f09cb44851d226b9e2.E_input', t60_funct_1)).
fof(t12_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>(r2_hidden(X1,k1_relat_1(X2))=>r2_hidden(k1_funct_1(X2,X1),k2_relat_1(X2)))), file('data/MPTP/proofs/ef0027bb2b4b44f09cb44851d226b9e2.E_input', t12_funct_1)).
fof(t56_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X2)&r2_hidden(X1,k1_relat_1(X2)))=>(X1=k1_funct_1(k2_funct_1(X2),k1_funct_1(X2,X1))&X1=k1_funct_1(k5_relat_1(X2,k2_funct_1(X2)),X1)))), file('data/MPTP/proofs/ef0027bb2b4b44f09cb44851d226b9e2.E_input', t56_funct_1)).
fof(t55_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))&k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))))), file('data/MPTP/proofs/ef0027bb2b4b44f09cb44851d226b9e2.E_input', t55_funct_1)).
fof(t9_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((k1_relat_1(X1)=k1_relat_1(X2)&![X3]:(r2_hidden(X3,k1_relat_1(X1))=>k1_funct_1(X1,X3)=k1_funct_1(X2,X3)))=>X1=X2))), file('data/MPTP/proofs/ef0027bb2b4b44f09cb44851d226b9e2.E_input', t9_funct_1)).
fof(dt_k2_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_relat_1(k2_funct_1(X1))&v1_funct_1(k2_funct_1(X1)))), file('data/MPTP/proofs/ef0027bb2b4b44f09cb44851d226b9e2.E_input', dt_k2_funct_1)).
fof(c_0_6, negated_conjecture, ~(![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>((v2_funct_1(X1)&(k1_relat_1(X1)=k2_relat_1(X2)&(k2_relat_1(X1)=k1_relat_1(X2)&![X3, X4]:((r2_hidden(X3,k1_relat_1(X1))&r2_hidden(X4,k1_relat_1(X2)))=>(k1_funct_1(X1,X3)=X4<=>k1_funct_1(X2,X4)=X3)))))=>X2=k2_funct_1(X1))))), inference(assume_negation,[status(cth)],[t60_funct_1])).
fof(c_0_7, negated_conjecture, ![X7, X8]:((v1_relat_1(esk1_0)&v1_funct_1(esk1_0))&((v1_relat_1(esk2_0)&v1_funct_1(esk2_0))&((v2_funct_1(esk1_0)&(k1_relat_1(esk1_0)=k2_relat_1(esk2_0)&(k2_relat_1(esk1_0)=k1_relat_1(esk2_0)&((k1_funct_1(esk1_0,X7)!=X8|k1_funct_1(esk2_0,X8)=X7|(~r2_hidden(X7,k1_relat_1(esk1_0))|~r2_hidden(X8,k1_relat_1(esk2_0))))&(k1_funct_1(esk2_0,X8)!=X7|k1_funct_1(esk1_0,X7)=X8|(~r2_hidden(X7,k1_relat_1(esk1_0))|~r2_hidden(X8,k1_relat_1(esk2_0))))))))&esk2_0!=k2_funct_1(esk1_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
fof(c_0_8, plain, ![X9, X10]:(~v1_relat_1(X10)|~v1_funct_1(X10)|(~r2_hidden(X9,k1_relat_1(X10))|r2_hidden(k1_funct_1(X10,X9),k2_relat_1(X10)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_funct_1])])).
cnf(c_0_9, negated_conjecture, (k1_funct_1(esk1_0,X2)=X1|k1_funct_1(esk2_0,X1)!=X2|~r2_hidden(X2,k1_relat_1(esk1_0))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, negated_conjecture, (k1_relat_1(esk1_0)=k2_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_14, plain, ![X14, X15]:((X14=k1_funct_1(k2_funct_1(X15),k1_funct_1(X15,X14))|(~v2_funct_1(X15)|~r2_hidden(X14,k1_relat_1(X15)))|(~v1_relat_1(X15)|~v1_funct_1(X15)))&(X14=k1_funct_1(k5_relat_1(X15,k2_funct_1(X15)),X14)|(~v2_funct_1(X15)|~r2_hidden(X14,k1_relat_1(X15)))|(~v1_relat_1(X15)|~v1_funct_1(X15)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_funct_1])])])).
cnf(c_0_15, negated_conjecture, (k1_funct_1(esk1_0,k1_funct_1(esk2_0,X1))=X1|~r2_hidden(k1_funct_1(esk2_0,X1),k1_relat_1(esk1_0))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_16, negated_conjecture, (r2_hidden(k1_funct_1(esk2_0,X1),k1_relat_1(esk1_0))|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_10, c_0_11]), c_0_12]), c_0_13])])).
fof(c_0_17, plain, ![X16]:((k2_relat_1(X16)=k1_relat_1(k2_funct_1(X16))|~v2_funct_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16)))&(k1_relat_1(X16)=k2_relat_1(k2_funct_1(X16))|~v2_funct_1(X16)|(~v1_relat_1(X16)|~v1_funct_1(X16)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_funct_1])])])).
fof(c_0_18, plain, ![X11, X12]:((r2_hidden(esk3_2(X11,X12),k1_relat_1(X11))|k1_relat_1(X11)!=k1_relat_1(X12)|X11=X12|(~v1_relat_1(X12)|~v1_funct_1(X12))|(~v1_relat_1(X11)|~v1_funct_1(X11)))&(k1_funct_1(X11,esk3_2(X11,X12))!=k1_funct_1(X12,esk3_2(X11,X12))|k1_relat_1(X11)!=k1_relat_1(X12)|X11=X12|(~v1_relat_1(X12)|~v1_funct_1(X12))|(~v1_relat_1(X11)|~v1_funct_1(X11)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_funct_1])])])])])).
cnf(c_0_19, plain, (X1=k1_funct_1(k2_funct_1(X2),k1_funct_1(X2,X1))|~v2_funct_1(X2)|~r2_hidden(X1,k1_relat_1(X2))|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, negated_conjecture, (k1_funct_1(esk1_0,k1_funct_1(esk2_0,X1))=X1|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21, negated_conjecture, (v2_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_23, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_24, negated_conjecture, (k2_relat_1(esk1_0)=k1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_25, plain, (k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, plain, (X1=X2|k1_funct_1(X1,esk3_2(X1,X2))!=k1_funct_1(X2,esk3_2(X1,X2))|k1_relat_1(X1)!=k1_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_27, negated_conjecture, (k1_funct_1(k2_funct_1(esk1_0),X1)=k1_funct_1(esk2_0,X1)|~r2_hidden(X1,k1_relat_1(esk2_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_19, c_0_20]), c_0_21]), c_0_22]), c_0_23])]), c_0_16])).
cnf(c_0_28, negated_conjecture, (k1_relat_1(k2_funct_1(esk1_0))=k1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_21]), c_0_22]), c_0_23])])).
cnf(c_0_29, negated_conjecture, (X1=k2_funct_1(esk1_0)|k1_funct_1(X1,esk3_2(X1,k2_funct_1(esk1_0)))!=k1_funct_1(esk2_0,esk3_2(X1,k2_funct_1(esk1_0)))|k1_relat_1(X1)!=k1_relat_1(esk2_0)|~r2_hidden(esk3_2(X1,k2_funct_1(esk1_0)),k1_relat_1(esk2_0))|~v1_funct_1(k2_funct_1(esk1_0))|~v1_funct_1(X1)|~v1_relat_1(k2_funct_1(esk1_0))|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])).
cnf(c_0_30, plain, (r2_hidden(esk3_2(X1,X2),k1_relat_1(X1))|X1=X2|k1_relat_1(X1)!=k1_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_31, negated_conjecture, (esk2_0!=k2_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_32, plain, ![X17]:((v1_relat_1(k2_funct_1(X17))|(~v1_relat_1(X17)|~v1_funct_1(X17)))&(v1_funct_1(k2_funct_1(X17))|(~v1_relat_1(X17)|~v1_funct_1(X17)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_1])])])).
cnf(c_0_33, negated_conjecture, (~v1_funct_1(k2_funct_1(esk1_0))|~v1_relat_1(k2_funct_1(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_30]), c_0_12]), c_0_13]), c_0_28])]), c_0_31])).
cnf(c_0_34, plain, (v1_funct_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_35, negated_conjecture, (~v1_relat_1(k2_funct_1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_22]), c_0_23])])).
cnf(c_0_36, plain, (v1_relat_1(k2_funct_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_36]), c_0_22]), c_0_23])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 25
# Proof object formula steps           : 13
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 19
# Processed clauses                    : 130
# ...of these trivial                  : 0
# ...subsumed                          : 37
# ...remaining for further processing  : 93
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 21
# Backward-rewritten                   : 0
# Generated clauses                    : 193
# ...of the previous two non-trivial   : 168
# Contextual simplify-reflections      : 33
# Paramodulations                      : 190
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 51
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 41
# Current number of unprocessed clauses: 74
# ...number of literals in the above   : 789
# Current number of archived formulas  : 0
# Current number of archived clauses   : 40
# Clause-clause subsumption calls (NU) : 1101
# Rec. Clause-clause subsumption calls : 201
# Non-unit clause-clause subsumptions  : 91
# Unit Clause-clause subsumption calls : 4
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9420

# -------------------------------------------------
# User time                : 0.007 s
# System time              : 0.007 s
# Total time               : 0.014 s
# Maximum resident set size: 28840 pages
