# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/ef9f3d919b9d44b4bdcee3acdbad4346.E_input', t3_xboole_0)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/ef9f3d919b9d44b4bdcee3acdbad4346.E_input', t65_xboole_1)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/ef9f3d919b9d44b4bdcee3acdbad4346.E_input', t7_tarski)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/ef9f3d919b9d44b4bdcee3acdbad4346.E_input', d4_relat_1)).
fof(d1_relat_1, axiom, ![X1]:(v1_relat_1(X1)<=>![X2]:~((r2_hidden(X2,X1)&![X3, X4]:~(X2=k4_tarski(X3,X4))))), file('data/MPTP/proofs/ef9f3d919b9d44b4bdcee3acdbad4346.E_input', d1_relat_1)).
fof(t19_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>~((r2_hidden(X1,k2_relat_1(X2))&![X3]:~(r2_hidden(X3,k1_relat_1(X2)))))), file('data/MPTP/proofs/ef9f3d919b9d44b4bdcee3acdbad4346.E_input', t19_relat_1)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)<=>r2_hidden(X3,X2))=>X1=X2), file('data/MPTP/proofs/ef9f3d919b9d44b4bdcee3acdbad4346.E_input', t2_tarski)).
fof(t60_relat_1, conjecture, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/ef9f3d919b9d44b4bdcee3acdbad4346.E_input', t60_relat_1)).
fof(c_0_8, plain, ![X10, X11, X13, X14, X15]:(((r2_hidden(esk2_2(X10,X11),X10)|r1_xboole_0(X10,X11))&(r2_hidden(esk2_2(X10,X11),X11)|r1_xboole_0(X10,X11)))&(~r2_hidden(X15,X13)|~r2_hidden(X15,X14)|~r1_xboole_0(X13,X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_9, plain, ![X9]:r1_xboole_0(X9,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_10, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, plain, ![X5, X6, X8]:((r2_hidden(esk1_2(X5,X6),X6)|~r2_hidden(X5,X6))&(~r2_hidden(X8,X6)|~r2_hidden(X8,esk1_2(X5,X6))|~r2_hidden(X5,X6))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
cnf(c_0_13, plain, (~r2_hidden(X1,k1_xboole_0)|~r2_hidden(X1,X2)), inference(spm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_14, plain, (r2_hidden(esk1_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
fof(c_0_15, plain, ![X30, X31, X32, X34, X35, X36, X37, X39]:(((~r2_hidden(X32,X31)|r2_hidden(k4_tarski(X32,esk8_3(X30,X31,X32)),X30)|X31!=k1_relat_1(X30))&(~r2_hidden(k4_tarski(X34,X35),X30)|r2_hidden(X34,X31)|X31!=k1_relat_1(X30)))&((~r2_hidden(esk9_2(X36,X37),X37)|~r2_hidden(k4_tarski(esk9_2(X36,X37),X39),X36)|X37=k1_relat_1(X36))&(r2_hidden(esk9_2(X36,X37),X37)|r2_hidden(k4_tarski(esk9_2(X36,X37),esk10_2(X36,X37)),X36)|X37=k1_relat_1(X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_16, plain, (~r2_hidden(esk1_2(X1,k1_xboole_0),X2)|~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17, plain, (r2_hidden(k4_tarski(X1,esk8_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_16, c_0_14])).
cnf(c_0_19, plain, (r2_hidden(k4_tarski(X1,esk8_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_17])).
fof(c_0_20, plain, ![X16, X17, X20, X22, X23]:((~v1_relat_1(X16)|(~r2_hidden(X17,X16)|X17=k4_tarski(esk3_2(X16,X17),esk4_2(X16,X17))))&((r2_hidden(esk5_1(X20),X20)|v1_relat_1(X20))&(esk5_1(X20)!=k4_tarski(X22,X23)|v1_relat_1(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_relat_1])])])])])])])).
cnf(c_0_21, plain, (~r2_hidden(X1,k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_18, c_0_19])).
fof(c_0_22, plain, ![X27, X28]:(~v1_relat_1(X28)|(~r2_hidden(X27,k2_relat_1(X28))|r2_hidden(esk7_2(X27,X28),k1_relat_1(X28)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t19_relat_1])])])])).
cnf(c_0_23, plain, (r2_hidden(esk5_1(X1),X1)|v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_24, plain, (~r2_hidden(X1,k1_relat_1(k1_relat_1(k1_xboole_0)))), inference(spm,[status(thm)],[c_0_21, c_0_19])).
cnf(c_0_25, plain, (r2_hidden(esk7_2(X2,X1),k1_relat_1(X1))|~v1_relat_1(X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (v1_relat_1(k1_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_21, c_0_23])).
fof(c_0_27, plain, ![X24, X25]:((~r2_hidden(esk6_2(X24,X25),X24)|~r2_hidden(esk6_2(X24,X25),X25)|X24=X25)&(r2_hidden(esk6_2(X24,X25),X24)|r2_hidden(esk6_2(X24,X25),X25)|X24=X25)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
cnf(c_0_28, plain, (~r2_hidden(X1,k2_relat_1(k1_relat_1(k1_xboole_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25]), c_0_26])])).
cnf(c_0_29, plain, (r2_hidden(esk6_2(X1,X2),X1)|r2_hidden(esk6_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30, plain, (k2_relat_1(k1_relat_1(k1_xboole_0))=X1|r2_hidden(esk6_2(k2_relat_1(k1_relat_1(k1_xboole_0)),X1),X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_31, negated_conjecture, ~((k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0)), inference(assume_negation,[status(cth)],[t60_relat_1])).
cnf(c_0_32, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_18, c_0_23])).
cnf(c_0_33, plain, (k2_relat_1(k1_relat_1(k1_xboole_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_18, c_0_30])).
fof(c_0_34, negated_conjecture, (k1_relat_1(k1_xboole_0)!=k1_xboole_0|k2_relat_1(k1_xboole_0)!=k1_xboole_0), inference(fof_nnf,[status(thm)],[c_0_31])).
cnf(c_0_35, plain, (~r2_hidden(X1,k2_relat_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_25]), c_0_32])])).
cnf(c_0_36, plain, (k1_xboole_0=X1|r2_hidden(esk6_2(k1_xboole_0,X1),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_30, c_0_33]), c_0_33])).
cnf(c_0_37, negated_conjecture, (k1_relat_1(k1_xboole_0)!=k1_xboole_0|k2_relat_1(k1_xboole_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_38, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_21, c_0_36])).
cnf(c_0_40, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38])]), c_0_39])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 41
# Proof object clause steps            : 24
# Proof object formula steps           : 17
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 17
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 17
# Processed clauses                    : 178
# ...of these trivial                  : 2
# ...subsumed                          : 58
# ...remaining for further processing  : 117
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 64
# Generated clauses                    : 378
# ...of the previous two non-trivial   : 400
# Contextual simplify-reflections      : 0
# Paramodulations                      : 373
# Factorizations                       : 2
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 34
#    Positive orientable unit clauses  : 5
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 28
# Current number of unprocessed clauses: 228
# ...number of literals in the above   : 368
# Current number of archived formulas  : 0
# Current number of archived clauses   : 81
# Clause-clause subsumption calls (NU) : 119
# Rec. Clause-clause subsumption calls : 86
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 144
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 37
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3655

# -------------------------------------------------
# User time                : 0.012 s
# System time              : 0.000 s
# Total time               : 0.012 s
# Maximum resident set size: 28616 pages
