# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_funct_2, axiom, ![X1, X2, X3]:(X3=k1_funct_2(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>?[X5]:((v1_relat_1(X5)&v1_funct_1(X5))&(X4=X5&(k1_relat_1(X5)=X1&r1_tarski(k2_relat_1(X5),X2)))))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', d2_funct_2)).
fof(d5_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k1_relat_1(X1))&X3=k1_funct_1(X1,X4))))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', d5_funct_1)).
fof(t121_funct_2, conjecture, ![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t121_funct_2)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', d1_xboole_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t2_subset)).
fof(existence_m1_subset_1, axiom, ![X1]:?[X2]:m1_subset_1(X2,X1), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', existence_m1_subset_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t113_zfmisc_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t5_subset)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t3_subset)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', cc4_relset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', d10_xboole_0)).
fof(t21_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t21_relat_1)).
fof(t65_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>(k1_relat_1(X1)=k1_xboole_0<=>k2_relat_1(X1)=k1_xboole_0)), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t65_relat_1)).
fof(dt_k6_partfun1, axiom, ![X1]:(v1_partfun1(k6_partfun1(X1),X1)&m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', dt_k6_partfun1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t7_xboole_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', fc1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', d3_tarski)).
fof(t152_zfmisc_1, axiom, ![X1, X2]:~(r2_hidden(X1,k2_zfmisc_1(X1,X2))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t152_zfmisc_1)).
fof(t11_relset_1, axiom, ![X1, X2, X3]:(v1_relat_1(X3)=>((r1_tarski(k1_relat_1(X3),X1)&r1_tarski(k2_relat_1(X3),X2))=>m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t11_relset_1)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_partfun1(X3,X1))=>(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', cc1_funct_2)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', t3_funct_2)).
fof(cc5_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&v1_partfun1(X3,X1))))), file('data/MPTP/proofs/efae297dfdc44cdbab9eda31d9b987f3.E_input', cc5_funct_2)).
fof(c_0_22, plain, ![X45, X46, X47, X48, X50, X51, X52, X53, X54, X56]:(((((v1_relat_1(esk7_4(X45,X46,X47,X48))|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46))&(v1_funct_1(esk7_4(X45,X46,X47,X48))|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46)))&((X48=esk7_4(X45,X46,X47,X48)|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46))&((k1_relat_1(esk7_4(X45,X46,X47,X48))=X45|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46))&(r1_tarski(k2_relat_1(esk7_4(X45,X46,X47,X48)),X46)|~r2_hidden(X48,X47)|X47!=k1_funct_2(X45,X46)))))&(~v1_relat_1(X51)|~v1_funct_1(X51)|(X50!=X51|(k1_relat_1(X51)!=X45|~r1_tarski(k2_relat_1(X51),X46)))|r2_hidden(X50,X47)|X47!=k1_funct_2(X45,X46)))&((~r2_hidden(esk8_3(X52,X53,X54),X54)|(~v1_relat_1(X56)|~v1_funct_1(X56)|(esk8_3(X52,X53,X54)!=X56|(k1_relat_1(X56)!=X52|~r1_tarski(k2_relat_1(X56),X53))))|X54=k1_funct_2(X52,X53))&(((v1_relat_1(esk9_3(X52,X53,X54))|r2_hidden(esk8_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53))&(v1_funct_1(esk9_3(X52,X53,X54))|r2_hidden(esk8_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53)))&((esk8_3(X52,X53,X54)=esk9_3(X52,X53,X54)|r2_hidden(esk8_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53))&((k1_relat_1(esk9_3(X52,X53,X54))=X52|r2_hidden(esk8_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53))&(r1_tarski(k2_relat_1(esk9_3(X52,X53,X54)),X53)|r2_hidden(esk8_3(X52,X53,X54),X54)|X54=k1_funct_2(X52,X53))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_funct_2])])])])])])).
fof(c_0_23, plain, ![X58, X59, X60, X62, X63, X64, X66]:((((r2_hidden(esk10_3(X58,X59,X60),k1_relat_1(X58))|~r2_hidden(X60,X59)|X59!=k2_relat_1(X58)|(~v1_relat_1(X58)|~v1_funct_1(X58)))&(X60=k1_funct_1(X58,esk10_3(X58,X59,X60))|~r2_hidden(X60,X59)|X59!=k2_relat_1(X58)|(~v1_relat_1(X58)|~v1_funct_1(X58))))&(~r2_hidden(X63,k1_relat_1(X58))|X62!=k1_funct_1(X58,X63)|r2_hidden(X62,X59)|X59!=k2_relat_1(X58)|(~v1_relat_1(X58)|~v1_funct_1(X58))))&((~r2_hidden(esk11_2(X58,X64),X64)|(~r2_hidden(X66,k1_relat_1(X58))|esk11_2(X58,X64)!=k1_funct_1(X58,X66))|X64=k2_relat_1(X58)|(~v1_relat_1(X58)|~v1_funct_1(X58)))&((r2_hidden(esk12_2(X58,X64),k1_relat_1(X58))|r2_hidden(esk11_2(X58,X64),X64)|X64=k2_relat_1(X58)|(~v1_relat_1(X58)|~v1_funct_1(X58)))&(esk11_2(X58,X64)=k1_funct_1(X58,esk12_2(X58,X64))|r2_hidden(esk11_2(X58,X64),X64)|X64=k2_relat_1(X58)|(~v1_relat_1(X58)|~v1_funct_1(X58)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_funct_1])])])])])])).
cnf(c_0_24, plain, (k1_relat_1(esk7_4(X1,X2,X3,X4))=X1|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_25, plain, (X1=esk7_4(X2,X3,X4,X1)|~r2_hidden(X1,X4)|X4!=k1_funct_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_26, negated_conjecture, ~(![X1, X2, X3]:(r2_hidden(X3,k1_funct_2(X1,X2))=>(v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))))), inference(assume_negation,[status(cth)],[t121_funct_2])).
cnf(c_0_27, plain, (v1_relat_1(esk7_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_28, plain, ![X10, X11, X12]:((~v1_xboole_0(X10)|~r2_hidden(X11,X10))&(r2_hidden(esk4_1(X12),X12)|v1_xboole_0(X12))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_29, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k1_relat_1(X2))|X3!=k1_funct_1(X2,X1)|X4!=k2_relat_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_30, plain, ![X43, X44]:(~m1_subset_1(X43,X44)|(v1_xboole_0(X44)|r2_hidden(X43,X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_31, plain, ![X15]:m1_subset_1(esk5_1(X15),X15), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_32, plain, (v1_funct_1(esk7_4(X1,X2,X3,X4))|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_33, plain, ![X27, X28]:((k2_zfmisc_1(X27,X28)!=k1_xboole_0|(X27=k1_xboole_0|X28=k1_xboole_0))&((X27!=k1_xboole_0|k2_zfmisc_1(X27,X28)=k1_xboole_0)&(X28!=k1_xboole_0|k2_zfmisc_1(X27,X28)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_34, plain, ![X29, X30, X31]:(~r2_hidden(X29,X30)|(~m1_subset_1(X30,k1_zfmisc_1(X31))|~v1_xboole_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_35, plain, ![X35, X36]:((~m1_subset_1(X35,k1_zfmisc_1(X36))|r1_tarski(X35,X36))&(~r1_tarski(X35,X36)|m1_subset_1(X35,k1_zfmisc_1(X36)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_36, plain, (k1_relat_1(esk7_4(X1,X2,k1_funct_2(X1,X2),X3))=X1|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_37, plain, (esk7_4(X1,X2,k1_funct_2(X1,X2),X3)=X3|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_25])).
fof(c_0_38, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))&(~v1_funct_1(esk3_0)|(~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
cnf(c_0_39, plain, (v1_relat_1(esk7_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_27])).
fof(c_0_40, plain, ![X17, X18, X19]:(~v1_xboole_0(X17)|(~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X18,X17)))|v1_xboole_0(X19))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
fof(c_0_41, plain, ![X68, X69]:(((r1_tarski(X68,X69)|X68!=X69)&(r1_tarski(X69,X68)|X68!=X69))&(~r1_tarski(X68,X69)|~r1_tarski(X69,X68)|X68=X69)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_42, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_43, plain, (r2_hidden(k1_funct_1(X1,X2),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_29])])).
cnf(c_0_44, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_45, plain, (m1_subset_1(esk5_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_46, plain, (v1_funct_1(esk7_4(X1,X2,k1_funct_2(X1,X2),X3))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_47, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_48, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_49, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_50, plain, ![X20]:(~v1_relat_1(X20)|r1_tarski(X20,k2_zfmisc_1(k1_relat_1(X20),k2_relat_1(X20)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_relat_1])])).
fof(c_0_51, plain, ![X9]:((k1_relat_1(X9)!=k1_xboole_0|k2_relat_1(X9)=k1_xboole_0|~v1_relat_1(X9))&(k2_relat_1(X9)!=k1_xboole_0|k1_relat_1(X9)=k1_xboole_0|~v1_relat_1(X9))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_relat_1])])])).
cnf(c_0_52, plain, (k1_relat_1(X1)=X2|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_53, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_54, plain, (v1_relat_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_39, c_0_37])).
fof(c_0_55, plain, ![X14]:(v1_partfun1(k6_partfun1(X14),X14)&m1_subset_1(k6_partfun1(X14),k1_zfmisc_1(k2_zfmisc_1(X14,X14)))), inference(variable_rename,[status(thm)],[dt_k6_partfun1])).
fof(c_0_56, plain, ![X72]:(X72=k1_xboole_0|r2_hidden(esk13_1(X72),X72)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_57, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_58, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_59, plain, (~v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k1_relat_1(X1))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_60, plain, (v1_xboole_0(X1)|r2_hidden(esk5_1(X1),X1)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_61, plain, (v1_funct_1(X1)|~r2_hidden(X1,k1_funct_2(X2,X3))), inference(spm,[status(thm)],[c_0_46, c_0_37])).
cnf(c_0_62, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_47])).
cnf(c_0_63, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_64, plain, (~r1_tarski(X1,X2)|~v1_xboole_0(X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_65, plain, (r1_tarski(X1,k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_66, plain, (k2_relat_1(X1)=k1_xboole_0|k1_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_67, negated_conjecture, (k1_relat_1(esk3_0)=esk1_0), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_68, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_54, c_0_53])).
fof(c_0_69, plain, ![X37, X38, X39, X40, X41]:((~r1_tarski(X37,X38)|(~r2_hidden(X39,X37)|r2_hidden(X39,X38)))&((r2_hidden(esk6_2(X40,X41),X40)|r1_tarski(X40,X41))&(~r2_hidden(esk6_2(X40,X41),X41)|r1_tarski(X40,X41)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_70, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_71, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_72, plain, (X1=k1_xboole_0|r2_hidden(esk13_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_73, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_xboole_0(X3)), inference(spm,[status(thm)],[c_0_57, c_0_49])).
cnf(c_0_74, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_58])).
cnf(c_0_75, plain, (v1_xboole_0(k1_relat_1(X1))|~v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_76, negated_conjecture, (v1_funct_1(esk3_0)), inference(spm,[status(thm)],[c_0_61, c_0_53])).
cnf(c_0_77, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_62]), c_0_63])])).
cnf(c_0_78, plain, (~v1_xboole_0(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_79, negated_conjecture, (k2_relat_1(esk3_0)=k1_xboole_0|esk1_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])])).
cnf(c_0_80, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_81, plain, (r1_tarski(k6_partfun1(X1),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_82, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_42, c_0_72])).
cnf(c_0_83, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_84, negated_conjecture, (v1_xboole_0(esk1_0)|~v1_xboole_0(k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_67]), c_0_68]), c_0_76])])).
fof(c_0_85, plain, ![X70, X71]:~r2_hidden(X70,k2_zfmisc_1(X70,X71)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[t152_zfmisc_1])])).
cnf(c_0_86, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_77, c_0_49])).
cnf(c_0_87, plain, (r1_tarski(k2_relat_1(esk7_4(X1,X2,X3,X4)),X2)|~r2_hidden(X4,X3)|X3!=k1_funct_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
fof(c_0_88, plain, ![X21, X22, X23]:(~v1_relat_1(X23)|(~r1_tarski(k1_relat_1(X23),X21)|~r1_tarski(k2_relat_1(X23),X22)|m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(X21,X22))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_relset_1])])).
cnf(c_0_89, negated_conjecture, (esk1_0!=k1_xboole_0|~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79]), c_0_67]), c_0_62]), c_0_63]), c_0_68])])).
cnf(c_0_90, plain, (r2_hidden(esk6_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_91, plain, (r2_hidden(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,k6_partfun1(X2))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_92, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_93, negated_conjecture, (v1_xboole_0(esk1_0)|esk1_0!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_79]), c_0_63])])).
cnf(c_0_94, plain, (~r2_hidden(X1,k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_95, plain, (r2_hidden(esk10_3(X1,X2,X3),k1_relat_1(X1))|~r2_hidden(X3,X2)|X2!=k2_relat_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_96, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_82, c_0_86])).
cnf(c_0_97, plain, (r1_tarski(k2_relat_1(esk7_4(X1,X2,k1_funct_2(X1,X2),X3)),X2)|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(er,[status(thm)],[c_0_87])).
cnf(c_0_98, negated_conjecture, (~v1_funct_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_99, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_relat_1(X1)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_100, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_101, negated_conjecture, (r1_tarski(esk3_0,X1)|esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_102, plain, (r1_tarski(k6_partfun1(X1),X2)|r2_hidden(esk6_2(k6_partfun1(X1),X2),k2_zfmisc_1(X1,X1))), inference(spm,[status(thm)],[c_0_91, c_0_90])).
cnf(c_0_103, negated_conjecture, (k2_zfmisc_1(X1,esk1_0)=k1_xboole_0|esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_104, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_94, c_0_62])).
cnf(c_0_105, plain, (r2_hidden(esk10_3(X1,k2_relat_1(X1),X2),k1_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_95])).
cnf(c_0_106, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_relat_1(esk7_4(X1,X2,k1_funct_2(X1,X2),X3)))|~r2_hidden(X3,k1_funct_2(X1,X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_36]), c_0_46]), c_0_39])).
cnf(c_0_107, plain, (k2_relat_1(esk7_4(X1,k1_xboole_0,k1_funct_2(X1,k1_xboole_0),X2))=k1_xboole_0|~r2_hidden(X2,k1_funct_2(X1,k1_xboole_0))), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_108, negated_conjecture, (~r1_tarski(k1_relat_1(esk3_0),esk1_0)|~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~v1_relat_1(esk3_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~v1_funct_1(esk3_0)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
fof(c_0_109, plain, ![X32, X33, X34]:((v1_funct_1(X34)|(~v1_funct_1(X34)|~v1_partfun1(X34,X32))|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))&(v1_funct_2(X34,X32,X33)|(~v1_funct_1(X34)|~v1_partfun1(X34,X32))|~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])])).
cnf(c_0_110, negated_conjecture, (X1=esk3_0|esk1_0!=k1_xboole_0|~r1_tarski(X1,esk3_0)), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_111, negated_conjecture, (r1_tarski(k6_partfun1(esk1_0),X1)|esk1_0!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_104])).
cnf(c_0_112, negated_conjecture, (r2_hidden(esk10_3(esk3_0,k2_relat_1(esk3_0),X1),esk1_0)|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_67]), c_0_68]), c_0_76])])).
cnf(c_0_113, plain, (v1_xboole_0(X1)|~r2_hidden(X2,k1_funct_2(X1,k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_63])])).
cnf(c_0_114, negated_conjecture, (~r1_tarski(k1_relat_1(esk3_0),esk1_0)|~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_108, c_0_76])]), c_0_68])])).
cnf(c_0_115, plain, (r1_tarski(k2_relat_1(X1),X2)|~r2_hidden(X1,k1_funct_2(X3,X2))), inference(spm,[status(thm)],[c_0_97, c_0_37])).
cnf(c_0_116, plain, (v1_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_117, plain, (v1_partfun1(k6_partfun1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_118, negated_conjecture, (k6_partfun1(esk1_0)=esk3_0|esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_110, c_0_111])).
fof(c_0_119, plain, ![X74]:((v1_funct_1(X74)|(~v1_relat_1(X74)|~v1_funct_1(X74)))&((v1_funct_2(X74,k1_relat_1(X74),k2_relat_1(X74))|(~v1_relat_1(X74)|~v1_funct_1(X74)))&(m1_subset_1(X74,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X74),k2_relat_1(X74))))|(~v1_relat_1(X74)|~v1_funct_1(X74))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
cnf(c_0_120, negated_conjecture, (~v1_xboole_0(esk1_0)|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_42, c_0_112])).
cnf(c_0_121, plain, (k1_funct_2(X1,k1_xboole_0)=k1_xboole_0|v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_113, c_0_72])).
cnf(c_0_122, negated_conjecture, (~r1_tarski(k2_relat_1(esk3_0),esk2_0)|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_67]), c_0_74])])).
cnf(c_0_123, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),esk2_0)), inference(spm,[status(thm)],[c_0_115, c_0_53])).
cnf(c_0_124, plain, (v1_funct_2(X1,X2,X3)|~r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_partfun1(X1,X2)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_116, c_0_49])).
cnf(c_0_125, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_117, c_0_118])).
fof(c_0_126, plain, ![X24, X25, X26]:((v1_funct_1(X26)|(~v1_funct_1(X26)|~v1_funct_2(X26,X24,X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))|v1_xboole_0(X25))&(v1_partfun1(X26,X24)|(~v1_funct_1(X26)|~v1_funct_2(X26,X24,X25))|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25)))|v1_xboole_0(X25))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_funct_2])])])])])).
cnf(c_0_127, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_128, plain, (v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_129, plain, (r2_hidden(X2,X5)|~v1_relat_1(X1)|~v1_funct_1(X1)|X2!=X1|k1_relat_1(X1)!=X3|~r1_tarski(k2_relat_1(X1),X4)|X5!=k1_funct_2(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_130, negated_conjecture, (k1_funct_2(esk1_0,k1_xboole_0)=k1_xboole_0|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_131, negated_conjecture, (~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_122, c_0_123])])).
cnf(c_0_132, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,X1)|esk1_0!=k1_xboole_0), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_76])]), c_0_101])).
cnf(c_0_133, plain, (v1_partfun1(X1,X2)|v1_xboole_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_126])).
cnf(c_0_134, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k2_relat_1(esk3_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_67]), c_0_68]), c_0_76])])).
cnf(c_0_135, negated_conjecture, (v1_funct_2(esk3_0,esk1_0,k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_67]), c_0_68]), c_0_76])])).
cnf(c_0_136, plain, (r2_hidden(X1,k1_funct_2(k1_relat_1(X1),X2))|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_129])])])).
cnf(c_0_137, plain, (k1_relat_1(X1)=k1_xboole_0|~v1_xboole_0(k2_relat_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_59, c_0_72])).
cnf(c_0_138, negated_conjecture, (k1_funct_2(esk1_0,k1_xboole_0)=k1_xboole_0|v1_xboole_0(k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_130, c_0_60])).
cnf(c_0_139, negated_conjecture, (esk1_0!=k1_xboole_0), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_140, plain, (k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_141, negated_conjecture, (~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))|~v1_funct_2(esk3_0,esk1_0,esk2_0)|~v1_funct_1(esk3_0)), inference(spm,[status(thm)],[c_0_98, c_0_49])).
cnf(c_0_142, negated_conjecture, (v1_partfun1(esk3_0,esk1_0)|v1_xboole_0(k2_relat_1(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_135]), c_0_76])])).
cnf(c_0_143, negated_conjecture, (r2_hidden(esk3_0,k1_funct_2(esk1_0,X1))|~r1_tarski(k2_relat_1(esk3_0),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_67]), c_0_68]), c_0_76])])).
cnf(c_0_144, negated_conjecture, (k1_funct_2(esk1_0,k1_xboole_0)=k1_xboole_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138]), c_0_67]), c_0_68]), c_0_76])]), c_0_139])).
cnf(c_0_145, plain, (k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)|~v1_funct_1(X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_140]), c_0_104])).
cnf(c_0_146, negated_conjecture, (~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))|~v1_funct_2(esk3_0,esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_141, c_0_76])])).
cnf(c_0_147, negated_conjecture, (v1_xboole_0(k2_relat_1(esk3_0))|v1_funct_2(esk3_0,esk1_0,X1)|~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_142]), c_0_76])])).
cnf(c_0_148, negated_conjecture, (~r1_tarski(k2_relat_1(esk3_0),k1_xboole_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_104])).
cnf(c_0_149, plain, (r1_tarski(k2_relat_1(X1),X2)|k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_145, c_0_90])).
cnf(c_0_150, negated_conjecture, (v1_xboole_0(k2_relat_1(esk3_0))|~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_151, negated_conjecture, (k2_relat_1(esk3_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_68]), c_0_76])])).
cnf(c_0_152, negated_conjecture, (~r1_tarski(esk3_0,k2_zfmisc_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_150]), c_0_151])).
cnf(c_0_153, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X3)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_70, c_0_99])).
cnf(c_0_154, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_67]), c_0_74]), c_0_123]), c_0_68])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 155
# Proof object clause steps            : 111
# Proof object formula steps           : 44
# Proof object conjectures             : 42
# Proof object clause conjectures      : 39
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 35
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 62
# Proof object simplifying inferences  : 79
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 53
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 50
# Processed clauses                    : 4947
# ...of these trivial                  : 10
# ...subsumed                          : 3873
# ...remaining for further processing  : 1064
# Other redundant clauses eliminated   : 65
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 97
# Backward-rewritten                   : 30
# Generated clauses                    : 20619
# ...of the previous two non-trivial   : 18163
# Contextual simplify-reflections      : 66
# Paramodulations                      : 20555
# Factorizations                       : 0
# Equation resolutions                 : 65
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 872
#    Positive orientable unit clauses  : 27
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 831
# Current number of unprocessed clauses: 12991
# ...number of literals in the above   : 56441
# Current number of archived formulas  : 0
# Current number of archived clauses   : 178
# Clause-clause subsumption calls (NU) : 130758
# Rec. Clause-clause subsumption calls : 79856
# Non-unit clause-clause subsumptions  : 2135
# Unit Clause-clause subsumption calls : 988
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 273180

# -------------------------------------------------
# User time                : 0.220 s
# System time              : 0.014 s
# Total time               : 0.233 s
# Maximum resident set size: 28660 pages
