# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_AE_CS_SP_PI_S0a
# and selection function SelectNegativeLiterals.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d7_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(X2=k4_relat_1(X1)<=>![X3, X4]:(r2_hidden(k4_tarski(X3,X4),X2)<=>r2_hidden(k4_tarski(X4,X3),X1))))), file('data/MPTP/proofs/f060729537a04a9e950f248c33ca37f4.E_input', d7_relat_1)).
fof(dt_k4_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>v1_relat_1(k4_relat_1(X1))), file('data/MPTP/proofs/f060729537a04a9e950f248c33ca37f4.E_input', dt_k4_relat_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/f060729537a04a9e950f248c33ca37f4.E_input', d4_relat_1)).
fof(t37_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k4_relat_1(X1))&k1_relat_1(X1)=k2_relat_1(k4_relat_1(X1)))), file('data/MPTP/proofs/f060729537a04a9e950f248c33ca37f4.E_input', t37_relat_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/f060729537a04a9e950f248c33ca37f4.E_input', d5_relat_1)).
fof(c_0_5, plain, ![X5, X6, X7, X8, X9, X10]:(((~r2_hidden(k4_tarski(X7,X8),X6)|r2_hidden(k4_tarski(X8,X7),X5)|X6!=k4_relat_1(X5)|~v1_relat_1(X6)|~v1_relat_1(X5))&(~r2_hidden(k4_tarski(X10,X9),X5)|r2_hidden(k4_tarski(X9,X10),X6)|X6!=k4_relat_1(X5)|~v1_relat_1(X6)|~v1_relat_1(X5)))&((~r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|~r2_hidden(k4_tarski(esk2_2(X5,X6),esk1_2(X5,X6)),X5)|X6=k4_relat_1(X5)|~v1_relat_1(X6)|~v1_relat_1(X5))&(r2_hidden(k4_tarski(esk1_2(X5,X6),esk2_2(X5,X6)),X6)|r2_hidden(k4_tarski(esk2_2(X5,X6),esk1_2(X5,X6)),X5)|X6=k4_relat_1(X5)|~v1_relat_1(X6)|~v1_relat_1(X5)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_relat_1])])])])])])).
fof(c_0_6, plain, ![X13]:(~v1_relat_1(X13)|v1_relat_1(k4_relat_1(X13))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_relat_1])])).
cnf(c_0_7, plain, (r2_hidden(k4_tarski(X2,X1),X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X3!=k4_relat_1(X4)|~v1_relat_1(X3)|~v1_relat_1(X4)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (v1_relat_1(k4_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, plain, ![X26, X27, X28, X30, X31, X32, X33, X35]:(((~r2_hidden(X28,X27)|r2_hidden(k4_tarski(X28,esk7_3(X26,X27,X28)),X26)|X27!=k1_relat_1(X26))&(~r2_hidden(k4_tarski(X30,X31),X26)|r2_hidden(X30,X27)|X27!=k1_relat_1(X26)))&((~r2_hidden(esk8_2(X32,X33),X33)|~r2_hidden(k4_tarski(esk8_2(X32,X33),X35),X32)|X33=k1_relat_1(X32))&(r2_hidden(esk8_2(X32,X33),X33)|r2_hidden(k4_tarski(esk8_2(X32,X33),esk9_2(X32,X33)),X32)|X33=k1_relat_1(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
fof(c_0_10, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k4_relat_1(X1))&k1_relat_1(X1)=k2_relat_1(k4_relat_1(X1))))), inference(assume_negation,[status(cth)],[t37_relat_1])).
fof(c_0_11, plain, ![X14, X15, X16, X18, X19, X20, X21, X23]:(((~r2_hidden(X16,X15)|r2_hidden(k4_tarski(esk3_3(X14,X15,X16),X16),X14)|X15!=k2_relat_1(X14))&(~r2_hidden(k4_tarski(X19,X18),X14)|r2_hidden(X18,X15)|X15!=k2_relat_1(X14)))&((~r2_hidden(esk4_2(X20,X21),X21)|~r2_hidden(k4_tarski(X23,esk4_2(X20,X21)),X20)|X21=k2_relat_1(X20))&(r2_hidden(esk4_2(X20,X21),X21)|r2_hidden(k4_tarski(esk5_2(X20,X21),esk4_2(X20,X21)),X20)|X21=k2_relat_1(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_12, plain, (r2_hidden(k4_tarski(X1,X2),X3)|~r2_hidden(k4_tarski(X2,X1),k4_relat_1(X3))|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_7]), c_0_8])).
cnf(c_0_13, plain, (r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(k4_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),X1)|X2=k1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, negated_conjecture, (v1_relat_1(esk6_0)&(k2_relat_1(esk6_0)!=k1_relat_1(k4_relat_1(esk6_0))|k1_relat_1(esk6_0)!=k2_relat_1(k4_relat_1(esk6_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_10])])])).
cnf(c_0_15, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, plain, (X1=k1_relat_1(k4_relat_1(X2))|r2_hidden(k4_tarski(esk9_2(k4_relat_1(X2),X1),esk8_2(k4_relat_1(X2),X1)),X2)|r2_hidden(esk8_2(k4_relat_1(X2),X1),X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, negated_conjecture, (v1_relat_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_15])).
cnf(c_0_19, negated_conjecture, (X1=k1_relat_1(k4_relat_1(esk6_0))|r2_hidden(k4_tarski(esk9_2(k4_relat_1(esk6_0),X1),esk8_2(k4_relat_1(esk6_0),X1)),esk6_0)|r2_hidden(esk8_2(k4_relat_1(esk6_0),X1),X1)), inference(spm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20, plain, (r2_hidden(esk4_2(X1,X2),X2)|r2_hidden(k4_tarski(esk5_2(X1,X2),esk4_2(X1,X2)),X1)|X2=k2_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_21, plain, (r2_hidden(k4_tarski(esk3_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_22, negated_conjecture, (X1=k1_relat_1(k4_relat_1(esk6_0))|r2_hidden(esk8_2(k4_relat_1(esk6_0),X1),k2_relat_1(esk6_0))|r2_hidden(esk8_2(k4_relat_1(esk6_0),X1),X1)), inference(spm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_23, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24, plain, (X1=k2_relat_1(k4_relat_1(X2))|r2_hidden(k4_tarski(esk4_2(k4_relat_1(X2),X1),esk5_2(k4_relat_1(X2),X1)),X2)|r2_hidden(esk4_2(k4_relat_1(X2),X1),X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_12, c_0_20])).
cnf(c_0_25, plain, (r2_hidden(k4_tarski(X2,X1),X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k4_relat_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_26, plain, (r2_hidden(k4_tarski(esk3_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (k1_relat_1(k4_relat_1(esk6_0))=k2_relat_1(esk6_0)|r2_hidden(esk8_2(k4_relat_1(esk6_0),k2_relat_1(esk6_0)),k2_relat_1(esk6_0))), inference(ef,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_23])).
cnf(c_0_29, negated_conjecture, (X1=k2_relat_1(k4_relat_1(esk6_0))|r2_hidden(k4_tarski(esk4_2(k4_relat_1(esk6_0),X1),esk5_2(k4_relat_1(esk6_0),X1)),esk6_0)|r2_hidden(esk4_2(k4_relat_1(esk6_0),X1),X1)), inference(spm,[status(thm)],[c_0_24, c_0_17])).
cnf(c_0_30, plain, (r2_hidden(k4_tarski(X1,X2),k4_relat_1(X3))|~r2_hidden(k4_tarski(X2,X1),X3)|~v1_relat_1(X3)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_25]), c_0_8])).
cnf(c_0_31, negated_conjecture, (k1_relat_1(k4_relat_1(esk6_0))=k2_relat_1(esk6_0)|r2_hidden(k4_tarski(esk3_3(esk6_0,k2_relat_1(esk6_0),esk8_2(k4_relat_1(esk6_0),k2_relat_1(esk6_0))),esk8_2(k4_relat_1(esk6_0),k2_relat_1(esk6_0))),esk6_0)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_32, plain, (r2_hidden(k4_tarski(X1,esk7_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_33, negated_conjecture, (X1=k2_relat_1(k4_relat_1(esk6_0))|r2_hidden(esk4_2(k4_relat_1(esk6_0),X1),k1_relat_1(esk6_0))|r2_hidden(esk4_2(k4_relat_1(esk6_0),X1),X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (X2=k1_relat_1(X1)|~r2_hidden(esk8_2(X1,X2),X2)|~r2_hidden(k4_tarski(esk8_2(X1,X2),X3),X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_35, negated_conjecture, (k1_relat_1(k4_relat_1(esk6_0))=k2_relat_1(esk6_0)|r2_hidden(k4_tarski(esk8_2(k4_relat_1(esk6_0),k2_relat_1(esk6_0)),esk3_3(esk6_0,k2_relat_1(esk6_0),esk8_2(k4_relat_1(esk6_0),k2_relat_1(esk6_0)))),k4_relat_1(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_17])])).
cnf(c_0_36, plain, (r2_hidden(k4_tarski(X1,esk7_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, (k2_relat_1(k4_relat_1(esk6_0))=k1_relat_1(esk6_0)|r2_hidden(esk4_2(k4_relat_1(esk6_0),k1_relat_1(esk6_0)),k1_relat_1(esk6_0))), inference(ef,[status(thm)],[c_0_33])).
cnf(c_0_38, negated_conjecture, (k2_relat_1(esk6_0)!=k1_relat_1(k4_relat_1(esk6_0))|k1_relat_1(esk6_0)!=k2_relat_1(k4_relat_1(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_39, negated_conjecture, (k1_relat_1(k4_relat_1(esk6_0))=k2_relat_1(esk6_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_27])).
cnf(c_0_40, negated_conjecture, (k2_relat_1(k4_relat_1(esk6_0))=k1_relat_1(esk6_0)|r2_hidden(k4_tarski(esk4_2(k4_relat_1(esk6_0),k1_relat_1(esk6_0)),esk7_3(esk6_0,k1_relat_1(esk6_0),esk4_2(k4_relat_1(esk6_0),k1_relat_1(esk6_0)))),esk6_0)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41, negated_conjecture, (k2_relat_1(k4_relat_1(esk6_0))!=k1_relat_1(esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39])])).
cnf(c_0_42, negated_conjecture, (r2_hidden(k4_tarski(esk4_2(k4_relat_1(esk6_0),k1_relat_1(esk6_0)),esk7_3(esk6_0,k1_relat_1(esk6_0),esk4_2(k4_relat_1(esk6_0),k1_relat_1(esk6_0)))),esk6_0)), inference(sr,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_43, plain, (X2=k2_relat_1(X1)|~r2_hidden(esk4_2(X1,X2),X2)|~r2_hidden(k4_tarski(X3,esk4_2(X1,X2)),X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_44, negated_conjecture, (r2_hidden(k4_tarski(esk7_3(esk6_0,k1_relat_1(esk6_0),esk4_2(k4_relat_1(esk6_0),k1_relat_1(esk6_0))),esk4_2(k4_relat_1(esk6_0),k1_relat_1(esk6_0))),k4_relat_1(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_42]), c_0_17])])).
cnf(c_0_45, negated_conjecture, (r2_hidden(esk4_2(k4_relat_1(esk6_0),k1_relat_1(esk6_0)),k1_relat_1(esk6_0))), inference(sr,[status(thm)],[c_0_37, c_0_41])).
cnf(c_0_46, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45])]), c_0_41]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 47
# Proof object clause steps            : 36
# Proof object formula steps           : 11
# Proof object conjectures             : 20
# Proof object clause conjectures      : 17
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 14
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 15
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 15
# Processed clauses                    : 1508
# ...of these trivial                  : 2
# ...subsumed                          : 530
# ...remaining for further processing  : 976
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 35
# Backward-rewritten                   : 144
# Generated clauses                    : 11059
# ...of the previous two non-trivial   : 11146
# Contextual simplify-reflections      : 7
# Paramodulations                      : 10704
# Factorizations                       : 22
# Equation resolutions                 : 329
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 793
#    Positive orientable unit clauses  : 70
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 721
# Current number of unprocessed clauses: 9620
# ...number of literals in the above   : 33848
# Current number of archived formulas  : 0
# Current number of archived clauses   : 183
# Clause-clause subsumption calls (NU) : 69842
# Rec. Clause-clause subsumption calls : 10100
# Non-unit clause-clause subsumptions  : 572
# Unit Clause-clause subsumption calls : 1360
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1895
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1882866

# -------------------------------------------------
# User time                : 0.349 s
# System time              : 0.000 s
# Total time               : 0.349 s
# Maximum resident set size: 28676 pages
