# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t69_enumset1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t70_enumset1)).
fof(t12_ordinal1, conjecture, ![X1, X2]:(k1_ordinal1(X1)=k1_ordinal1(X2)=>X1=X2), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t12_ordinal1)).
fof(t10_ordinal1, axiom, ![X1]:r2_hidden(X1,k1_ordinal1(X1)), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t10_ordinal1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t72_enumset1)).
fof(t5_xboole_0, axiom, ![X1, X2, X3]:~((r1_xboole_0(X1,X2)&(r2_hidden(X3,k2_xboole_0(X1,X2))&(~((r2_hidden(X3,X1)&~(r2_hidden(X3,X2))))&~((r2_hidden(X3,X2)&~(r2_hidden(X3,X1)))))))), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t5_xboole_0)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', d10_xboole_0)).
fof(t12_setfam_1, axiom, ![X1, X2]:k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t12_setfam_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', d1_tarski)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t83_xboole_1)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t65_zfmisc_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t7_ordinal1)).
fof(t51_xboole_1, axiom, ![X1, X2]:k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t51_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t48_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', commutativity_k2_tarski)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t43_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t7_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', t28_xboole_1)).
fof(antisymmetry_r2_hidden, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', antisymmetry_r2_hidden)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', l1_zfmisc_1)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/f17b20cbe6a443118725630157db992a.E_input', d5_xboole_0)).
fof(c_0_24, plain, ![X5]:k1_ordinal1(X5)=k2_xboole_0(X5,k1_tarski(X5)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_25, plain, ![X6]:k2_tarski(X6,X6)=k1_tarski(X6), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_26, plain, ![X7, X8]:k3_tarski(k2_tarski(X7,X8))=k2_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_27, plain, ![X9, X10]:k1_enumset1(X9,X9,X10)=k2_tarski(X9,X10), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_28, negated_conjecture, ~(![X1, X2]:(k1_ordinal1(X1)=k1_ordinal1(X2)=>X1=X2)), inference(assume_negation,[status(cth)],[t12_ordinal1])).
fof(c_0_29, plain, ![X13]:r2_hidden(X13,k1_ordinal1(X13)), inference(variable_rename,[status(thm)],[t10_ordinal1])).
cnf(c_0_30, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_32, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_33, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_34, plain, ![X14, X15, X16]:k2_enumset1(X14,X14,X15,X16)=k1_enumset1(X14,X15,X16), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_35, plain, ![X17, X18, X19, X20]:k3_enumset1(X17,X17,X18,X19,X20)=k2_enumset1(X17,X18,X19,X20), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_36, negated_conjecture, (k1_ordinal1(esk1_0)=k1_ordinal1(esk2_0)&esk1_0!=esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_28])])])).
fof(c_0_37, plain, ![X21, X22, X23]:(((r2_hidden(X23,X22)|r2_hidden(X23,X21)|~r2_hidden(X23,k2_xboole_0(X21,X22))|~r1_xboole_0(X21,X22))&(~r2_hidden(X23,X21)|r2_hidden(X23,X21)|~r2_hidden(X23,k2_xboole_0(X21,X22))|~r1_xboole_0(X21,X22)))&((r2_hidden(X23,X22)|~r2_hidden(X23,X22)|~r2_hidden(X23,k2_xboole_0(X21,X22))|~r1_xboole_0(X21,X22))&(~r2_hidden(X23,X21)|~r2_hidden(X23,X22)|~r2_hidden(X23,k2_xboole_0(X21,X22))|~r1_xboole_0(X21,X22)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t5_xboole_0])])])])).
cnf(c_0_38, plain, (r2_hidden(X1,k1_ordinal1(X1))), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_39, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_40, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_43, negated_conjecture, (k1_ordinal1(esk1_0)=k1_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X3,X2))|~r1_xboole_0(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, plain, (r2_hidden(X1,k3_tarski(k3_enumset1(X1,X1,X1,X1,k3_enumset1(X1,X1,X1,X1,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39]), c_0_33]), c_0_40]), c_0_41]), c_0_41]), c_0_42]), c_0_42])).
cnf(c_0_46, negated_conjecture, (k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))=k3_tarski(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_39]), c_0_39]), c_0_33]), c_0_33]), c_0_40]), c_0_40]), c_0_41]), c_0_41]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_42]), c_0_42])).
fof(c_0_47, plain, ![X24, X25]:(((r1_tarski(X24,X25)|X24!=X25)&(r1_tarski(X25,X24)|X24!=X25))&(~r1_tarski(X24,X25)|~r1_tarski(X25,X24)|X24=X25)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
fof(c_0_48, plain, ![X26, X27]:k1_setfam_1(k2_tarski(X26,X27))=k3_xboole_0(X26,X27), inference(variable_rename,[status(thm)],[t12_setfam_1])).
fof(c_0_49, plain, ![X28, X29, X30, X31, X32, X33]:(((~r2_hidden(X30,X29)|X30=X28|X29!=k1_tarski(X28))&(X31!=X28|r2_hidden(X31,X29)|X29!=k1_tarski(X28)))&((~r2_hidden(esk3_2(X32,X33),X33)|esk3_2(X32,X33)!=X32|X33=k1_tarski(X32))&(r2_hidden(esk3_2(X32,X33),X33)|esk3_2(X32,X33)=X32|X33=k1_tarski(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_50, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X2)|~r1_xboole_0(X3,X2)|~r2_hidden(X1,k3_tarski(k3_enumset1(X3,X3,X3,X3,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_51, negated_conjecture, (r2_hidden(esk1_0,k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
fof(c_0_52, plain, ![X35, X36]:((~r1_xboole_0(X35,X36)|k4_xboole_0(X35,X36)=X35)&(k4_xboole_0(X35,X36)!=X35|r1_xboole_0(X35,X36))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_53, plain, ![X37, X38]:((k4_xboole_0(X37,k1_tarski(X38))!=X37|~r2_hidden(X38,X37))&(r2_hidden(X38,X37)|k4_xboole_0(X37,k1_tarski(X38))=X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
fof(c_0_54, plain, ![X39, X40]:(~r2_hidden(X39,X40)|~r1_tarski(X40,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_55, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_56, plain, ![X44, X45]:k2_xboole_0(k3_xboole_0(X44,X45),k4_xboole_0(X44,X45))=X44, inference(variable_rename,[status(thm)],[t51_xboole_1])).
cnf(c_0_57, plain, (k1_setfam_1(k2_tarski(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_58, plain, ![X46, X47]:k4_xboole_0(X46,k4_xboole_0(X46,X47))=k3_xboole_0(X46,X47), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_59, plain, ![X48, X49]:k2_tarski(X48,X49)=k2_tarski(X49,X48), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_60, plain, ![X41, X42, X43]:(~r1_tarski(X41,k2_xboole_0(X42,X43))|r1_tarski(k4_xboole_0(X41,X42),X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
fof(c_0_61, plain, ![X50, X51]:r1_tarski(X50,k2_xboole_0(X50,X51)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_62, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_63, negated_conjecture, (r2_hidden(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|r2_hidden(esk1_0,esk2_0)|~r1_xboole_0(esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_64, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_65, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_66, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_67, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_55])).
fof(c_0_68, plain, ![X52, X53]:(~r1_tarski(X52,X53)|k3_xboole_0(X52,X53)=X52), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_69, plain, (k2_xboole_0(k3_xboole_0(X1,X2),k4_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_70, plain, (k1_setfam_1(k1_enumset1(X1,X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_57, c_0_33])).
cnf(c_0_71, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_72, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_73, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_74, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_75, plain, (X1=X3|X2!=k3_enumset1(X3,X3,X3,X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_31]), c_0_33]), c_0_41]), c_0_42])).
cnf(c_0_76, negated_conjecture, (r2_hidden(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|r2_hidden(esk1_0,esk2_0)|k4_xboole_0(esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))!=esk2_0), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_77, plain, (k4_xboole_0(X2,k3_enumset1(X1,X1,X1,X1,X1))=X2|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_31]), c_0_33]), c_0_41]), c_0_42])).
cnf(c_0_78, plain, (~r2_hidden(X1,X1)), inference(spm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_79, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_80, plain, (k3_tarski(k3_enumset1(k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)),k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)),k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)),k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2)),k4_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_40]), c_0_70]), c_0_70]), c_0_41]), c_0_41]), c_0_41]), c_0_42]), c_0_42]), c_0_42]), c_0_42])).
cnf(c_0_81, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_70]), c_0_41]), c_0_42])).
cnf(c_0_82, plain, (k3_enumset1(X1,X1,X1,X1,X2)=k3_enumset1(X2,X2,X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_33]), c_0_33]), c_0_41]), c_0_41]), c_0_42]), c_0_42])).
cnf(c_0_83, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k3_enumset1(X2,X2,X2,X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_40]), c_0_41]), c_0_42])).
cnf(c_0_84, plain, (r1_tarski(X1,k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_40]), c_0_41]), c_0_42])).
fof(c_0_85, plain, ![X54, X55]:(~r2_hidden(X54,X55)|~r2_hidden(X55,X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_hidden])])])).
cnf(c_0_86, plain, (X1=X2|~r2_hidden(X1,k3_enumset1(X2,X2,X2,X2,X2))), inference(er,[status(thm)],[c_0_75])).
cnf(c_0_87, negated_conjecture, (r2_hidden(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))|r2_hidden(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])).
cnf(c_0_88, negated_conjecture, (esk1_0!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_89, plain, (k1_setfam_1(k3_enumset1(X1,X1,X1,X1,X2))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_70]), c_0_41]), c_0_42])).
cnf(c_0_90, plain, (k3_tarski(k3_enumset1(k4_xboole_0(X1,X2),k4_xboole_0(X1,X2),k4_xboole_0(X1,X2),k4_xboole_0(X1,X2),k4_xboole_0(X1,k4_xboole_0(X1,X2))))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_81]), c_0_81]), c_0_81]), c_0_81]), c_0_82])).
cnf(c_0_91, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k3_enumset1(X3,X3,X3,X3,X2)))), inference(spm,[status(thm)],[c_0_83, c_0_82])).
cnf(c_0_92, negated_conjecture, (r1_tarski(esk1_0,k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(spm,[status(thm)],[c_0_84, c_0_46])).
cnf(c_0_93, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_94, negated_conjecture, (r2_hidden(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])).
cnf(c_0_95, negated_conjecture, (r1_tarski(k4_xboole_0(X1,esk1_0),k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))|~r1_tarski(X1,k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0))))), inference(spm,[status(thm)],[c_0_83, c_0_46])).
fof(c_0_96, plain, ![X56, X57]:((~r1_tarski(k1_tarski(X56),X57)|r2_hidden(X56,X57))&(~r2_hidden(X56,X57)|r1_tarski(k1_tarski(X56),X57))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
fof(c_0_97, plain, ![X58, X59, X60, X61, X62, X63, X64, X65]:((((r2_hidden(X61,X58)|~r2_hidden(X61,X60)|X60!=k4_xboole_0(X58,X59))&(~r2_hidden(X61,X59)|~r2_hidden(X61,X60)|X60!=k4_xboole_0(X58,X59)))&(~r2_hidden(X62,X58)|r2_hidden(X62,X59)|r2_hidden(X62,X60)|X60!=k4_xboole_0(X58,X59)))&((~r2_hidden(esk4_3(X63,X64,X65),X65)|(~r2_hidden(esk4_3(X63,X64,X65),X63)|r2_hidden(esk4_3(X63,X64,X65),X64))|X65=k4_xboole_0(X63,X64))&((r2_hidden(esk4_3(X63,X64,X65),X63)|r2_hidden(esk4_3(X63,X64,X65),X65)|X65=k4_xboole_0(X63,X64))&(~r2_hidden(esk4_3(X63,X64,X65),X64)|r2_hidden(esk4_3(X63,X64,X65),X65)|X65=k4_xboole_0(X63,X64))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_98, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|~r1_tarski(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_89, c_0_82]), c_0_81])).
cnf(c_0_99, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X1|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_81, c_0_89])).
cnf(c_0_100, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_84, c_0_90])).
cnf(c_0_101, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)),esk2_0)), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_102, negated_conjecture, (~r2_hidden(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_103, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_104, negated_conjecture, (r1_tarski(k4_xboole_0(esk2_0,esk1_0),k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_95, c_0_84])).
cnf(c_0_105, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_106, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_107, plain, (k4_xboole_0(X1,X1)=k4_xboole_0(X1,X2)|~r1_tarski(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_100])])).
cnf(c_0_108, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_77]), c_0_102])).
cnf(c_0_109, negated_conjecture, (k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k4_xboole_0(esk2_0,esk1_0)|~r1_tarski(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0),k4_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_110, plain, (r1_tarski(k3_enumset1(X1,X1,X1,X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_31]), c_0_33]), c_0_41]), c_0_42])).
cnf(c_0_111, plain, (r2_hidden(X1,k4_xboole_0(X2,X3))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_106])).
cnf(c_0_112, negated_conjecture, (k4_xboole_0(esk1_0,esk1_0)=k4_xboole_0(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_113, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_81])).
cnf(c_0_114, negated_conjecture, (k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,esk1_0)=k4_xboole_0(esk2_0,esk1_0)|~r2_hidden(esk1_0,k4_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_109, c_0_110])).
cnf(c_0_115, plain, (r2_hidden(k4_xboole_0(X1,X2),X2)|~r2_hidden(k4_xboole_0(X1,X2),X1)), inference(spm,[status(thm)],[c_0_78, c_0_111])).
cnf(c_0_116, negated_conjecture, (k4_xboole_0(esk2_0,k4_xboole_0(esk2_0,esk1_0))=esk1_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_112]), c_0_67])]), c_0_113])).
cnf(c_0_117, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_118, plain, (~r2_hidden(k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)),X1)), inference(spm,[status(thm)],[c_0_66, c_0_84])).
cnf(c_0_119, negated_conjecture, (k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))=k3_tarski(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,k4_xboole_0(esk2_0,esk1_0)))|~r2_hidden(esk1_0,k4_xboole_0(esk2_0,esk1_0))), inference(spm,[status(thm)],[c_0_46, c_0_114])).
cnf(c_0_120, negated_conjecture, (r2_hidden(esk1_0,k4_xboole_0(esk2_0,esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_94])])).
cnf(c_0_121, negated_conjecture, (k3_tarski(k3_enumset1(esk1_0,esk1_0,esk1_0,esk1_0,k4_xboole_0(esk2_0,esk1_0)))=esk2_0), inference(spm,[status(thm)],[c_0_90, c_0_116])).
cnf(c_0_122, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k3_enumset1(X2,X2,X2,X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_117, c_0_31]), c_0_33]), c_0_41]), c_0_42])).
cnf(c_0_123, plain, (~r2_hidden(k3_tarski(k3_enumset1(X1,X1,X1,X1,X2)),X2)), inference(spm,[status(thm)],[c_0_118, c_0_82])).
cnf(c_0_124, negated_conjecture, (k3_tarski(k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,k3_enumset1(esk2_0,esk2_0,esk2_0,esk2_0,esk2_0)))=esk2_0), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_119, c_0_120])]), c_0_121])).
cnf(c_0_125, plain, (r2_hidden(X1,k3_enumset1(X1,X1,X1,X1,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_122])])).
cnf(c_0_126, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_125])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 127
# Proof object clause steps            : 78
# Proof object formula steps           : 49
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 98
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 32
# Processed clauses                    : 707
# ...of these trivial                  : 11
# ...subsumed                          : 447
# ...remaining for further processing  : 249
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 63
# Generated clauses                    : 2227
# ...of the previous two non-trivial   : 1933
# Contextual simplify-reflections      : 4
# Paramodulations                      : 2214
# Factorizations                       : 2
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 136
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 27
#    Non-unit-clauses                  : 72
# Current number of unprocessed clauses: 1245
# ...number of literals in the above   : 3441
# Current number of archived formulas  : 0
# Current number of archived clauses   : 113
# Clause-clause subsumption calls (NU) : 3884
# Rec. Clause-clause subsumption calls : 2357
# Non-unit clause-clause subsumptions  : 131
# Unit Clause-clause subsumption calls : 2287
# Rewrite failures with RHS unbound    : 32
# BW rewrite match attempts            : 164
# BW rewrite match successes           : 72
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 28575

# -------------------------------------------------
# User time                : 0.030 s
# System time              : 0.000 s
# Total time               : 0.030 s
# Maximum resident set size: 28640 pages
