# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_filter_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v11_lattices(X1)&l3_lattices(X1))))&(~(v2_struct_0(X2))&(v10_lattices(X2)&(v11_lattices(X2)&l3_lattices(X2)))))<=>(~(v2_struct_0(k7_filter_1(X1,X2)))&(v10_lattices(k7_filter_1(X1,X2))&(v11_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2)))))))), file('data/MPTP/proofs/f27a38f0efb242aaa18a5325bc4d2b92.E_input', t39_filter_1)).
fof(cc5_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&v17_lattices(X1))=>(~(v2_struct_0(X1))&(v11_lattices(X1)&(v15_lattices(X1)&v16_lattices(X1)))))), file('data/MPTP/proofs/f27a38f0efb242aaa18a5325bc4d2b92.E_input', cc5_lattices)).
fof(t47_filter_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))&(~(v2_struct_0(X2))&(v10_lattices(X2)&(v17_lattices(X2)&l3_lattices(X2)))))<=>(~(v2_struct_0(k7_filter_1(X1,X2)))&(v10_lattices(k7_filter_1(X1,X2))&(v17_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2)))))))), file('data/MPTP/proofs/f27a38f0efb242aaa18a5325bc4d2b92.E_input', t47_filter_1)).
fof(cc6_lattices, axiom, ![X1]:(l3_lattices(X1)=>((~(v2_struct_0(X1))&(v11_lattices(X1)&(v15_lattices(X1)&v16_lattices(X1))))=>(~(v2_struct_0(X1))&v17_lattices(X1)))), file('data/MPTP/proofs/f27a38f0efb242aaa18a5325bc4d2b92.E_input', cc6_lattices)).
fof(dt_k7_filter_1, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l3_lattices(X1))&(~(v2_struct_0(X2))&l3_lattices(X2)))=>(v3_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2)))), file('data/MPTP/proofs/f27a38f0efb242aaa18a5325bc4d2b92.E_input', dt_k7_filter_1)).
fof(t46_filter_1, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v15_lattices(X1)&(v16_lattices(X1)&l3_lattices(X1)))))&(~(v2_struct_0(X2))&(v10_lattices(X2)&(v15_lattices(X2)&(v16_lattices(X2)&l3_lattices(X2))))))<=>(~(v2_struct_0(k7_filter_1(X1,X2)))&(v10_lattices(k7_filter_1(X1,X2))&(v15_lattices(k7_filter_1(X1,X2))&(v16_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2))))))))), file('data/MPTP/proofs/f27a38f0efb242aaa18a5325bc4d2b92.E_input', t46_filter_1)).
fof(c_0_6, plain, ![X2, X1]:(epred1_2(X1,X2)<=>(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v11_lattices(X1)&l3_lattices(X1))))&(~(v2_struct_0(X2))&(v10_lattices(X2)&(v11_lattices(X2)&l3_lattices(X2)))))<=>(~(v2_struct_0(k7_filter_1(X1,X2)))&(v10_lattices(k7_filter_1(X1,X2))&(v11_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2))))))), introduced(definition)).
fof(c_0_7, plain, ![X2, X1]:(epred1_2(X1,X2)=>(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v11_lattices(X1)&l3_lattices(X1))))&(~(v2_struct_0(X2))&(v10_lattices(X2)&(v11_lattices(X2)&l3_lattices(X2)))))<=>(~(v2_struct_0(k7_filter_1(X1,X2)))&(v10_lattices(k7_filter_1(X1,X2))&(v11_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2))))))), inference(split_equiv,[status(thm)],[c_0_6])).
fof(c_0_8, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>epred1_2(X1,X2))), inference(apply_def,[status(thm)],[t39_filter_1, c_0_6])).
fof(c_0_9, plain, ![X13, X14]:(((~v2_struct_0(k7_filter_1(X14,X13))|(v2_struct_0(X14)|(~v10_lattices(X14)|(~v11_lattices(X14)|~l3_lattices(X14)))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v11_lattices(X13)|~l3_lattices(X13)))))|~epred1_2(X14,X13))&((v10_lattices(k7_filter_1(X14,X13))|(v2_struct_0(X14)|(~v10_lattices(X14)|(~v11_lattices(X14)|~l3_lattices(X14)))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v11_lattices(X13)|~l3_lattices(X13)))))|~epred1_2(X14,X13))&((v11_lattices(k7_filter_1(X14,X13))|(v2_struct_0(X14)|(~v10_lattices(X14)|(~v11_lattices(X14)|~l3_lattices(X14)))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v11_lattices(X13)|~l3_lattices(X13)))))|~epred1_2(X14,X13))&(l3_lattices(k7_filter_1(X14,X13))|(v2_struct_0(X14)|(~v10_lattices(X14)|(~v11_lattices(X14)|~l3_lattices(X14)))|(v2_struct_0(X13)|(~v10_lattices(X13)|(~v11_lattices(X13)|~l3_lattices(X13)))))|~epred1_2(X14,X13)))))&(((~v2_struct_0(X14)|(v2_struct_0(k7_filter_1(X14,X13))|(~v10_lattices(k7_filter_1(X14,X13))|(~v11_lattices(k7_filter_1(X14,X13))|~l3_lattices(k7_filter_1(X14,X13)))))|~epred1_2(X14,X13))&((v10_lattices(X14)|(v2_struct_0(k7_filter_1(X14,X13))|(~v10_lattices(k7_filter_1(X14,X13))|(~v11_lattices(k7_filter_1(X14,X13))|~l3_lattices(k7_filter_1(X14,X13)))))|~epred1_2(X14,X13))&((v11_lattices(X14)|(v2_struct_0(k7_filter_1(X14,X13))|(~v10_lattices(k7_filter_1(X14,X13))|(~v11_lattices(k7_filter_1(X14,X13))|~l3_lattices(k7_filter_1(X14,X13)))))|~epred1_2(X14,X13))&(l3_lattices(X14)|(v2_struct_0(k7_filter_1(X14,X13))|(~v10_lattices(k7_filter_1(X14,X13))|(~v11_lattices(k7_filter_1(X14,X13))|~l3_lattices(k7_filter_1(X14,X13)))))|~epred1_2(X14,X13)))))&((~v2_struct_0(X13)|(v2_struct_0(k7_filter_1(X14,X13))|(~v10_lattices(k7_filter_1(X14,X13))|(~v11_lattices(k7_filter_1(X14,X13))|~l3_lattices(k7_filter_1(X14,X13)))))|~epred1_2(X14,X13))&((v10_lattices(X13)|(v2_struct_0(k7_filter_1(X14,X13))|(~v10_lattices(k7_filter_1(X14,X13))|(~v11_lattices(k7_filter_1(X14,X13))|~l3_lattices(k7_filter_1(X14,X13)))))|~epred1_2(X14,X13))&((v11_lattices(X13)|(v2_struct_0(k7_filter_1(X14,X13))|(~v10_lattices(k7_filter_1(X14,X13))|(~v11_lattices(k7_filter_1(X14,X13))|~l3_lattices(k7_filter_1(X14,X13)))))|~epred1_2(X14,X13))&(l3_lattices(X13)|(v2_struct_0(k7_filter_1(X14,X13))|(~v10_lattices(k7_filter_1(X14,X13))|(~v11_lattices(k7_filter_1(X14,X13))|~l3_lattices(k7_filter_1(X14,X13)))))|~epred1_2(X14,X13))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_7])])])])).
fof(c_0_10, plain, ![X3, X4]:(v2_struct_0(X3)|(~v10_lattices(X3)|~l3_lattices(X3))|(v2_struct_0(X4)|(~v10_lattices(X4)|~l3_lattices(X4))|epred1_2(X3,X4))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])).
cnf(c_0_11, plain, (v10_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v10_lattices(X1)|~v11_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~v11_lattices(X2)|~l3_lattices(X2)|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (v2_struct_0(X1)|v2_struct_0(X2)|epred1_2(X1,X2)|~v10_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_13, plain, ![X5]:((~v2_struct_0(X5)|(v2_struct_0(X5)|~v17_lattices(X5))|~l3_lattices(X5))&((v11_lattices(X5)|(v2_struct_0(X5)|~v17_lattices(X5))|~l3_lattices(X5))&((v15_lattices(X5)|(v2_struct_0(X5)|~v17_lattices(X5))|~l3_lattices(X5))&(v16_lattices(X5)|(v2_struct_0(X5)|~v17_lattices(X5))|~l3_lattices(X5))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_lattices])])])])).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v17_lattices(X1)&l3_lattices(X1))))&(~(v2_struct_0(X2))&(v10_lattices(X2)&(v17_lattices(X2)&l3_lattices(X2)))))<=>(~(v2_struct_0(k7_filter_1(X1,X2)))&(v10_lattices(k7_filter_1(X1,X2))&(v17_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2))))))))), inference(assume_negation,[status(cth)],[t47_filter_1])).
cnf(c_0_15, plain, (v10_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v11_lattices(X2)|~v11_lattices(X1)|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (v11_lattices(X1)|v2_struct_0(X1)|~v17_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_17, negated_conjecture, ((~v2_struct_0(esk1_0)&(v10_lattices(esk1_0)&l3_lattices(esk1_0)))&((~v2_struct_0(esk2_0)&(v10_lattices(esk2_0)&l3_lattices(esk2_0)))&((v2_struct_0(esk1_0)|(~v10_lattices(esk1_0)|(~v17_lattices(esk1_0)|~l3_lattices(esk1_0)))|(v2_struct_0(esk2_0)|(~v10_lattices(esk2_0)|(~v17_lattices(esk2_0)|~l3_lattices(esk2_0))))|(v2_struct_0(k7_filter_1(esk1_0,esk2_0))|(~v10_lattices(k7_filter_1(esk1_0,esk2_0))|(~v17_lattices(k7_filter_1(esk1_0,esk2_0))|~l3_lattices(k7_filter_1(esk1_0,esk2_0))))))&((((~v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v2_struct_0(esk1_0))&((v10_lattices(k7_filter_1(esk1_0,esk2_0))|~v2_struct_0(esk1_0))&((v17_lattices(k7_filter_1(esk1_0,esk2_0))|~v2_struct_0(esk1_0))&(l3_lattices(k7_filter_1(esk1_0,esk2_0))|~v2_struct_0(esk1_0)))))&(((~v2_struct_0(k7_filter_1(esk1_0,esk2_0))|v10_lattices(esk1_0))&((v10_lattices(k7_filter_1(esk1_0,esk2_0))|v10_lattices(esk1_0))&((v17_lattices(k7_filter_1(esk1_0,esk2_0))|v10_lattices(esk1_0))&(l3_lattices(k7_filter_1(esk1_0,esk2_0))|v10_lattices(esk1_0)))))&(((~v2_struct_0(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0))&((v10_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0))&((v17_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0))&(l3_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0)))))&((~v2_struct_0(k7_filter_1(esk1_0,esk2_0))|l3_lattices(esk1_0))&((v10_lattices(k7_filter_1(esk1_0,esk2_0))|l3_lattices(esk1_0))&((v17_lattices(k7_filter_1(esk1_0,esk2_0))|l3_lattices(esk1_0))&(l3_lattices(k7_filter_1(esk1_0,esk2_0))|l3_lattices(esk1_0))))))))&(((~v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v2_struct_0(esk2_0))&((v10_lattices(k7_filter_1(esk1_0,esk2_0))|~v2_struct_0(esk2_0))&((v17_lattices(k7_filter_1(esk1_0,esk2_0))|~v2_struct_0(esk2_0))&(l3_lattices(k7_filter_1(esk1_0,esk2_0))|~v2_struct_0(esk2_0)))))&(((~v2_struct_0(k7_filter_1(esk1_0,esk2_0))|v10_lattices(esk2_0))&((v10_lattices(k7_filter_1(esk1_0,esk2_0))|v10_lattices(esk2_0))&((v17_lattices(k7_filter_1(esk1_0,esk2_0))|v10_lattices(esk2_0))&(l3_lattices(k7_filter_1(esk1_0,esk2_0))|v10_lattices(esk2_0)))))&(((~v2_struct_0(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0))&((v10_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0))&((v17_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0))&(l3_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0)))))&((~v2_struct_0(k7_filter_1(esk1_0,esk2_0))|l3_lattices(esk2_0))&((v10_lattices(k7_filter_1(esk1_0,esk2_0))|l3_lattices(esk2_0))&((v17_lattices(k7_filter_1(esk1_0,esk2_0))|l3_lattices(esk2_0))&(l3_lattices(k7_filter_1(esk1_0,esk2_0))|l3_lattices(esk2_0)))))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])])).
cnf(c_0_18, plain, (v10_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v17_lattices(X2)|~v11_lattices(X1)|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X2)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_19, negated_conjecture, (v10_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_20, negated_conjecture, (l3_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, negated_conjecture, (v10_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (v10_lattices(k7_filter_1(esk1_0,esk2_0))|v10_lattices(k7_filter_1(X1,esk2_0))|v2_struct_0(X1)|~v11_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_18, c_0_19]), c_0_20]), c_0_21])]), c_0_22])).
cnf(c_0_24, negated_conjecture, (v2_struct_0(esk1_0)|v2_struct_0(esk2_0)|v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v10_lattices(esk1_0)|~v17_lattices(esk1_0)|~l3_lattices(esk1_0)|~v10_lattices(esk2_0)|~v17_lattices(esk2_0)|~l3_lattices(esk2_0)|~v10_lattices(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(k7_filter_1(esk1_0,esk2_0))|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (v10_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_26, negated_conjecture, (l3_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_28, negated_conjecture, (v10_lattices(k7_filter_1(esk1_0,esk2_0))|v10_lattices(k7_filter_1(X1,esk2_0))|v2_struct_0(X1)|~v17_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X1)), inference(spm,[status(thm)],[c_0_23, c_0_16])).
cnf(c_0_29, negated_conjecture, (v10_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_30, plain, ![X12]:((~v2_struct_0(X12)|(v2_struct_0(X12)|(~v11_lattices(X12)|(~v15_lattices(X12)|~v16_lattices(X12))))|~l3_lattices(X12))&(v17_lattices(X12)|(v2_struct_0(X12)|(~v11_lattices(X12)|(~v15_lattices(X12)|~v16_lattices(X12))))|~l3_lattices(X12))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc6_lattices])])])])).
cnf(c_0_31, plain, (v11_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v10_lattices(X1)|~v11_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~v11_lattices(X2)|~l3_lattices(X2)|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_32, plain, ![X8, X9]:((v3_lattices(k7_filter_1(X8,X9))|(v2_struct_0(X8)|~l3_lattices(X8)|(v2_struct_0(X9)|~l3_lattices(X9))))&(l3_lattices(k7_filter_1(X8,X9))|(v2_struct_0(X8)|~l3_lattices(X8)|(v2_struct_0(X9)|~l3_lattices(X9))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_filter_1])])])])).
cnf(c_0_33, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v2_struct_0(k7_filter_1(X1,X2))|~v10_lattices(X1)|~v11_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~v11_lattices(X2)|~l3_lattices(X2)|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_34, plain, ![X2, X1]:(epred2_2(X1,X2)<=>(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v15_lattices(X1)&(v16_lattices(X1)&l3_lattices(X1)))))&(~(v2_struct_0(X2))&(v10_lattices(X2)&(v15_lattices(X2)&(v16_lattices(X2)&l3_lattices(X2))))))<=>(~(v2_struct_0(k7_filter_1(X1,X2)))&(v10_lattices(k7_filter_1(X1,X2))&(v15_lattices(k7_filter_1(X1,X2))&(v16_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2)))))))), introduced(definition)).
cnf(c_0_35, negated_conjecture, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(esk1_0)|~v17_lattices(esk2_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))|~v10_lattices(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_21]), c_0_26]), c_0_20])]), c_0_27]), c_0_22])).
cnf(c_0_36, negated_conjecture, (v10_lattices(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_26]), c_0_25])]), c_0_27])).
cnf(c_0_37, plain, (v17_lattices(X1)|v2_struct_0(X1)|~v11_lattices(X1)|~v15_lattices(X1)|~v16_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (v11_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v11_lattices(X2)|~v11_lattices(X1)|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[c_0_31, c_0_12])).
cnf(c_0_39, plain, (l3_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~l3_lattices(X1)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v11_lattices(X2)|~v11_lattices(X1)|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X2)|~v10_lattices(X1)|~v2_struct_0(k7_filter_1(X1,X2))), inference(csr,[status(thm)],[c_0_33, c_0_12])).
cnf(c_0_41, plain, (v11_lattices(X1)|v2_struct_0(k7_filter_1(X2,X1))|~v10_lattices(k7_filter_1(X2,X1))|~v11_lattices(k7_filter_1(X2,X1))|~l3_lattices(k7_filter_1(X2,X1))|~epred1_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_42, plain, ![X2, X1]:(epred2_2(X1,X2)=>(((~(v2_struct_0(X1))&(v10_lattices(X1)&(v15_lattices(X1)&(v16_lattices(X1)&l3_lattices(X1)))))&(~(v2_struct_0(X2))&(v10_lattices(X2)&(v15_lattices(X2)&(v16_lattices(X2)&l3_lattices(X2))))))<=>(~(v2_struct_0(k7_filter_1(X1,X2)))&(v10_lattices(k7_filter_1(X1,X2))&(v15_lattices(k7_filter_1(X1,X2))&(v16_lattices(k7_filter_1(X1,X2))&l3_lattices(k7_filter_1(X1,X2)))))))), inference(split_equiv,[status(thm)],[c_0_34])).
fof(c_0_43, axiom, ![X1]:((~(v2_struct_0(X1))&(v10_lattices(X1)&l3_lattices(X1)))=>![X2]:((~(v2_struct_0(X2))&(v10_lattices(X2)&l3_lattices(X2)))=>epred2_2(X1,X2))), inference(apply_def,[status(thm)],[t46_filter_1, c_0_34])).
cnf(c_0_44, negated_conjecture, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(esk1_0)|~v17_lattices(esk2_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36])])).
cnf(c_0_45, plain, (v17_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v16_lattices(k7_filter_1(X1,X2))|~v15_lattices(k7_filter_1(X1,X2))|~v11_lattices(X2)|~v11_lattices(X1)|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39]), c_0_40])).
cnf(c_0_46, plain, (v11_lattices(X1)|v2_struct_0(k7_filter_1(X2,X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v11_lattices(k7_filter_1(X2,X1))|~l3_lattices(X1)|~l3_lattices(X2)|~v10_lattices(k7_filter_1(X2,X1))|~v10_lattices(X1)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_12]), c_0_39])).
fof(c_0_47, plain, ![X15, X16]:(((~v2_struct_0(k7_filter_1(X16,X15))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v15_lattices(X16)|(~v16_lattices(X16)|~l3_lattices(X16))))|(v2_struct_0(X15)|(~v10_lattices(X15)|(~v15_lattices(X15)|(~v16_lattices(X15)|~l3_lattices(X15))))))|~epred2_2(X16,X15))&((v10_lattices(k7_filter_1(X16,X15))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v15_lattices(X16)|(~v16_lattices(X16)|~l3_lattices(X16))))|(v2_struct_0(X15)|(~v10_lattices(X15)|(~v15_lattices(X15)|(~v16_lattices(X15)|~l3_lattices(X15))))))|~epred2_2(X16,X15))&((v15_lattices(k7_filter_1(X16,X15))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v15_lattices(X16)|(~v16_lattices(X16)|~l3_lattices(X16))))|(v2_struct_0(X15)|(~v10_lattices(X15)|(~v15_lattices(X15)|(~v16_lattices(X15)|~l3_lattices(X15))))))|~epred2_2(X16,X15))&((v16_lattices(k7_filter_1(X16,X15))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v15_lattices(X16)|(~v16_lattices(X16)|~l3_lattices(X16))))|(v2_struct_0(X15)|(~v10_lattices(X15)|(~v15_lattices(X15)|(~v16_lattices(X15)|~l3_lattices(X15))))))|~epred2_2(X16,X15))&(l3_lattices(k7_filter_1(X16,X15))|(v2_struct_0(X16)|(~v10_lattices(X16)|(~v15_lattices(X16)|(~v16_lattices(X16)|~l3_lattices(X16))))|(v2_struct_0(X15)|(~v10_lattices(X15)|(~v15_lattices(X15)|(~v16_lattices(X15)|~l3_lattices(X15))))))|~epred2_2(X16,X15))))))&(((~v2_struct_0(X16)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))&((v10_lattices(X16)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))&((v15_lattices(X16)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))&((v16_lattices(X16)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))&(l3_lattices(X16)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))))))&((~v2_struct_0(X15)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))&((v10_lattices(X15)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))&((v15_lattices(X15)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))&((v16_lattices(X15)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15))&(l3_lattices(X15)|(v2_struct_0(k7_filter_1(X16,X15))|(~v10_lattices(k7_filter_1(X16,X15))|(~v15_lattices(k7_filter_1(X16,X15))|(~v16_lattices(k7_filter_1(X16,X15))|~l3_lattices(k7_filter_1(X16,X15))))))|~epred2_2(X16,X15)))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_42])])])])).
fof(c_0_48, plain, ![X10, X11]:(v2_struct_0(X10)|(~v10_lattices(X10)|~l3_lattices(X10))|(v2_struct_0(X11)|(~v10_lattices(X11)|~l3_lattices(X11))|epred2_2(X10,X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_43])])])])).
cnf(c_0_49, negated_conjecture, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v16_lattices(k7_filter_1(esk1_0,esk2_0))|~v15_lattices(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(esk1_0)|~v17_lattices(esk2_0)|~v11_lattices(esk2_0)|~v11_lattices(esk1_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45]), c_0_20]), c_0_26]), c_0_21]), c_0_25])]), c_0_27]), c_0_22])).
cnf(c_0_50, plain, (v11_lattices(X1)|v2_struct_0(k7_filter_1(X2,X1))|v2_struct_0(X1)|v2_struct_0(X2)|~v17_lattices(k7_filter_1(X2,X1))|~l3_lattices(X1)|~l3_lattices(X2)|~v10_lattices(k7_filter_1(X2,X1))|~v10_lattices(X1)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_16]), c_0_39])).
cnf(c_0_51, negated_conjecture, (v17_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_52, negated_conjecture, (v17_lattices(esk2_0)|~v2_struct_0(k7_filter_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_53, plain, (v15_lattices(X1)|v2_struct_0(k7_filter_1(X2,X1))|~v10_lattices(k7_filter_1(X2,X1))|~v15_lattices(k7_filter_1(X2,X1))|~v16_lattices(k7_filter_1(X2,X1))|~l3_lattices(k7_filter_1(X2,X1))|~epred2_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_54, plain, (v2_struct_0(X1)|v2_struct_0(X2)|epred2_2(X1,X2)|~v10_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~l3_lattices(X2)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_55, plain, (v16_lattices(X1)|v2_struct_0(X1)|~v17_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_56, negated_conjecture, (l3_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_57, plain, (v15_lattices(X1)|v2_struct_0(X1)|~v17_lattices(X1)|~l3_lattices(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_58, negated_conjecture, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v16_lattices(k7_filter_1(esk1_0,esk2_0))|~v15_lattices(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(esk1_0)|~v17_lattices(esk2_0)|~v11_lattices(esk2_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_16]), c_0_26])]), c_0_27])).
cnf(c_0_59, plain, (v16_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v10_lattices(X1)|~v15_lattices(X1)|~v16_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~v15_lattices(X2)|~v16_lattices(X2)|~l3_lattices(X2)|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_60, negated_conjecture, (v17_lattices(esk2_0)|v11_lattices(esk2_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51]), c_0_20]), c_0_26]), c_0_36]), c_0_21]), c_0_25])]), c_0_22]), c_0_27]), c_0_52])).
cnf(c_0_61, plain, (v15_lattices(X1)|v2_struct_0(k7_filter_1(X2,X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v16_lattices(k7_filter_1(X2,X1))|~v15_lattices(k7_filter_1(X2,X1))|~l3_lattices(X1)|~l3_lattices(X2)|~v10_lattices(k7_filter_1(X2,X1))|~v10_lattices(X1)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_39])).
cnf(c_0_62, negated_conjecture, (v16_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_51]), c_0_56]), c_0_52])).
cnf(c_0_63, negated_conjecture, (v15_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_51]), c_0_56]), c_0_52])).
cnf(c_0_64, negated_conjecture, (v15_lattices(esk2_0)|~v2_struct_0(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_52]), c_0_20])]), c_0_22])).
cnf(c_0_65, negated_conjecture, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v16_lattices(k7_filter_1(esk1_0,esk2_0))|~v15_lattices(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(esk1_0)|~v17_lattices(esk2_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_16]), c_0_20])]), c_0_22])).
cnf(c_0_66, plain, (v16_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v16_lattices(X2)|~v16_lattices(X1)|~v15_lattices(X2)|~v15_lattices(X1)|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[c_0_59, c_0_54])).
cnf(c_0_67, negated_conjecture, (v17_lattices(esk2_0)|~v16_lattices(esk2_0)|~v15_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_60]), c_0_20])]), c_0_22])).
cnf(c_0_68, plain, (v15_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v10_lattices(X1)|~v15_lattices(X1)|~v16_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~v15_lattices(X2)|~v16_lattices(X2)|~l3_lattices(X2)|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_69, plain, (v15_lattices(X1)|v2_struct_0(k7_filter_1(X1,X2))|~v10_lattices(k7_filter_1(X1,X2))|~v15_lattices(k7_filter_1(X1,X2))|~v16_lattices(k7_filter_1(X1,X2))|~l3_lattices(k7_filter_1(X1,X2))|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_70, negated_conjecture, (v17_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_71, negated_conjecture, (l3_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_72, negated_conjecture, (v17_lattices(esk1_0)|~v2_struct_0(k7_filter_1(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_73, negated_conjecture, (v15_lattices(esk2_0)|v17_lattices(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_20]), c_0_26]), c_0_36]), c_0_21]), c_0_25])]), c_0_27]), c_0_22]), c_0_63]), c_0_64])).
cnf(c_0_74, plain, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v16_lattices(esk2_0)|~v16_lattices(esk1_0)|~v15_lattices(k7_filter_1(esk1_0,esk2_0))|~v15_lattices(esk2_0)|~v15_lattices(esk1_0)|~v17_lattices(esk1_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_20]), c_0_26]), c_0_21]), c_0_25])]), c_0_22]), c_0_27]), c_0_67])).
cnf(c_0_75, plain, (v15_lattices(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v16_lattices(X2)|~v16_lattices(X1)|~v15_lattices(X2)|~v15_lattices(X1)|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[c_0_68, c_0_54])).
cnf(c_0_76, plain, (v15_lattices(X1)|v2_struct_0(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v16_lattices(k7_filter_1(X1,X2))|~v15_lattices(k7_filter_1(X1,X2))|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(k7_filter_1(X1,X2))|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_54]), c_0_39])).
cnf(c_0_77, negated_conjecture, (v16_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_70]), c_0_71]), c_0_72])).
cnf(c_0_78, negated_conjecture, (v15_lattices(k7_filter_1(esk1_0,esk2_0))|v17_lattices(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_70]), c_0_71]), c_0_72])).
cnf(c_0_79, negated_conjecture, (v15_lattices(esk1_0)|~v2_struct_0(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_72]), c_0_26])]), c_0_27])).
cnf(c_0_80, plain, (v16_lattices(X1)|v2_struct_0(k7_filter_1(X2,X1))|~v10_lattices(k7_filter_1(X2,X1))|~v15_lattices(k7_filter_1(X2,X1))|~v16_lattices(k7_filter_1(X2,X1))|~l3_lattices(k7_filter_1(X2,X1))|~epred2_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_81, negated_conjecture, (v15_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_73]), c_0_20])]), c_0_22])).
cnf(c_0_82, plain, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v16_lattices(esk2_0)|~v16_lattices(esk1_0)|~v15_lattices(esk2_0)|~v15_lattices(esk1_0)|~v17_lattices(esk1_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_20]), c_0_26]), c_0_21]), c_0_25])]), c_0_22]), c_0_27])).
cnf(c_0_83, negated_conjecture, (v15_lattices(esk1_0)|v17_lattices(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_20]), c_0_26]), c_0_36]), c_0_21]), c_0_25])]), c_0_27]), c_0_22]), c_0_78]), c_0_79])).
cnf(c_0_84, plain, (v16_lattices(X1)|v2_struct_0(k7_filter_1(X2,X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v16_lattices(k7_filter_1(X2,X1))|~v15_lattices(k7_filter_1(X2,X1))|~l3_lattices(X1)|~l3_lattices(X2)|~v10_lattices(k7_filter_1(X2,X1))|~v10_lattices(X1)|~v10_lattices(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_54]), c_0_39])).
cnf(c_0_85, negated_conjecture, (v16_lattices(esk2_0)|~v2_struct_0(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_52]), c_0_20])]), c_0_22])).
cnf(c_0_86, negated_conjecture, (v17_lattices(esk2_0)|~v16_lattices(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_81])])).
cnf(c_0_87, plain, (v16_lattices(X1)|v2_struct_0(k7_filter_1(X1,X2))|~v10_lattices(k7_filter_1(X1,X2))|~v15_lattices(k7_filter_1(X1,X2))|~v16_lattices(k7_filter_1(X1,X2))|~l3_lattices(k7_filter_1(X1,X2))|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_88, plain, (v11_lattices(X1)|v2_struct_0(k7_filter_1(X1,X2))|~v10_lattices(k7_filter_1(X1,X2))|~v11_lattices(k7_filter_1(X1,X2))|~l3_lattices(k7_filter_1(X1,X2))|~epred1_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_89, plain, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v16_lattices(esk2_0)|~v16_lattices(esk1_0)|~v15_lattices(esk1_0)|~v17_lattices(esk1_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_81])])).
cnf(c_0_90, negated_conjecture, (v15_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_83]), c_0_26])]), c_0_27])).
cnf(c_0_91, negated_conjecture, (v17_lattices(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_62]), c_0_20]), c_0_26]), c_0_36]), c_0_21]), c_0_25])]), c_0_27]), c_0_22]), c_0_63]), c_0_85]), c_0_86])).
cnf(c_0_92, plain, (v16_lattices(X1)|v2_struct_0(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v16_lattices(k7_filter_1(X1,X2))|~v15_lattices(k7_filter_1(X1,X2))|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(k7_filter_1(X1,X2))|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_54]), c_0_39])).
cnf(c_0_93, negated_conjecture, (v16_lattices(esk1_0)|~v2_struct_0(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_72]), c_0_26])]), c_0_27])).
cnf(c_0_94, plain, (v11_lattices(X1)|v2_struct_0(k7_filter_1(X1,X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v11_lattices(k7_filter_1(X1,X2))|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(k7_filter_1(X1,X2))|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_12]), c_0_39])).
cnf(c_0_95, plain, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v16_lattices(esk2_0)|~v16_lattices(esk1_0)|~v17_lattices(esk1_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_90])])).
cnf(c_0_96, negated_conjecture, (v16_lattices(esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_91]), c_0_20])]), c_0_22])).
cnf(c_0_97, negated_conjecture, (v16_lattices(esk1_0)|v17_lattices(esk1_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_77]), c_0_20]), c_0_26]), c_0_36]), c_0_21]), c_0_25])]), c_0_27]), c_0_22]), c_0_78]), c_0_93])).
cnf(c_0_98, plain, (v11_lattices(X1)|v2_struct_0(k7_filter_1(X1,X2))|v2_struct_0(X2)|v2_struct_0(X1)|~v17_lattices(k7_filter_1(X1,X2))|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(k7_filter_1(X1,X2))|~v10_lattices(X2)|~v10_lattices(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_16]), c_0_39])).
cnf(c_0_99, plain, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v16_lattices(esk1_0)|~v17_lattices(esk1_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_96])])).
cnf(c_0_100, negated_conjecture, (v16_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_97]), c_0_26])]), c_0_27])).
cnf(c_0_101, negated_conjecture, (v17_lattices(esk1_0)|v11_lattices(esk1_0)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_70]), c_0_20]), c_0_26]), c_0_36]), c_0_21]), c_0_25])]), c_0_22]), c_0_27]), c_0_72])).
cnf(c_0_102, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v2_struct_0(k7_filter_1(X1,X2))|~v10_lattices(X1)|~v15_lattices(X1)|~v16_lattices(X1)|~l3_lattices(X1)|~v10_lattices(X2)|~v15_lattices(X2)|~v16_lattices(X2)|~l3_lattices(X2)|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_103, plain, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~v17_lattices(esk1_0)|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_104, negated_conjecture, (v17_lattices(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_101]), c_0_100]), c_0_90]), c_0_26])]), c_0_27])).
cnf(c_0_105, plain, (v2_struct_0(X1)|v2_struct_0(X2)|~v16_lattices(X2)|~v16_lattices(X1)|~v15_lattices(X2)|~v15_lattices(X1)|~l3_lattices(X2)|~l3_lattices(X1)|~v10_lattices(X2)|~v10_lattices(X1)|~v2_struct_0(k7_filter_1(X1,X2))), inference(csr,[status(thm)],[c_0_102, c_0_54])).
cnf(c_0_106, plain, (v2_struct_0(k7_filter_1(esk1_0,esk2_0))|~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_107, plain, (~l3_lattices(k7_filter_1(esk1_0,esk2_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_96]), c_0_100]), c_0_81]), c_0_90]), c_0_20]), c_0_26]), c_0_21]), c_0_25])]), c_0_22]), c_0_27])).
cnf(c_0_108, plain, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_39]), c_0_20]), c_0_26])]), c_0_22]), c_0_27]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 109
# Proof object clause steps            : 88
# Proof object formula steps           : 21
# Proof object conjectures             : 47
# Proof object clause conjectures      : 44
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 40
# Proof object simplifying inferences  : 188
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 76
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 74
# Processed clauses                    : 395
# ...of these trivial                  : 19
# ...subsumed                          : 123
# ...remaining for further processing  : 253
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 44
# Backward-rewritten                   : 79
# Generated clauses                    : 518
# ...of the previous two non-trivial   : 367
# Contextual simplify-reflections      : 104
# Paramodulations                      : 518
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 82
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 65
# Current number of unprocessed clauses: 50
# ...number of literals in the above   : 624
# Current number of archived formulas  : 0
# Current number of archived clauses   : 171
# Clause-clause subsumption calls (NU) : 8507
# Rec. Clause-clause subsumption calls : 1501
# Non-unit clause-clause subsumptions  : 240
# Unit Clause-clause subsumption calls : 294
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 10
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20601

# -------------------------------------------------
# User time                : 0.020 s
# System time              : 0.003 s
# Total time               : 0.023 s
# Maximum resident set size: 28652 pages
