# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t70_enumset1)).
fof(t95_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2)), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t95_xboole_1)).
fof(t75_zfmisc_1, conjecture, ![X1, X2, X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t75_zfmisc_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', idempotence_k3_xboole_0)).
fof(t31_zfmisc_1, axiom, ![X1]:k3_tarski(k1_tarski(X1))=X1, file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t31_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t69_enumset1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t100_xboole_1)).
fof(t92_xboole_1, axiom, ![X1]:k5_xboole_0(X1,X1)=k1_xboole_0, file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t92_xboole_1)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t91_xboole_1)).
fof(t94_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t94_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t28_xboole_1)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', commutativity_k5_xboole_0)).
fof(l45_zfmisc_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_tarski(X2,X3))<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3))))))), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', l45_zfmisc_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t7_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t5_boole)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t2_boole)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', t10_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/f2966ea1eebf4009aa108f1bbc124eb9.E_input', d10_xboole_0)).
fof(c_0_19, plain, ![X22, X23]:k3_tarski(k2_tarski(X22,X23))=k2_xboole_0(X22,X23), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_20, plain, ![X28, X29]:k1_enumset1(X28,X28,X29)=k2_tarski(X28,X29), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_21, plain, ![X15, X16]:k3_xboole_0(X15,X16)=k5_xboole_0(k5_xboole_0(X15,X16),k2_xboole_0(X15,X16)), inference(variable_rename,[status(thm)],[t95_xboole_1])).
cnf(c_0_22, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_23, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_24, negated_conjecture, ~(![X1, X2, X3]:(k4_xboole_0(X1,k2_tarski(X2,X3))=k1_xboole_0<=>~((~(X1=k1_xboole_0)&(~(X1=k1_tarski(X2))&(~(X1=k1_tarski(X3))&~(X1=k2_tarski(X2,X3)))))))), inference(assume_negation,[status(cth)],[t75_zfmisc_1])).
fof(c_0_25, plain, ![X37]:k3_xboole_0(X37,X37)=X37, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_26, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_28, plain, ![X30]:k3_tarski(k1_tarski(X30))=X30, inference(variable_rename,[status(thm)],[t31_zfmisc_1])).
fof(c_0_29, plain, ![X12]:k2_tarski(X12,X12)=k1_tarski(X12), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_30, negated_conjecture, (((esk1_0!=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk2_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&((esk1_0!=k1_tarski(esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0)&(esk1_0!=k2_tarski(esk2_0,esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0))))&(k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0|(esk1_0=k1_xboole_0|(esk1_0=k1_tarski(esk2_0)|(esk1_0=k1_tarski(esk3_0)|esk1_0=k2_tarski(esk2_0,esk3_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_24])])])])])).
fof(c_0_31, plain, ![X38, X39]:k4_xboole_0(X38,X39)=k5_xboole_0(X38,k3_xboole_0(X38,X39)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k3_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_26, c_0_27])).
fof(c_0_34, plain, ![X14]:k5_xboole_0(X14,X14)=k1_xboole_0, inference(variable_rename,[status(thm)],[t92_xboole_1])).
cnf(c_0_35, plain, (k3_tarski(k1_tarski(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_37, negated_conjecture, (k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))=k1_xboole_0|esk1_0=k1_xboole_0|esk1_0=k1_tarski(esk2_0)|esk1_0=k1_tarski(esk3_0)|esk1_0=k2_tarski(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, ![X9, X10, X11]:k5_xboole_0(k5_xboole_0(X9,X10),X11)=k5_xboole_0(X9,k5_xboole_0(X10,X11)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_40, plain, (k5_xboole_0(k5_xboole_0(X1,X1),k3_tarski(k1_enumset1(X1,X1,X1)))=X1), inference(rw,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_41, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_42, plain, (k3_tarski(k1_enumset1(X1,X1,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36]), c_0_23])).
fof(c_0_43, plain, ![X24, X25]:k2_xboole_0(X24,X25)=k5_xboole_0(k5_xboole_0(X24,X25),k3_xboole_0(X24,X25)), inference(variable_rename,[status(thm)],[t94_xboole_1])).
fof(c_0_44, plain, ![X26, X27]:(~r1_tarski(X26,X27)|k3_xboole_0(X26,X27)=X26), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_45, negated_conjecture, (esk1_0=k1_xboole_0|esk1_0=k1_enumset1(esk2_0,esk2_0,esk2_0)|esk1_0=k1_enumset1(esk2_0,esk2_0,esk3_0)|esk1_0=k1_enumset1(esk3_0,esk3_0,esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)))))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_36]), c_0_36]), c_0_23]), c_0_23]), c_0_23]), c_0_23]), c_0_38]), c_0_33])).
cnf(c_0_46, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41]), c_0_42])).
cnf(c_0_48, negated_conjecture, (esk1_0!=k1_tarski(esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_49, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_51, plain, ![X17, X18]:k5_xboole_0(X17,X18)=k5_xboole_0(X18,X17), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
fof(c_0_52, plain, ![X19, X20, X21]:((~r1_tarski(X19,k2_tarski(X20,X21))|(X19=k1_xboole_0|(X19=k1_tarski(X20)|(X19=k1_tarski(X21)|X19=k2_tarski(X20,X21)))))&((X19!=k1_xboole_0|r1_tarski(X19,k2_tarski(X20,X21)))&((X19!=k1_tarski(X20)|r1_tarski(X19,k2_tarski(X20,X21)))&((X19!=k1_tarski(X21)|r1_tarski(X19,k2_tarski(X20,X21)))&(X19!=k2_tarski(X20,X21)|r1_tarski(X19,k2_tarski(X20,X21))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l45_zfmisc_1])])])])).
fof(c_0_53, plain, ![X34, X35]:r1_tarski(X34,k2_xboole_0(X34,X35)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
cnf(c_0_54, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))))=k1_xboole_0|k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_55, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_41]), c_0_47])).
fof(c_0_56, plain, ![X13]:k5_xboole_0(X13,k1_xboole_0)=X13, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_57, negated_conjecture, (esk1_0!=k1_enumset1(esk3_0,esk3_0,esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_36]), c_0_23]), c_0_23]), c_0_38]), c_0_33])).
cnf(c_0_58, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k5_xboole_0(k5_xboole_0(X1,X2),k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k1_enumset1(X1,X1,X2))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_27]), c_0_33])).
cnf(c_0_59, plain, (k5_xboole_0(k5_xboole_0(X1,X2),k3_tarski(k1_enumset1(X1,X1,X2)))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_50, c_0_33])).
cnf(c_0_60, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (r1_tarski(X1,k2_tarski(X3,X2))|X1!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_63, negated_conjecture, (k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))=k1_xboole_0|k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_64, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_65, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|X1=k1_tarski(X3)|X1=k2_tarski(X2,X3)|~r1_tarski(X1,k2_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_66, plain, ![X36]:k3_xboole_0(X36,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
cnf(c_0_67, negated_conjecture, (esk1_0!=k2_tarski(esk2_0,esk3_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_68, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))))!=k1_xboole_0|k1_enumset1(esk3_0,esk3_0,esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_57, c_0_46])).
cnf(c_0_69, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k1_enumset1(X1,X1,X2))))))=k3_tarski(k1_enumset1(X1,X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_46]), c_0_46])).
cnf(c_0_70, plain, (k5_xboole_0(X1,k5_xboole_0(X2,k3_tarski(k1_enumset1(X1,X1,X2))))=X1|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_59, c_0_46])).
cnf(c_0_71, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(spm,[status(thm)],[c_0_55, c_0_60])).
cnf(c_0_72, plain, (r1_tarski(X1,k1_enumset1(X3,X3,X2))|X1!=k1_enumset1(X2,X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_36]), c_0_23]), c_0_23])).
cnf(c_0_73, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X1,X1,X2)))), inference(rw,[status(thm)],[c_0_62, c_0_27])).
cnf(c_0_74, negated_conjecture, (k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)))=k1_enumset1(esk2_0,esk2_0,esk3_0)|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_63]), c_0_64])).
cnf(c_0_75, plain, (X1=k1_xboole_0|X1=k1_enumset1(X3,X3,X3)|X1=k1_enumset1(X2,X2,X3)|X1=k1_enumset1(X2,X2,X2)|~r1_tarski(X1,k1_enumset1(X2,X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_36]), c_0_36]), c_0_23]), c_0_23]), c_0_23]), c_0_23])).
cnf(c_0_76, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_77, negated_conjecture, (esk1_0!=k1_tarski(esk2_0)|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_78, negated_conjecture, (esk1_0!=k1_enumset1(esk2_0,esk2_0,esk3_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_23]), c_0_23]), c_0_38]), c_0_33])).
cnf(c_0_79, negated_conjecture, (k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))!=k1_xboole_0|k1_enumset1(esk3_0,esk3_0,esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_68, c_0_55])).
cnf(c_0_80, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=X2|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).
cnf(c_0_81, plain, (r1_tarski(k1_enumset1(X1,X1,X1),k1_enumset1(X2,X2,X1))), inference(er,[status(thm)],[c_0_72])).
cnf(c_0_82, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|k1_enumset1(esk3_0,esk3_0,esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74]), c_0_75])).
fof(c_0_83, plain, ![X31, X32, X33]:(~r1_tarski(X31,X32)|r1_tarski(X31,k2_xboole_0(X33,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_84, plain, (k5_xboole_0(k5_xboole_0(X1,k1_xboole_0),k3_tarski(k1_enumset1(X1,X1,k1_xboole_0)))=k1_xboole_0), inference(rw,[status(thm)],[c_0_76, c_0_33])).
cnf(c_0_85, negated_conjecture, (esk1_0!=k1_enumset1(esk2_0,esk2_0,esk2_0)|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_77, c_0_36]), c_0_23]), c_0_23]), c_0_38]), c_0_33])).
cnf(c_0_86, plain, (r1_tarski(X1,k2_tarski(X2,X3))|X1!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_87, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))))!=k1_xboole_0|k1_enumset1(esk2_0,esk2_0,esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_78, c_0_46])).
cnf(c_0_88, negated_conjecture, (k1_enumset1(esk3_0,esk3_0,esk3_0)!=esk1_0|~r1_tarski(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80]), c_0_41])])).
cnf(c_0_89, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_xboole_0=esk1_0|r1_tarski(esk1_0,k1_enumset1(X1,X1,esk3_0))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_90, negated_conjecture, (esk1_0!=k1_xboole_0|k4_xboole_0(esk1_0,k2_tarski(esk2_0,esk3_0))!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_91, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_92, plain, (k5_xboole_0(X1,k3_tarski(k1_enumset1(X1,X1,k1_xboole_0)))=k1_xboole_0), inference(rw,[status(thm)],[c_0_84, c_0_64])).
fof(c_0_93, plain, ![X7, X8]:(((r1_tarski(X7,X8)|X7!=X8)&(r1_tarski(X8,X7)|X7!=X8))&(~r1_tarski(X7,X8)|~r1_tarski(X8,X7)|X7=X8)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_94, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))))!=k1_xboole_0|k1_enumset1(esk2_0,esk2_0,esk2_0)!=esk1_0), inference(rw,[status(thm)],[c_0_85, c_0_46])).
cnf(c_0_95, plain, (r1_tarski(X1,k1_enumset1(X2,X2,X3))|X1!=k1_enumset1(X2,X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_36]), c_0_23]), c_0_23])).
cnf(c_0_96, negated_conjecture, (k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))!=k1_xboole_0|k1_enumset1(esk2_0,esk2_0,esk3_0)!=esk1_0), inference(rw,[status(thm)],[c_0_87, c_0_55])).
cnf(c_0_97, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_enumset1(esk2_0,esk2_0,esk3_0)=esk1_0|k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_82])).
cnf(c_0_98, negated_conjecture, (esk1_0!=k1_xboole_0|k5_xboole_0(esk1_0,k5_xboole_0(k5_xboole_0(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0)))))!=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_23]), c_0_38]), c_0_33])).
cnf(c_0_99, plain, (r1_tarski(X1,k3_tarski(k1_enumset1(X3,X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_91, c_0_27])).
cnf(c_0_100, plain, (k3_tarski(k1_enumset1(X1,X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_92]), c_0_64])).
cnf(c_0_101, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_102, negated_conjecture, (k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))!=k1_xboole_0|k1_enumset1(esk2_0,esk2_0,esk2_0)!=esk1_0), inference(rw,[status(thm)],[c_0_94, c_0_55])).
cnf(c_0_103, plain, (r1_tarski(k1_enumset1(X1,X1,X1),k1_enumset1(X1,X1,X2))), inference(er,[status(thm)],[c_0_95])).
cnf(c_0_104, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)=esk1_0|k1_xboole_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_42]), c_0_41])])).
cnf(c_0_105, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))))!=k1_xboole_0|k1_xboole_0!=esk1_0), inference(rw,[status(thm)],[c_0_98, c_0_46])).
cnf(c_0_106, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_107, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_101])).
cnf(c_0_108, negated_conjecture, (k1_enumset1(esk2_0,esk2_0,esk2_0)!=esk1_0|~r1_tarski(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_80]), c_0_41])])).
cnf(c_0_109, negated_conjecture, (k1_xboole_0=esk1_0|r1_tarski(esk1_0,k1_enumset1(esk2_0,esk2_0,X1))), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_110, negated_conjecture, (k5_xboole_0(k1_enumset1(esk2_0,esk2_0,esk3_0),k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))))!=k1_xboole_0|k1_xboole_0!=esk1_0), inference(rw,[status(thm)],[c_0_105, c_0_55])).
cnf(c_0_111, plain, (r1_tarski(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_112, negated_conjecture, (k1_xboole_0=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_104])).
cnf(c_0_113, negated_conjecture, (k1_xboole_0!=esk1_0|~r1_tarski(esk1_0,k1_enumset1(esk2_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_80]), c_0_41])])).
cnf(c_0_114, plain, (r1_tarski(esk1_0,X1)), inference(rw,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_115, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_113, c_0_112])]), c_0_114])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 116
# Proof object clause steps            : 77
# Proof object formula steps           : 39
# Proof object conjectures             : 34
# Proof object clause conjectures      : 31
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 16
# Proof object simplifying inferences  : 88
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 19
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 24
# Processed clauses                    : 180
# ...of these trivial                  : 8
# ...subsumed                          : 62
# ...remaining for further processing  : 109
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 40
# Generated clauses                    : 453
# ...of the previous two non-trivial   : 307
# Contextual simplify-reflections      : 3
# Paramodulations                      : 447
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 8
# Current number of unprocessed clauses: 156
# ...number of literals in the above   : 260
# Current number of archived formulas  : 0
# Current number of archived clauses   : 79
# Clause-clause subsumption calls (NU) : 174
# Rec. Clause-clause subsumption calls : 85
# Non-unit clause-clause subsumptions  : 46
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 88
# BW rewrite match successes           : 61
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5352

# -------------------------------------------------
# User time                : 0.013 s
# System time              : 0.000 s
# Total time               : 0.013 s
# Maximum resident set size: 28628 pages
