# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t88_funct_2, conjecture, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&(v3_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>(r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,X2,k2_funct_2(X1,X2)),k6_partfun1(X1))&r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,k2_funct_2(X1,X2),X2),k6_partfun1(X1)))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t88_funct_2)).
fof(cc2_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v3_funct_2(X3,X1,X2))=>(v1_funct_1(X3)&(v2_funct_1(X3)&v2_funct_2(X3,X2))))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', cc2_funct_2)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', cc1_relset_1)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t5_subset)).
fof(rc1_funct_2, axiom, ![X1, X2]:?[X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&(v1_relat_1(X3)&(v4_relat_1(X3,X1)&(v5_relat_1(X3,X2)&(v1_funct_1(X3)&v1_funct_2(X3,X1,X2)))))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', rc1_funct_2)).
fof(d3_funct_2, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>(v2_funct_2(X2,X1)<=>k2_relat_1(X2)=X1)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', d3_funct_2)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t7_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', d1_xboole_0)).
fof(dt_k2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k2_relset_1(X1,X2,X3),k1_zfmisc_1(X2))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', dt_k2_relset_1)).
fof(redefinition_k7_relset_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k7_relset_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', redefinition_k7_relset_1)).
fof(t38_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(k7_relset_1(X1,X2,X3,X1)=k2_relset_1(X1,X2,X3)&k8_relset_1(X1,X2,X3,X2)=k1_relset_1(X1,X2,X3))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t38_relset_1)).
fof(t64_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>((k1_relat_1(X1)=k1_xboole_0|k2_relat_1(X1)=k1_xboole_0)=>X1=k1_xboole_0)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t64_relat_1)).
fof(t113_zfmisc_1, axiom, ![X1, X2]:(k2_zfmisc_1(X1,X2)=k1_xboole_0<=>(X1=k1_xboole_0|X2=k1_xboole_0)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t113_zfmisc_1)).
fof(t8_boole, axiom, ![X1, X2]:~((v1_xboole_0(X1)&(~(X1=X2)&v1_xboole_0(X2)))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t8_boole)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', cc4_relset_1)).
fof(redefinition_k2_funct_2, axiom, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&(v3_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>k2_funct_2(X1,X2)=k2_funct_1(X2)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', redefinition_k2_funct_2)).
fof(t69_xboole_1, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>~((r1_tarski(X2,X1)&r1_xboole_0(X2,X1)))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t69_xboole_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t65_xboole_1)).
fof(t4_subset_1, axiom, ![X1]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t4_subset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', fc1_xboole_0)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', fc6_relat_1)).
fof(t3_funct_2, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v1_funct_1(X1)&(v1_funct_2(X1,k1_relat_1(X1),k2_relat_1(X1))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t3_funct_2)).
fof(t55_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>(v2_funct_1(X1)=>(k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))&k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t55_funct_1)).
fof(dt_k2_funct_2, axiom, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&(v3_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>(v1_funct_1(k2_funct_2(X1,X2))&(v1_funct_2(k2_funct_2(X1,X2),X1,X1)&(v3_funct_2(k2_funct_2(X1,X2),X1,X1)&m1_subset_1(k2_funct_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', dt_k2_funct_2)).
fof(d18_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>(v4_relat_1(X2,X1)<=>r1_tarski(k1_relat_1(X2),X1))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', d18_relat_1)).
fof(t146_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t146_relat_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', cc2_relat_1)).
fof(t54_relset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))=>(![X4]:(r2_hidden(X4,X1)=>k11_relat_1(X2,X4)=k11_relat_1(X3,X4))=>r2_relset_1(X1,X1,X2,X3)))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t54_relset_1)).
fof(t29_relset_1, axiom, ![X1]:m1_subset_1(k6_relat_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t29_relset_1)).
fof(redefinition_k6_partfun1, axiom, ![X1]:k6_partfun1(X1)=k6_relat_1(X1), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', redefinition_k6_partfun1)).
fof(t35_funct_2, axiom, ![X1, X2, X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>((k2_relset_1(X1,X2,X3)=X2&v2_funct_1(X3))=>(X2=k1_xboole_0|(k5_relat_1(X3,k2_funct_1(X3))=k6_partfun1(X1)&k5_relat_1(k2_funct_1(X3),X3)=k6_partfun1(X2))))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', t35_funct_2)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', dt_k1_partfun1)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:(((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&(v1_funct_1(X6)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k5_relat_1(X5,X6)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', redefinition_k1_partfun1)).
fof(reflexivity_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>r2_relset_1(X1,X2,X3,X3)), file('data/MPTP/proofs/f367eda2f5514c8c877c08106a549233.E_input', reflexivity_r2_relset_1)).
fof(c_0_35, negated_conjecture, ~(![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,X1,X1)&(v3_funct_2(X2,X1,X1)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))))=>(r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,X2,k2_funct_2(X1,X2)),k6_partfun1(X1))&r2_relset_1(X1,X1,k1_partfun1(X1,X1,X1,X1,k2_funct_2(X1,X2),X2),k6_partfun1(X1))))), inference(assume_negation,[status(cth)],[t88_funct_2])).
fof(c_0_36, plain, ![X55, X56, X57]:((v1_funct_1(X57)|(~v1_funct_1(X57)|~v3_funct_2(X57,X55,X56))|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56))))&((v2_funct_1(X57)|(~v1_funct_1(X57)|~v3_funct_2(X57,X55,X56))|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56))))&(v2_funct_2(X57,X56)|(~v1_funct_1(X57)|~v3_funct_2(X57,X55,X56))|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_funct_2])])])).
fof(c_0_37, negated_conjecture, ((v1_funct_1(esk2_0)&(v1_funct_2(esk2_0,esk1_0,esk1_0)&(v3_funct_2(esk2_0,esk1_0,esk1_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0))))))&(~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k2_funct_2(esk1_0,esk2_0)),k6_partfun1(esk1_0))|~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,k2_funct_2(esk1_0,esk2_0),esk2_0),k6_partfun1(esk1_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).
fof(c_0_38, plain, ![X65, X66, X67]:((v4_relat_1(X67,X65)|~m1_subset_1(X67,k1_zfmisc_1(k2_zfmisc_1(X65,X66))))&(v5_relat_1(X67,X66)|~m1_subset_1(X67,k1_zfmisc_1(k2_zfmisc_1(X65,X66))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_39, plain, ![X36, X37, X38]:(~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(X36,X37)))|v1_relat_1(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_40, plain, ![X48, X49, X50]:(~r2_hidden(X48,X49)|(~m1_subset_1(X49,k1_zfmisc_1(X50))|~v1_xboole_0(X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_41, plain, ![X25, X26]:(m1_subset_1(esk4_2(X25,X26),k1_zfmisc_1(k2_zfmisc_1(X25,X26)))&(v1_relat_1(esk4_2(X25,X26))&(v4_relat_1(esk4_2(X25,X26),X25)&(v5_relat_1(esk4_2(X25,X26),X26)&(v1_funct_1(esk4_2(X25,X26))&v1_funct_2(esk4_2(X25,X26),X25,X26)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_funct_2])])).
fof(c_0_42, plain, ![X75, X76]:((~v2_funct_2(X76,X75)|k2_relat_1(X76)=X75|(~v1_relat_1(X76)|~v5_relat_1(X76,X75)))&(k2_relat_1(X76)!=X75|v2_funct_2(X76,X75)|(~v1_relat_1(X76)|~v5_relat_1(X76,X75)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_funct_2])])])).
cnf(c_0_43, plain, (v2_funct_2(X1,X2)|~v1_funct_1(X1)|~v3_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, negated_conjecture, (v3_funct_2(esk2_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_46, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_47, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_48, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_51, plain, ![X23]:(X23=k1_xboole_0|r2_hidden(esk3_1(X23),X23)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
fof(c_0_52, plain, ![X87, X88, X89]:((~v1_xboole_0(X87)|~r2_hidden(X88,X87))&(r2_hidden(esk6_1(X89),X89)|v1_xboole_0(X89))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_53, plain, ![X71, X72, X73]:(~m1_subset_1(X73,k1_zfmisc_1(k2_zfmisc_1(X71,X72)))|m1_subset_1(k2_relset_1(X71,X72,X73),k1_zfmisc_1(X72))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_relset_1])])).
fof(c_0_54, plain, ![X28, X29, X30, X31]:(~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|k7_relset_1(X28,X29,X30,X31)=k9_relat_1(X30,X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_relset_1])])).
fof(c_0_55, plain, ![X68, X69, X70]:((k7_relset_1(X68,X69,X70,X68)=k2_relset_1(X68,X69,X70)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69))))&(k8_relset_1(X68,X69,X70,X69)=k1_relset_1(X68,X69,X70)|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X68,X69))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t38_relset_1])])])).
fof(c_0_56, plain, ![X86]:((k1_relat_1(X86)!=k1_xboole_0|X86=k1_xboole_0|~v1_relat_1(X86))&(k2_relat_1(X86)!=k1_xboole_0|X86=k1_xboole_0|~v1_relat_1(X86))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_relat_1])])])).
cnf(c_0_57, plain, (k2_relat_1(X1)=X2|~v2_funct_2(X1,X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_58, negated_conjecture, (v2_funct_2(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_45]), c_0_46])])).
cnf(c_0_59, negated_conjecture, (v5_relat_1(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_47, c_0_45])).
cnf(c_0_60, negated_conjecture, (v1_relat_1(esk2_0)), inference(spm,[status(thm)],[c_0_48, c_0_45])).
cnf(c_0_61, plain, (~r2_hidden(X1,esk4_2(X2,X3))|~v1_xboole_0(k2_zfmisc_1(X2,X3))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_62, plain, (X1=k1_xboole_0|r2_hidden(esk3_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_63, plain, ![X53, X54]:((k2_zfmisc_1(X53,X54)!=k1_xboole_0|(X53=k1_xboole_0|X54=k1_xboole_0))&((X53!=k1_xboole_0|k2_zfmisc_1(X53,X54)=k1_xboole_0)&(X54!=k1_xboole_0|k2_zfmisc_1(X53,X54)=k1_xboole_0))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t113_zfmisc_1])])])).
fof(c_0_64, plain, ![X13, X14]:(~v1_xboole_0(X13)|(X13=X14|~v1_xboole_0(X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_boole])])])).
cnf(c_0_65, plain, (r2_hidden(esk6_1(X1),X1)|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_66, plain, ![X45, X46, X47]:(~v1_xboole_0(X45)|(~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X46,X45)))|v1_xboole_0(X47))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_67, plain, (m1_subset_1(k2_relset_1(X2,X3,X1),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_68, plain, (k7_relset_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_69, plain, (k7_relset_1(X1,X2,X3,X1)=k2_relset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_70, plain, (X1=k1_xboole_0|k2_relat_1(X1)!=k1_xboole_0|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_71, negated_conjecture, (k2_relat_1(esk2_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59]), c_0_60])])).
cnf(c_0_72, plain, (v1_funct_1(esk4_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_73, plain, (esk4_2(X1,X2)=k1_xboole_0|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_74, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_75, plain, ![X17, X18]:(~v1_funct_1(X18)|(~v1_funct_2(X18,X17,X17)|(~v3_funct_2(X18,X17,X17)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(X17,X17)))))|k2_funct_2(X17,X18)=k2_funct_1(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_funct_2])])).
fof(c_0_76, plain, ![X94, X95]:(v1_xboole_0(X95)|(~r1_tarski(X95,X94)|~r1_xboole_0(X95,X94))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t69_xboole_1])])])).
fof(c_0_77, plain, ![X77]:r1_xboole_0(X77,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
cnf(c_0_78, plain, (X1=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_79, plain, (v1_xboole_0(esk4_2(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_61, c_0_65])).
cnf(c_0_80, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_81, plain, (~r2_hidden(X1,k2_relset_1(X2,X3,X4))|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(spm,[status(thm)],[c_0_49, c_0_67])).
cnf(c_0_82, plain, (k2_relset_1(X1,X2,X3)=k9_relat_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_83, plain, (k2_zfmisc_1(X2,X1)=k1_xboole_0|X1!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_84, negated_conjecture, (esk2_0=k1_xboole_0|k1_xboole_0!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_60])])).
fof(c_0_85, plain, ![X58]:m1_subset_1(k1_xboole_0,k1_zfmisc_1(X58)), inference(variable_rename,[status(thm)],[t4_subset_1])).
cnf(c_0_86, plain, (v1_funct_1(k1_xboole_0)|~v1_xboole_0(k2_zfmisc_1(X1,X2))), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_87, plain, (k2_zfmisc_1(k1_xboole_0,X1)=k1_xboole_0), inference(er,[status(thm)],[c_0_74])).
cnf(c_0_88, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_89, plain, (v5_relat_1(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_90, plain, ![X11, X12]:v1_relat_1(k2_zfmisc_1(X11,X12)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_91, plain, ![X93]:((v1_funct_1(X93)|(~v1_relat_1(X93)|~v1_funct_1(X93)))&((v1_funct_2(X93,k1_relat_1(X93),k2_relat_1(X93))|(~v1_relat_1(X93)|~v1_funct_1(X93)))&(m1_subset_1(X93,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X93),k2_relat_1(X93))))|(~v1_relat_1(X93)|~v1_funct_1(X93))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_funct_2])])])).
fof(c_0_92, plain, ![X74]:((k2_relat_1(X74)=k1_relat_1(k2_funct_1(X74))|~v2_funct_1(X74)|(~v1_relat_1(X74)|~v1_funct_1(X74)))&(k1_relat_1(X74)=k2_relat_1(k2_funct_1(X74))|~v2_funct_1(X74)|(~v1_relat_1(X74)|~v1_funct_1(X74)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t55_funct_1])])])).
cnf(c_0_93, plain, (k2_funct_2(X2,X1)=k2_funct_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X2)|~v3_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_94, negated_conjecture, (v1_funct_2(esk2_0,esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_95, plain, (v2_funct_1(X1)|~v1_funct_1(X1)|~v3_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_96, plain, ![X91, X92]:((v1_funct_1(k2_funct_2(X91,X92))|(~v1_funct_1(X92)|(~v1_funct_2(X92,X91,X91)|(~v3_funct_2(X92,X91,X91)|~m1_subset_1(X92,k1_zfmisc_1(k2_zfmisc_1(X91,X91)))))))&((v1_funct_2(k2_funct_2(X91,X92),X91,X91)|(~v1_funct_1(X92)|(~v1_funct_2(X92,X91,X91)|(~v3_funct_2(X92,X91,X91)|~m1_subset_1(X92,k1_zfmisc_1(k2_zfmisc_1(X91,X91)))))))&((v3_funct_2(k2_funct_2(X91,X92),X91,X91)|(~v1_funct_1(X92)|(~v1_funct_2(X92,X91,X91)|(~v3_funct_2(X92,X91,X91)|~m1_subset_1(X92,k1_zfmisc_1(k2_zfmisc_1(X91,X91)))))))&(m1_subset_1(k2_funct_2(X91,X92),k1_zfmisc_1(k2_zfmisc_1(X91,X91)))|(~v1_funct_1(X92)|(~v1_funct_2(X92,X91,X91)|(~v3_funct_2(X92,X91,X91)|~m1_subset_1(X92,k1_zfmisc_1(k2_zfmisc_1(X91,X91)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_funct_2])])])).
cnf(c_0_97, plain, (v1_xboole_0(X1)|~r1_tarski(X1,X2)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_98, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_99, plain, ![X9, X10]:((~v4_relat_1(X10,X9)|r1_tarski(k1_relat_1(X10),X9)|~v1_relat_1(X10))&(~r1_tarski(k1_relat_1(X10),X9)|v4_relat_1(X10,X9)|~v1_relat_1(X10))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d18_relat_1])])])).
cnf(c_0_100, plain, (X1=esk4_2(X2,X3)|~v1_xboole_0(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_101, negated_conjecture, (v1_xboole_0(esk2_0)|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_80, c_0_45])).
cnf(c_0_102, plain, (~r2_hidden(X1,k9_relat_1(X2,X3))|~v1_xboole_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(spm,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_103, plain, (k2_zfmisc_1(X1,k1_xboole_0)=k1_xboole_0), inference(er,[status(thm)],[c_0_83])).
fof(c_0_104, plain, ![X32]:(~v1_relat_1(X32)|k9_relat_1(X32,k1_relat_1(X32))=k2_relat_1(X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t146_relat_1])])).
cnf(c_0_105, negated_conjecture, (v3_funct_2(k1_xboole_0,esk1_0,esk1_0)|k1_xboole_0!=esk1_0), inference(spm,[status(thm)],[c_0_44, c_0_84])).
cnf(c_0_106, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_107, plain, (v1_funct_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88])])).
cnf(c_0_108, plain, (v5_relat_1(k1_xboole_0,X1)|~v1_xboole_0(k2_zfmisc_1(X2,X1))), inference(spm,[status(thm)],[c_0_89, c_0_73])).
cnf(c_0_109, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_110, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_relat_1(X1),k2_relat_1(X1))))|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_111, plain, (k1_relat_1(X1)=k2_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_112, negated_conjecture, (k2_funct_1(esk2_0)=k2_funct_2(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_44]), c_0_45]), c_0_94]), c_0_46])])).
cnf(c_0_113, negated_conjecture, (v2_funct_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_44]), c_0_45]), c_0_46])])).
cnf(c_0_114, plain, (k2_relat_1(X1)=k1_relat_1(k2_funct_1(X1))|~v2_funct_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_115, plain, (v1_funct_1(k2_funct_2(X1,X2))|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X1)|~v3_funct_2(X2,X1,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_116, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_117, plain, (v1_xboole_0(X1)|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_118, plain, (r1_tarski(k1_relat_1(X1),X2)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_119, plain, (v4_relat_1(esk4_2(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_120, negated_conjecture, (esk4_2(X1,X2)=esk2_0|~v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_121, plain, (~r2_hidden(X1,k9_relat_1(X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103]), c_0_88])])).
cnf(c_0_122, plain, (k9_relat_1(X1,k1_relat_1(X1))=k2_relat_1(X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_123, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_48, c_0_87])).
cnf(c_0_124, negated_conjecture, (v2_funct_2(k1_xboole_0,esk1_0)|k1_xboole_0!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_105]), c_0_106]), c_0_107])])).
cnf(c_0_125, plain, (v5_relat_1(k1_xboole_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_87]), c_0_88])])).
cnf(c_0_126, plain, (v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_109, c_0_87])).
cnf(c_0_127, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(k2_zfmisc_1(k1_relat_1(X2),k2_relat_1(X2)))|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(spm,[status(thm)],[c_0_49, c_0_110])).
cnf(c_0_128, negated_conjecture, (k2_relat_1(k2_funct_2(esk1_0,esk2_0))=k1_relat_1(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_113]), c_0_60]), c_0_46])])).
cnf(c_0_129, negated_conjecture, (k1_relat_1(k2_funct_2(esk1_0,esk2_0))=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_112]), c_0_71]), c_0_113]), c_0_60]), c_0_46])])).
cnf(c_0_130, negated_conjecture, (v1_funct_1(k2_funct_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_44]), c_0_45]), c_0_94]), c_0_46])])).
cnf(c_0_131, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_116, c_0_62])).
cnf(c_0_132, plain, (v1_xboole_0(k1_relat_1(X1))|~v4_relat_1(X1,k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_133, negated_conjecture, (v4_relat_1(esk2_0,X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(esk1_0)), inference(spm,[status(thm)],[c_0_119, c_0_120])).
fof(c_0_134, plain, ![X15, X16]:(~v1_relat_1(X15)|(~m1_subset_1(X16,k1_zfmisc_1(X15))|v1_relat_1(X16))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
cnf(c_0_135, plain, (~r2_hidden(X1,k2_relat_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122]), c_0_123])).
cnf(c_0_136, negated_conjecture, (k2_relat_1(k1_xboole_0)=esk1_0|k1_xboole_0!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_124]), c_0_125]), c_0_126])])).
cnf(c_0_137, negated_conjecture, (~r2_hidden(X1,k2_funct_2(esk1_0,esk2_0))|~v1_xboole_0(k2_zfmisc_1(esk1_0,k1_relat_1(esk2_0)))|~v1_relat_1(k2_funct_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128]), c_0_129]), c_0_130])])).
cnf(c_0_138, plain, (k1_relat_1(X1)=k1_xboole_0|~v4_relat_1(X1,k1_xboole_0)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_139, negated_conjecture, (v4_relat_1(esk2_0,X1)|~v1_xboole_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_103]), c_0_88])])).
cnf(c_0_140, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_141, plain, (m1_subset_1(k2_funct_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X1)|~v3_funct_2(X2,X1,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_142, negated_conjecture, (k1_xboole_0!=esk1_0|~r2_hidden(X1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136]), c_0_106])])).
cnf(c_0_143, negated_conjecture, (v1_xboole_0(k2_funct_2(esk1_0,esk2_0))|~v1_xboole_0(k2_zfmisc_1(esk1_0,k1_relat_1(esk2_0)))|~v1_relat_1(k2_funct_2(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_137, c_0_65])).
cnf(c_0_144, negated_conjecture, (k1_relat_1(esk2_0)=k1_xboole_0|~v1_xboole_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_60])])).
cnf(c_0_145, plain, (v1_relat_1(k2_funct_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v3_funct_2(X2,X1,X1)|~v1_funct_2(X2,X1,X1)|~v1_funct_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_109])])).
cnf(c_0_146, negated_conjecture, (X1=esk2_0|~v1_xboole_0(esk1_0)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_78, c_0_101])).
cnf(c_0_147, negated_conjecture, (v1_xboole_0(esk1_0)|k1_xboole_0!=esk1_0), inference(spm,[status(thm)],[c_0_142, c_0_65])).
cnf(c_0_148, negated_conjecture, (v1_xboole_0(k2_funct_2(esk1_0,esk2_0))|~v1_xboole_0(esk1_0)|~v1_relat_1(k2_funct_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144]), c_0_103]), c_0_88])])).
cnf(c_0_149, negated_conjecture, (v1_relat_1(k2_funct_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_44]), c_0_45]), c_0_94]), c_0_46])])).
cnf(c_0_150, plain, (v3_funct_2(k2_funct_2(X1,X2),X1,X1)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X1)|~v3_funct_2(X2,X1,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_151, negated_conjecture, (X1=esk2_0|k1_xboole_0!=esk1_0|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_146, c_0_147])).
cnf(c_0_152, negated_conjecture, (v1_xboole_0(k2_funct_2(esk1_0,esk2_0))|~v1_xboole_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_149])])).
fof(c_0_153, plain, ![X61, X62, X63]:((r2_hidden(esk5_3(X61,X62,X63),X61)|r2_relset_1(X61,X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X61)))|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X61,X61))))&(k11_relat_1(X62,esk5_3(X61,X62,X63))!=k11_relat_1(X63,esk5_3(X61,X62,X63))|r2_relset_1(X61,X61,X62,X63)|~m1_subset_1(X63,k1_zfmisc_1(k2_zfmisc_1(X61,X61)))|~m1_subset_1(X62,k1_zfmisc_1(k2_zfmisc_1(X61,X61))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t54_relset_1])])])])])).
fof(c_0_154, plain, ![X51]:m1_subset_1(k6_relat_1(X51),k1_zfmisc_1(k2_zfmisc_1(X51,X51))), inference(variable_rename,[status(thm)],[t29_relset_1])).
fof(c_0_155, plain, ![X52]:k6_partfun1(X52)=k6_relat_1(X52), inference(variable_rename,[status(thm)],[redefinition_k6_partfun1])).
cnf(c_0_156, plain, (v2_funct_2(k2_funct_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v3_funct_2(X2,X1,X1)|~v1_funct_2(X2,X1,X1)|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_150]), c_0_115]), c_0_141])).
cnf(c_0_157, plain, (v5_relat_1(k2_funct_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v3_funct_2(X2,X1,X1)|~v1_funct_2(X2,X1,X1)|~v1_funct_1(X2)), inference(spm,[status(thm)],[c_0_47, c_0_141])).
cnf(c_0_158, negated_conjecture, (~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k2_funct_2(esk1_0,esk2_0)),k6_partfun1(esk1_0))|~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,k2_funct_2(esk1_0,esk2_0),esk2_0),k6_partfun1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_159, negated_conjecture, (k2_funct_2(esk1_0,esk2_0)=esk2_0|k1_xboole_0!=esk1_0), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_147])).
cnf(c_0_160, plain, (r2_hidden(esk5_3(X1,X2,X3),X1)|r2_relset_1(X1,X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_161, plain, (m1_subset_1(k6_relat_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_162, plain, (k6_partfun1(X1)=k6_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_155])).
fof(c_0_163, plain, ![X33, X34, X35]:((k5_relat_1(X35,k2_funct_1(X35))=k6_partfun1(X33)|X34=k1_xboole_0|(k2_relset_1(X33,X34,X35)!=X34|~v2_funct_1(X35))|(~v1_funct_1(X35)|(~v1_funct_2(X35,X33,X34)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(X33,X34))))))&(k5_relat_1(k2_funct_1(X35),X35)=k6_partfun1(X34)|X34=k1_xboole_0|(k2_relset_1(X33,X34,X35)!=X34|~v2_funct_1(X35))|(~v1_funct_1(X35)|(~v1_funct_2(X35,X33,X34)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(X33,X34))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_funct_2])])])).
cnf(c_0_164, plain, (k2_relat_1(k2_funct_2(X1,X2))=X1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v3_funct_2(X2,X1,X1)|~v1_funct_2(X2,X1,X1)|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_156]), c_0_145]), c_0_157])).
cnf(c_0_165, negated_conjecture, (k1_xboole_0!=esk1_0|~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0),k6_partfun1(esk1_0))), inference(spm,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_166, plain, (r2_relset_1(X1,X1,X2,X3)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(spm,[status(thm)],[c_0_116, c_0_160])).
cnf(c_0_167, plain, (m1_subset_1(k6_partfun1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(rw,[status(thm)],[c_0_161, c_0_162])).
fof(c_0_168, plain, ![X80, X81, X82, X83, X84, X85]:((v1_funct_1(k1_partfun1(X80,X81,X82,X83,X84,X85))|(~v1_funct_1(X84)|~m1_subset_1(X84,k1_zfmisc_1(k2_zfmisc_1(X80,X81)))|(~v1_funct_1(X85)|~m1_subset_1(X85,k1_zfmisc_1(k2_zfmisc_1(X82,X83))))))&(m1_subset_1(k1_partfun1(X80,X81,X82,X83,X84,X85),k1_zfmisc_1(k2_zfmisc_1(X80,X83)))|(~v1_funct_1(X84)|~m1_subset_1(X84,k1_zfmisc_1(k2_zfmisc_1(X80,X81)))|(~v1_funct_1(X85)|~m1_subset_1(X85,k1_zfmisc_1(k2_zfmisc_1(X82,X83))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
fof(c_0_169, plain, ![X39, X40, X41, X42, X43, X44]:(~v1_funct_1(X43)|~m1_subset_1(X43,k1_zfmisc_1(k2_zfmisc_1(X39,X40)))|(~v1_funct_1(X44)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(X41,X42))))|k1_partfun1(X39,X40,X41,X42,X43,X44)=k5_relat_1(X43,X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_170, negated_conjecture, (m1_subset_1(k2_funct_2(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k1_relat_1(esk2_0))))|~v1_relat_1(k2_funct_2(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_128]), c_0_129]), c_0_130])])).
cnf(c_0_171, plain, (k5_relat_1(k2_funct_1(X1),X1)=k6_partfun1(X2)|X2=k1_xboole_0|k2_relset_1(X3,X2,X1)!=X2|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_172, negated_conjecture, (k1_relat_1(esk2_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_44]), c_0_128]), c_0_45]), c_0_94]), c_0_46])])).
cnf(c_0_173, negated_conjecture, (k1_xboole_0!=esk1_0|~m1_subset_1(k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_167])]), c_0_147])).
cnf(c_0_174, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_175, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k5_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_176, negated_conjecture, (m1_subset_1(k2_funct_2(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k1_relat_1(esk2_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_170, c_0_149])])).
cnf(c_0_177, plain, (k5_relat_1(k2_funct_1(X1),X1)=k6_partfun1(k9_relat_1(X1,X2))|k9_relat_1(X1,X2)=k1_xboole_0|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k9_relat_1(X1,X2))))|~v1_funct_2(X1,X2,k9_relat_1(X1,X2))|~v1_funct_1(X1)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_82])])).
cnf(c_0_178, negated_conjecture, (k9_relat_1(esk2_0,esk1_0)=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_172]), c_0_71]), c_0_60])])).
cnf(c_0_179, negated_conjecture, (k1_xboole_0!=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_45]), c_0_46])])).
cnf(c_0_180, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_partfun1(X2)|X3=k1_xboole_0|k2_relset_1(X2,X3,X1)!=X3|~v2_funct_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_181, negated_conjecture, (~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k2_funct_2(esk1_0,esk2_0)),k6_partfun1(esk1_0))|~r2_relset_1(esk1_0,esk1_0,k5_relat_1(k2_funct_2(esk1_0,esk2_0),esk2_0),k6_partfun1(esk1_0))|~m1_subset_1(k2_funct_2(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_175]), c_0_45]), c_0_46]), c_0_130])])).
cnf(c_0_182, negated_conjecture, (m1_subset_1(k2_funct_2(esk1_0,esk2_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,esk1_0)))), inference(rw,[status(thm)],[c_0_176, c_0_172])).
cnf(c_0_183, negated_conjecture, (k5_relat_1(k2_funct_2(esk1_0,esk2_0),esk2_0)=k6_partfun1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178]), c_0_112]), c_0_113]), c_0_45]), c_0_94]), c_0_46])]), c_0_179])).
cnf(c_0_184, plain, (k5_relat_1(X1,k2_funct_1(X1))=k6_partfun1(X2)|k9_relat_1(X1,X2)=k1_xboole_0|~v2_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k9_relat_1(X1,X2))))|~v1_funct_2(X1,X2,k9_relat_1(X1,X2))|~v1_funct_1(X1)), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_82])])).
fof(c_0_185, plain, ![X19, X20, X21, X22]:(~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X19,X20)))|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X19,X20)))|r2_relset_1(X19,X20,X21,X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r2_relset_1])])).
cnf(c_0_186, negated_conjecture, (~r2_relset_1(esk1_0,esk1_0,k1_partfun1(esk1_0,esk1_0,esk1_0,esk1_0,esk2_0,k2_funct_2(esk1_0,esk2_0)),k6_partfun1(esk1_0))|~r2_relset_1(esk1_0,esk1_0,k6_partfun1(esk1_0),k6_partfun1(esk1_0))), inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_181, c_0_182])]), c_0_183])).
cnf(c_0_187, negated_conjecture, (k5_relat_1(esk2_0,k2_funct_2(esk1_0,esk2_0))=k6_partfun1(esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_178]), c_0_112]), c_0_113]), c_0_45]), c_0_94]), c_0_46])]), c_0_179])).
cnf(c_0_188, plain, (r2_relset_1(X2,X3,X1,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_189, negated_conjecture, (~r2_relset_1(esk1_0,esk1_0,k6_partfun1(esk1_0),k6_partfun1(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_175]), c_0_187]), c_0_182]), c_0_45]), c_0_130]), c_0_46])])).
cnf(c_0_190, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(spm,[status(thm)],[c_0_188, c_0_50])).
cnf(c_0_191, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_167])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 192
# Proof object clause steps            : 122
# Proof object formula steps           : 70
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 49
# Proof object initial formulas used   : 35
# Proof object generating inferences   : 66
# Proof object simplifying inferences  : 120
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 37
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 66
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 63
# Processed clauses                    : 8057
# ...of these trivial                  : 33
# ...subsumed                          : 6863
# ...remaining for further processing  : 1161
# Other redundant clauses eliminated   : 26
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 144
# Backward-rewritten                   : 157
# Generated clauses                    : 20235
# ...of the previous two non-trivial   : 18455
# Contextual simplify-reflections      : 60
# Paramodulations                      : 20208
# Factorizations                       : 0
# Equation resolutions                 : 27
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 794
#    Positive orientable unit clauses  : 69
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 712
# Current number of unprocessed clauses: 9606
# ...number of literals in the above   : 41709
# Current number of archived formulas  : 0
# Current number of archived clauses   : 365
# Clause-clause subsumption calls (NU) : 130213
# Rec. Clause-clause subsumption calls : 96037
# Non-unit clause-clause subsumptions  : 2411
# Unit Clause-clause subsumption calls : 3443
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 44
# BW rewrite match successes           : 21
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 282579

# -------------------------------------------------
# User time                : 0.245 s
# System time              : 0.003 s
# Total time               : 0.249 s
# Maximum resident set size: 28848 pages
