# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t24_yellow_1, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1)), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t24_yellow_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t3_subset)).
fof(d12_setfam_1, axiom, ![X1, X2]:(m1_setfam_1(X2,X1)<=>r1_tarski(X1,k3_tarski(X2))), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', d12_setfam_1)).
fof(t56_setfam_1, axiom, ![X1, X2, X3]:((r1_tarski(k3_tarski(X1),X2)&r2_hidden(X3,X1))=>r1_tarski(X3,X2)), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t56_setfam_1)).
fof(t95_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>r1_tarski(k3_tarski(X1),k3_tarski(X2))), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t95_zfmisc_1)).
fof(d1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v2_pre_topc(X1)<=>(r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))&(![X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))=>(r1_tarski(X2,u1_pre_topc(X1))=>r2_hidden(k5_setfam_1(u1_struct_0(X1),X2),u1_pre_topc(X1))))&![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((r2_hidden(X2,u1_pre_topc(X1))&r2_hidden(X3,u1_pre_topc(X1)))=>r2_hidden(k9_subset_1(u1_struct_0(X1),X2,X3),u1_pre_topc(X1))))))))), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', d1_pre_topc)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t5_subset)).
fof(t99_zfmisc_1, axiom, ![X1]:k3_tarski(k1_zfmisc_1(X1))=X1, file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t99_zfmisc_1)).
fof(t94_zfmisc_1, axiom, ![X1, X2]:(![X3]:(r2_hidden(X3,X1)=>r1_tarski(X3,X2))=>r1_tarski(k3_tarski(X1),X2)), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t94_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', d2_subset_1)).
fof(t60_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(m1_setfam_1(X2,X1)<=>k5_setfam_1(X1,X2)=X1)), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t60_setfam_1)).
fof(redefinition_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2)), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', redefinition_k5_setfam_1)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', dt_u1_pre_topc)).
fof(t14_yellow_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>(r2_hidden(k3_tarski(X1),X1)=>k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1))), file('data/MPTP/proofs/f39cd37a9a4446379cb6eba3258320ab.E_input', t14_yellow_1)).
fof(c_0_14, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>k4_yellow_0(k2_yellow_1(u1_pre_topc(X1)))=u1_struct_0(X1))), inference(assume_negation,[status(cth)],[t24_yellow_1])).
fof(c_0_15, plain, ![X27, X28]:((~m1_subset_1(X27,k1_zfmisc_1(X28))|r1_tarski(X27,X28))&(~r1_tarski(X27,X28)|m1_subset_1(X27,k1_zfmisc_1(X28)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_16, plain, ![X25, X26]:((~m1_setfam_1(X26,X25)|r1_tarski(X25,k3_tarski(X26)))&(~r1_tarski(X25,k3_tarski(X26))|m1_setfam_1(X26,X25))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d12_setfam_1])])).
fof(c_0_17, plain, ![X22, X23, X24]:(~r1_tarski(k3_tarski(X22),X23)|~r2_hidden(X24,X22)|r1_tarski(X24,X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t56_setfam_1])])).
fof(c_0_18, plain, ![X32, X33]:(~r1_tarski(X32,X33)|r1_tarski(k3_tarski(X32),k3_tarski(X33))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t95_zfmisc_1])])).
fof(c_0_19, plain, ![X5, X6, X7, X8]:(((r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|~v2_pre_topc(X5)|~l1_pre_topc(X5))&((~m1_subset_1(X6,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|(~r1_tarski(X6,u1_pre_topc(X5))|r2_hidden(k5_setfam_1(u1_struct_0(X5),X6),u1_pre_topc(X5)))|~v2_pre_topc(X5)|~l1_pre_topc(X5))&(~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X5)))|(~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X5)))|(~r2_hidden(X7,u1_pre_topc(X5))|~r2_hidden(X8,u1_pre_topc(X5))|r2_hidden(k9_subset_1(u1_struct_0(X5),X7,X8),u1_pre_topc(X5))))|~v2_pre_topc(X5)|~l1_pre_topc(X5))))&(((m1_subset_1(esk3_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&((m1_subset_1(esk4_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(((r2_hidden(esk3_1(X5),u1_pre_topc(X5))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(r2_hidden(esk4_1(X5),u1_pre_topc(X5))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))&(~r2_hidden(k9_subset_1(u1_struct_0(X5),esk3_1(X5),esk4_1(X5)),u1_pre_topc(X5))|m1_subset_1(esk2_1(X5),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X5))))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))))&(((m1_subset_1(esk3_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&((m1_subset_1(esk4_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(((r2_hidden(esk3_1(X5),u1_pre_topc(X5))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(r2_hidden(esk4_1(X5),u1_pre_topc(X5))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))&(~r2_hidden(k9_subset_1(u1_struct_0(X5),esk3_1(X5),esk4_1(X5)),u1_pre_topc(X5))|r1_tarski(esk2_1(X5),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))))&((m1_subset_1(esk3_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&((m1_subset_1(esk4_1(X5),k1_zfmisc_1(u1_struct_0(X5)))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(((r2_hidden(esk3_1(X5),u1_pre_topc(X5))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5))&(r2_hidden(esk4_1(X5),u1_pre_topc(X5))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))&(~r2_hidden(k9_subset_1(u1_struct_0(X5),esk3_1(X5),esk4_1(X5)),u1_pre_topc(X5))|~r2_hidden(k5_setfam_1(u1_struct_0(X5),esk2_1(X5)),u1_pre_topc(X5))|~r2_hidden(u1_struct_0(X5),u1_pre_topc(X5))|v2_pre_topc(X5)|~l1_pre_topc(X5)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_pre_topc])])])])])).
fof(c_0_20, negated_conjecture, ((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_14])])])])).
fof(c_0_21, plain, ![X29, X30, X31]:(~r2_hidden(X29,X30)|(~m1_subset_1(X30,k1_zfmisc_1(X31))|~v1_xboole_0(X31))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
cnf(c_0_22, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (r1_tarski(X2,k3_tarski(X1))|~m1_setfam_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (r1_tarski(X3,X2)|~r1_tarski(k3_tarski(X1),X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_25, plain, (r1_tarski(k3_tarski(X1),k3_tarski(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (r2_hidden(u1_struct_0(X1),u1_pre_topc(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_30, plain, (m1_subset_1(X1,k1_zfmisc_1(k3_tarski(X2)))|~m1_setfam_1(X2,X1)), inference(spm,[status(thm)],[c_0_22, c_0_23])).
fof(c_0_31, plain, ![X13]:k3_tarski(k1_zfmisc_1(X13))=X13, inference(variable_rename,[status(thm)],[t99_zfmisc_1])).
fof(c_0_32, plain, ![X19, X20]:((r2_hidden(esk5_2(X19,X20),X19)|r1_tarski(k3_tarski(X19),X20))&(~r1_tarski(esk5_2(X19,X20),X20)|r1_tarski(k3_tarski(X19),X20))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t94_zfmisc_1])])])])).
fof(c_0_33, plain, ![X16, X17]:(((~m1_subset_1(X17,X16)|r2_hidden(X17,X16)|v1_xboole_0(X16))&(~r2_hidden(X17,X16)|m1_subset_1(X17,X16)|v1_xboole_0(X16)))&((~m1_subset_1(X17,X16)|v1_xboole_0(X17)|~v1_xboole_0(X16))&(~v1_xboole_0(X17)|m1_subset_1(X17,X16)|~v1_xboole_0(X16)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
cnf(c_0_34, plain, (r1_tarski(X1,k3_tarski(X2))|~r1_tarski(X3,X2)|~r2_hidden(X1,X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_35, negated_conjecture, (r2_hidden(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_26, c_0_27]), c_0_28])])).
cnf(c_0_36, plain, (~v1_xboole_0(k3_tarski(X1))|~m1_setfam_1(X1,X2)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_37, plain, (k3_tarski(k1_zfmisc_1(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (m1_setfam_1(X2,X1)|~r1_tarski(X1,k3_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_39, plain, (r1_tarski(k3_tarski(X1),X2)|~r1_tarski(esk5_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_41, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (r2_hidden(esk5_2(X1,X2),X1)|r1_tarski(k3_tarski(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_43, plain, ![X14, X15]:((~m1_setfam_1(X15,X14)|k5_setfam_1(X14,X15)=X14|~m1_subset_1(X15,k1_zfmisc_1(k1_zfmisc_1(X14))))&(k5_setfam_1(X14,X15)!=X14|m1_setfam_1(X15,X14)|~m1_subset_1(X15,k1_zfmisc_1(k1_zfmisc_1(X14))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_setfam_1])])])).
cnf(c_0_44, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),k3_tarski(X1))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_45, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_46, plain, (~v1_xboole_0(X1)|~m1_setfam_1(k1_zfmisc_1(X1),X2)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_47, plain, (m1_setfam_1(k1_zfmisc_1(X1),X2)|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_37])).
cnf(c_0_48, plain, (r1_tarski(k3_tarski(X1),X2)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_49, plain, (v1_xboole_0(X1)|r1_tarski(k3_tarski(X1),X2)|m1_subset_1(esk5_2(X1,X2),X1)), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_50, plain, ![X34, X35]:(~m1_subset_1(X35,k1_zfmisc_1(k1_zfmisc_1(X34)))|k5_setfam_1(X34,X35)=k3_tarski(X35)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
cnf(c_0_51, plain, (k5_setfam_1(X2,X1)=X2|~m1_setfam_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, negated_conjecture, (m1_setfam_1(X1,u1_struct_0(esk1_0))|~r1_tarski(u1_pre_topc(esk1_0),X1)), inference(spm,[status(thm)],[c_0_38, c_0_44])).
cnf(c_0_53, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_zfmisc_1(k3_tarski(X2)))|~m1_setfam_1(X2,X1)), inference(spm,[status(thm)],[c_0_45, c_0_30])).
cnf(c_0_54, plain, (~v1_xboole_0(X1)|~r1_tarski(X2,X1)|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r1_tarski(X1,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49]), c_0_37])).
cnf(c_0_56, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_57, negated_conjecture, (k5_setfam_1(u1_struct_0(esk1_0),X1)=u1_struct_0(esk1_0)|~r1_tarski(u1_pre_topc(esk1_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_51, c_0_52])).
fof(c_0_58, plain, ![X12]:(~l1_pre_topc(X12)|m1_subset_1(u1_pre_topc(X12),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
cnf(c_0_59, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_zfmisc_1(X2))|~m1_setfam_1(k1_zfmisc_1(X2),X1)), inference(spm,[status(thm)],[c_0_53, c_0_37])).
cnf(c_0_60, plain, (v1_xboole_0(k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_54, c_0_55])).
fof(c_0_61, plain, ![X18]:(v1_xboole_0(X18)|(~r2_hidden(k3_tarski(X18),X18)|k4_yellow_0(k2_yellow_1(X18))=k3_tarski(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_yellow_1])])])).
cnf(c_0_62, negated_conjecture, (k3_tarski(X1)=u1_struct_0(esk1_0)|~r1_tarski(u1_pre_topc(esk1_0),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(spm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_63, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_64, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(spm,[status(thm)],[c_0_59, c_0_47])).
cnf(c_0_65, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_pre_topc(esk1_0)))|~v1_xboole_0(u1_pre_topc(esk1_0))), inference(spm,[status(thm)],[c_0_60, c_0_35])).
cnf(c_0_66, plain, (v1_xboole_0(X1)|k4_yellow_0(k2_yellow_1(X1))=k3_tarski(X1)|~r2_hidden(k3_tarski(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_67, negated_conjecture, (k3_tarski(u1_pre_topc(esk1_0))=u1_struct_0(esk1_0)|~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_28])])).
cnf(c_0_68, negated_conjecture, (k4_yellow_0(k2_yellow_1(u1_pre_topc(esk1_0)))!=u1_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_69, negated_conjecture, (v1_xboole_0(X1)|~v1_xboole_0(u1_pre_topc(esk1_0))|~r1_tarski(X1,u1_pre_topc(esk1_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_70, negated_conjecture, (v1_xboole_0(u1_pre_topc(esk1_0))|~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_35])]), c_0_68])).
cnf(c_0_71, negated_conjecture, (v1_xboole_0(X1)|~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))|~r1_tarski(X1,u1_pre_topc(esk1_0))), inference(spm,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_72, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r1_tarski(k3_tarski(X1),X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_60, c_0_42])).
cnf(c_0_73, negated_conjecture, (~m1_setfam_1(X1,X2)|~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))|~r1_tarski(k3_tarski(X1),u1_pre_topc(esk1_0))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_36, c_0_71])).
cnf(c_0_74, plain, (v1_xboole_0(k1_zfmisc_1(X1))|~v1_xboole_0(X2)|~v1_xboole_0(X1)|~r2_hidden(X3,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_54, c_0_72])).
cnf(c_0_75, negated_conjecture, (~m1_setfam_1(k1_zfmisc_1(X1),X2)|~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))|~r1_tarski(X1,u1_pre_topc(esk1_0))|~r2_hidden(X3,X2)), inference(spm,[status(thm)],[c_0_73, c_0_37])).
cnf(c_0_76, plain, (v1_xboole_0(k1_zfmisc_1(k3_tarski(X1)))|m1_setfam_1(X1,k3_tarski(X1))), inference(spm,[status(thm)],[c_0_38, c_0_55])).
cnf(c_0_77, plain, (v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_xboole_0(k1_zfmisc_1(X1))|~v1_xboole_0(X2)|~r2_hidden(X3,X1)), inference(spm,[status(thm)],[c_0_74, c_0_37])).
cnf(c_0_78, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(X1))|~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))|~r1_tarski(X1,u1_pre_topc(esk1_0))|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_37]), c_0_37])).
cnf(c_0_79, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_pre_topc(esk1_0))))|~v1_xboole_0(k1_zfmisc_1(u1_pre_topc(esk1_0)))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_77, c_0_35])).
cnf(c_0_80, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_pre_topc(esk1_0)))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_35]), c_0_55])).
cnf(c_0_81, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_pre_topc(esk1_0))))|~v1_xboole_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_80])])).
cnf(c_0_82, plain, (m1_setfam_1(X1,k3_tarski(X2))|~r1_tarski(X2,X1)), inference(spm,[status(thm)],[c_0_38, c_0_25])).
cnf(c_0_83, plain, (m1_subset_1(X1,X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_84, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_zfmisc_1(u1_pre_topc(esk1_0))))), inference(spm,[status(thm)],[c_0_81, c_0_80])).
cnf(c_0_85, plain, (r1_tarski(X1,X2)|~m1_setfam_1(k1_zfmisc_1(X2),X1)), inference(spm,[status(thm)],[c_0_23, c_0_37])).
cnf(c_0_86, plain, (m1_setfam_1(X1,X2)|~r1_tarski(k1_zfmisc_1(X2),X1)), inference(spm,[status(thm)],[c_0_82, c_0_37])).
cnf(c_0_87, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_pre_topc(esk1_0))))|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_88, plain, (r1_tarski(X1,X2)|~r1_tarski(k1_zfmisc_1(X1),k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_89, negated_conjecture, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_87]), c_0_80])])).
cnf(c_0_90, plain, (r1_tarski(X1,X2)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(spm,[status(thm)],[c_0_88, c_0_40])).
cnf(c_0_91, negated_conjecture, (~v1_xboole_0(u1_pre_topc(esk1_0))), inference(spm,[status(thm)],[c_0_89, c_0_35])).
cnf(c_0_92, negated_conjecture, (v1_xboole_0(X1)|~r1_tarski(X1,u1_pre_topc(esk1_0))), inference(spm,[status(thm)],[c_0_64, c_0_80])).
cnf(c_0_93, negated_conjecture, (r1_tarski(X1,u1_pre_topc(esk1_0))|~v1_xboole_0(k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_90, c_0_87])).
cnf(c_0_94, negated_conjecture, (~r1_tarski(u1_pre_topc(esk1_0),u1_pre_topc(esk1_0))), inference(spm,[status(thm)],[c_0_91, c_0_92])).
cnf(c_0_95, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_80]), c_0_94]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 96
# Proof object clause steps            : 67
# Proof object formula steps           : 29
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 21
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 45
# Proof object simplifying inferences  : 16
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 41
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 41
# Processed clauses                    : 650
# ...of these trivial                  : 0
# ...subsumed                          : 219
# ...remaining for further processing  : 431
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 54
# Backward-rewritten                   : 21
# Generated clauses                    : 1778
# ...of the previous two non-trivial   : 1640
# Contextual simplify-reflections      : 17
# Paramodulations                      : 1774
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 314
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 11
#    Non-unit-clauses                  : 289
# Current number of unprocessed clauses: 1032
# ...number of literals in the above   : 3748
# Current number of archived formulas  : 0
# Current number of archived clauses   : 117
# Clause-clause subsumption calls (NU) : 11515
# Rec. Clause-clause subsumption calls : 6531
# Non-unit clause-clause subsumptions  : 277
# Unit Clause-clause subsumption calls : 269
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 31445

# -------------------------------------------------
# User time                : 0.036 s
# System time              : 0.000 s
# Total time               : 0.036 s
# Maximum resident set size: 28676 pages
