# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_C36_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', t69_enumset1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', redefinition_k6_domain_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', d2_subset_1)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', d1_xboole_0)).
fof(t36_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', t36_tex_2)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', dt_l1_pre_topc)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', fc2_struct_0)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', dt_k6_domain_1)).
fof(t10_tsep_1, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>?[X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&m1_pre_topc(X3,X1)))&X2=u1_struct_0(X3)))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', t10_tsep_1)).
fof(t26_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&m1_pre_topc(X2,X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>(v2_tex_2(X3,X1)<=>v1_tdlat_3(X2)))))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', t26_tex_2)).
fof(d4_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((~(v2_struct_0(X3))&(v1_pre_topc(X3)&m1_pre_topc(X3,X1)))=>(X3=k1_tex_2(X1,X2)<=>u1_struct_0(X3)=k6_domain_1(u1_struct_0(X1),X2))))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', d4_tex_2)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', cc1_pre_topc)).
fof(t24_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v2_pre_topc(k1_tex_2(X1,X2))=>(v1_tdlat_3(k1_tex_2(X1,X2))&v2_tdlat_3(k1_tex_2(X1,X2)))))), file('data/MPTP/proofs/f3fed5ba8e9d414bb0fcb60c0003a7a7.E_input', t24_tex_2)).
fof(c_0_14, plain, ![X4, X5, X6, X7, X8, X9]:(((~r2_hidden(X6,X5)|X6=X4|X5!=k1_tarski(X4))&(X7!=X4|r2_hidden(X7,X5)|X5!=k1_tarski(X4)))&((~r2_hidden(esk1_2(X8,X9),X9)|esk1_2(X8,X9)!=X8|X9=k1_tarski(X8))&(r2_hidden(esk1_2(X8,X9),X9)|esk1_2(X8,X9)=X8|X9=k1_tarski(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_15, plain, ![X11]:k2_tarski(X11,X11)=k1_tarski(X11), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_16, plain, ![X12, X13]:(v1_xboole_0(X12)|~m1_subset_1(X13,X12)|k6_domain_1(X12,X13)=k1_tarski(X13)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_17, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_18, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_19, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, (r2_hidden(esk1_2(X1,X2),X2)|esk1_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X14, X15]:(((~m1_subset_1(X15,X14)|r2_hidden(X15,X14)|v1_xboole_0(X14))&(~r2_hidden(X15,X14)|m1_subset_1(X15,X14)|v1_xboole_0(X14)))&((~m1_subset_1(X15,X14)|v1_xboole_0(X15)|~v1_xboole_0(X14))&(~v1_xboole_0(X15)|m1_subset_1(X15,X14)|~v1_xboole_0(X14)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_22, plain, ![X16, X17, X18]:((~v1_xboole_0(X16)|~r2_hidden(X17,X16))&(r2_hidden(esk2_1(X18),X18)|v1_xboole_0(X18))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_23, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
fof(c_0_24, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)))), inference(assume_negation,[status(cth)],[t36_tex_2])).
cnf(c_0_25, plain, (k6_domain_1(X1,X2)=k2_tarski(X2,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_19, c_0_18])).
cnf(c_0_26, plain, (X2=k2_tarski(X1,X1)|esk1_2(X1,X2)=X1|r2_hidden(esk1_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_20, c_0_18])).
cnf(c_0_27, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X1)), inference(er,[status(thm)],[c_0_23])).
fof(c_0_30, plain, ![X22]:(~l1_pre_topc(X22)|l1_struct_0(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_31, negated_conjecture, ((~v2_struct_0(esk3_0)&(v2_pre_topc(esk3_0)&l1_pre_topc(esk3_0)))&(m1_subset_1(esk4_0,u1_struct_0(esk3_0))&~v2_tex_2(k6_domain_1(u1_struct_0(esk3_0),esk4_0),esk3_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_24])])])])).
cnf(c_0_32, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_33, plain, (esk1_2(X1,X2)=X1|k6_domain_1(X3,X1)=X2|v1_xboole_0(X3)|r2_hidden(esk1_2(X1,X2),X2)|~m1_subset_1(X1,X3)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_35, plain, (r2_hidden(X1,k2_tarski(X1,X1))), inference(er,[status(thm)],[c_0_29])).
fof(c_0_36, plain, ![X23]:(v2_struct_0(X23)|~l1_struct_0(X23)|~v1_xboole_0(u1_struct_0(X23))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_37, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_38, negated_conjecture, (l1_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_39, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_32, c_0_18])).
cnf(c_0_40, plain, (k6_domain_1(X1,X2)=X3|esk1_2(X2,X3)=X2|r2_hidden(esk1_2(X2,X3),X3)|~r2_hidden(X2,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_28])).
cnf(c_0_41, plain, (v1_xboole_0(X1)|r2_hidden(X2,k6_domain_1(X1,X2))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_35, c_0_25])).
cnf(c_0_42, negated_conjecture, (~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_43, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, negated_conjecture, (l1_struct_0(esk3_0)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_45, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_39])).
cnf(c_0_46, plain, (k6_domain_1(k6_domain_1(X1,X2),X2)=X3|esk1_2(X2,X3)=X2|v1_xboole_0(X1)|r2_hidden(esk1_2(X2,X3),X3)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_48, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43]), c_0_44])])).
cnf(c_0_49, plain, (X1=X2|v1_xboole_0(X3)|~m1_subset_1(X2,X3)|~r2_hidden(X1,k6_domain_1(X3,X2))), inference(spm,[status(thm)],[c_0_45, c_0_25])).
cnf(c_0_50, negated_conjecture, (k6_domain_1(k6_domain_1(u1_struct_0(esk3_0),esk4_0),esk4_0)=X1|esk1_2(esk4_0,X1)=esk4_0|r2_hidden(esk1_2(esk4_0,X1),X1)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_48])).
cnf(c_0_51, negated_conjecture, (k6_domain_1(k6_domain_1(u1_struct_0(esk3_0),esk4_0),esk4_0)=k6_domain_1(X1,X2)|esk1_2(esk4_0,k6_domain_1(X1,X2))=esk4_0|esk1_2(esk4_0,k6_domain_1(X1,X2))=X2|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_52, plain, (~v1_xboole_0(k2_tarski(X1,X1))), inference(spm,[status(thm)],[c_0_28, c_0_35])).
cnf(c_0_53, negated_conjecture, (k6_domain_1(k6_domain_1(u1_struct_0(esk3_0),esk4_0),esk4_0)=k6_domain_1(u1_struct_0(esk3_0),esk4_0)|esk1_2(esk4_0,k6_domain_1(u1_struct_0(esk3_0),esk4_0))=esk4_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_47]), c_0_48])).
cnf(c_0_54, plain, (esk1_2(X1,X2)=X1|~v1_xboole_0(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_26]), c_0_28])).
fof(c_0_55, plain, ![X24, X25]:(v1_xboole_0(X24)|~m1_subset_1(X25,X24)|m1_subset_1(k6_domain_1(X24,X25),k1_zfmisc_1(X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_56, plain, (X2=k1_tarski(X1)|~r2_hidden(esk1_2(X1,X2),X2)|esk1_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_57, negated_conjecture, (esk1_2(esk4_0,k6_domain_1(u1_struct_0(esk3_0),esk4_0))=esk4_0|k6_domain_1(u1_struct_0(esk3_0),esk4_0)=k2_tarski(esk4_0,esk4_0)|~m1_subset_1(esk4_0,k6_domain_1(u1_struct_0(esk3_0),esk4_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_53]), c_0_54])).
fof(c_0_58, plain, ![X26, X27]:(((~v2_struct_0(esk5_2(X26,X27))|(v1_xboole_0(X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26))))|(v2_struct_0(X26)|~l1_pre_topc(X26)))&((v1_pre_topc(esk5_2(X26,X27))|(v1_xboole_0(X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26))))|(v2_struct_0(X26)|~l1_pre_topc(X26)))&(m1_pre_topc(esk5_2(X26,X27),X26)|(v1_xboole_0(X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26))))|(v2_struct_0(X26)|~l1_pre_topc(X26)))))&(X27=u1_struct_0(esk5_2(X26,X27))|(v1_xboole_0(X27)|~m1_subset_1(X27,k1_zfmisc_1(u1_struct_0(X26))))|(v2_struct_0(X26)|~l1_pre_topc(X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_tsep_1])])])])])])).
cnf(c_0_59, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60, plain, (X2=k2_tarski(X1,X1)|esk1_2(X1,X2)!=X1|~r2_hidden(esk1_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_56, c_0_18])).
cnf(c_0_61, negated_conjecture, (k6_domain_1(u1_struct_0(esk3_0),esk4_0)=k2_tarski(esk4_0,esk4_0)|esk1_2(esk4_0,k6_domain_1(u1_struct_0(esk3_0),esk4_0))=esk4_0|~r2_hidden(esk4_0,k6_domain_1(u1_struct_0(esk3_0),esk4_0))), inference(spm,[status(thm)],[c_0_57, c_0_34])).
fof(c_0_62, plain, ![X32, X33, X34]:((~v2_tex_2(X34,X32)|v1_tdlat_3(X33)|X34!=u1_struct_0(X33)|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(X32)))|(v2_struct_0(X33)|~m1_pre_topc(X33,X32))|(v2_struct_0(X32)|~l1_pre_topc(X32)))&(~v1_tdlat_3(X33)|v2_tex_2(X34,X32)|X34!=u1_struct_0(X33)|~m1_subset_1(X34,k1_zfmisc_1(u1_struct_0(X32)))|(v2_struct_0(X33)|~m1_pre_topc(X33,X32))|(v2_struct_0(X32)|~l1_pre_topc(X32)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_tex_2])])])])])).
cnf(c_0_63, plain, (m1_pre_topc(esk5_2(X1,X2),X1)|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_64, plain, (m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(spm,[status(thm)],[c_0_59, c_0_25])).
cnf(c_0_65, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|~v2_struct_0(esk5_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_66, plain, (v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~v1_xboole_0(k6_domain_1(X1,X2))), inference(spm,[status(thm)],[c_0_52, c_0_25])).
cnf(c_0_67, negated_conjecture, (k6_domain_1(u1_struct_0(esk3_0),esk4_0)=k2_tarski(esk4_0,esk4_0)|~r2_hidden(esk4_0,k6_domain_1(u1_struct_0(esk3_0),esk4_0))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
fof(c_0_68, plain, ![X29, X30, X31]:((X31!=k1_tex_2(X29,X30)|u1_struct_0(X31)=k6_domain_1(u1_struct_0(X29),X30)|(v2_struct_0(X31)|(~v1_pre_topc(X31)|~m1_pre_topc(X31,X29)))|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|~l1_pre_topc(X29)))&(u1_struct_0(X31)!=k6_domain_1(u1_struct_0(X29),X30)|X31=k1_tex_2(X29,X30)|(v2_struct_0(X31)|(~v1_pre_topc(X31)|~m1_pre_topc(X31,X29)))|~m1_subset_1(X30,u1_struct_0(X29))|(v2_struct_0(X29)|~l1_pre_topc(X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d4_tex_2])])])])])).
cnf(c_0_69, plain, (X1=u1_struct_0(esk5_2(X2,X1))|v1_xboole_0(X1)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_70, plain, (v1_pre_topc(esk5_2(X1,X2))|v1_xboole_0(X2)|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_71, plain, (v2_tex_2(X2,X3)|v2_struct_0(X1)|v2_struct_0(X3)|~v1_tdlat_3(X1)|X2!=u1_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_pre_topc(X1,X3)|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_72, negated_conjecture, (~v2_tex_2(k6_domain_1(u1_struct_0(esk3_0),esk4_0),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_73, plain, ![X35, X36]:(~v2_pre_topc(X35)|~l1_pre_topc(X35)|(~m1_pre_topc(X36,X35)|v2_pre_topc(X36))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
cnf(c_0_74, plain, (m1_pre_topc(esk5_2(X1,k2_tarski(X2,X2)),X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64]), c_0_52])).
cnf(c_0_75, plain, (v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_struct_0(esk5_2(X1,k6_domain_1(u1_struct_0(X1),X2)))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_59]), c_0_66])).
cnf(c_0_76, negated_conjecture, (k6_domain_1(u1_struct_0(esk3_0),esk4_0)=k2_tarski(esk4_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_41]), c_0_47])]), c_0_48])).
cnf(c_0_77, plain, (X1=k1_tex_2(X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|u1_struct_0(X1)!=k6_domain_1(u1_struct_0(X2),X3)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_78, plain, (u1_struct_0(esk5_2(X1,k2_tarski(X2,X2)))=k2_tarski(X2,X2)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_64]), c_0_52])).
cnf(c_0_79, plain, (v1_pre_topc(esk5_2(X1,k2_tarski(X2,X2)))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_64]), c_0_52])).
cnf(c_0_80, plain, (v2_tex_2(k6_domain_1(u1_struct_0(X1),X2),X1)|v2_struct_0(X3)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|k6_domain_1(u1_struct_0(X1),X2)!=u1_struct_0(X3)|~v1_tdlat_3(X3)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_71, c_0_59])).
cnf(c_0_81, negated_conjecture, (~v2_tex_2(k2_tarski(esk4_0,esk4_0),esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_25]), c_0_47])]), c_0_48])).
fof(c_0_82, plain, ![X37, X38]:((v1_tdlat_3(k1_tex_2(X37,X38))|~v2_pre_topc(k1_tex_2(X37,X38))|~m1_subset_1(X38,u1_struct_0(X37))|(v2_struct_0(X37)|~l1_pre_topc(X37)))&(v2_tdlat_3(k1_tex_2(X37,X38))|~v2_pre_topc(k1_tex_2(X37,X38))|~m1_subset_1(X38,u1_struct_0(X37))|(v2_struct_0(X37)|~l1_pre_topc(X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_tex_2])])])])])).
cnf(c_0_83, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_84, negated_conjecture, (m1_pre_topc(esk5_2(esk3_0,k2_tarski(esk4_0,esk4_0)),esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_47]), c_0_38])]), c_0_42]), c_0_48])).
cnf(c_0_85, negated_conjecture, (v2_pre_topc(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_86, negated_conjecture, (~v2_struct_0(esk5_2(esk3_0,k2_tarski(esk4_0,esk4_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_38]), c_0_47])]), c_0_42]), c_0_48])).
cnf(c_0_87, negated_conjecture, (X1=k1_tex_2(esk3_0,esk4_0)|v2_struct_0(X1)|u1_struct_0(X1)!=k2_tarski(esk4_0,esk4_0)|~m1_pre_topc(X1,esk3_0)|~v1_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_76]), c_0_38]), c_0_47])]), c_0_42])).
cnf(c_0_88, negated_conjecture, (u1_struct_0(esk5_2(esk3_0,k2_tarski(esk4_0,esk4_0)))=k2_tarski(esk4_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_47]), c_0_38])]), c_0_42]), c_0_48])).
cnf(c_0_89, negated_conjecture, (v1_pre_topc(esk5_2(esk3_0,k2_tarski(esk4_0,esk4_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_47]), c_0_38])]), c_0_42]), c_0_48])).
cnf(c_0_90, negated_conjecture, (v2_struct_0(X1)|u1_struct_0(X1)!=k2_tarski(esk4_0,esk4_0)|~v1_tdlat_3(X1)|~m1_pre_topc(X1,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_76]), c_0_38]), c_0_47])]), c_0_81]), c_0_42]), c_0_48])).
cnf(c_0_91, plain, (v1_tdlat_3(k1_tex_2(X1,X2))|v2_struct_0(X1)|~v2_pre_topc(k1_tex_2(X1,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_92, negated_conjecture, (v2_pre_topc(esk5_2(esk3_0,k2_tarski(esk4_0,esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84]), c_0_38]), c_0_85])])).
cnf(c_0_93, negated_conjecture, (esk5_2(esk3_0,k2_tarski(esk4_0,esk4_0))=k1_tex_2(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87]), c_0_88]), c_0_84]), c_0_89])])).
cnf(c_0_94, negated_conjecture, (~v1_tdlat_3(esk5_2(esk3_0,k2_tarski(esk4_0,esk4_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_90]), c_0_88]), c_0_84])])).
cnf(c_0_95, negated_conjecture, (v1_tdlat_3(k1_tex_2(esk3_0,esk4_0))|~v2_pre_topc(k1_tex_2(esk3_0,esk4_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_47]), c_0_38])]), c_0_42])).
cnf(c_0_96, negated_conjecture, (v2_pre_topc(k1_tex_2(esk3_0,esk4_0))), inference(rw,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_97, negated_conjecture, (~v1_tdlat_3(k1_tex_2(esk3_0,esk4_0))), inference(rw,[status(thm)],[c_0_94, c_0_93])).
cnf(c_0_98, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_95, c_0_96])]), c_0_97]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 99
# Proof object clause steps            : 70
# Proof object formula steps           : 29
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 36
# Proof object simplifying inferences  : 69
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 14
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 31
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 30
# Processed clauses                    : 493
# ...of these trivial                  : 1
# ...subsumed                          : 238
# ...remaining for further processing  : 253
# Other redundant clauses eliminated   : 44
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 29
# Generated clauses                    : 4314
# ...of the previous two non-trivial   : 4140
# Contextual simplify-reflections      : 26
# Paramodulations                      : 4249
# Factorizations                       : 14
# Equation resolutions                 : 51
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 185
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 12
#    Non-unit-clauses                  : 159
# Current number of unprocessed clauses: 3702
# ...number of literals in the above   : 24945
# Current number of archived formulas  : 0
# Current number of archived clauses   : 68
# Clause-clause subsumption calls (NU) : 7560
# Rec. Clause-clause subsumption calls : 1790
# Non-unit clause-clause subsumptions  : 240
# Unit Clause-clause subsumption calls : 71
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 92605

# -------------------------------------------------
# User time                : 0.071 s
# System time              : 0.006 s
# Total time               : 0.078 s
# Maximum resident set size: 28744 pages
