# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t33_ordinal1, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2)))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', t33_ordinal1)).
fof(d1_ordinal1, axiom, ![X1]:k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1)), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', d1_ordinal1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', t69_enumset1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', t70_enumset1)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', d3_xboole_0)).
fof(t26_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r1_ordinal1(X1,X2)|r2_hidden(X2,X1)))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', t26_ordinal1)).
fof(t29_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>v3_ordinal1(k1_ordinal1(X1))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', t29_ordinal1)).
fof(l1_zfmisc_1, axiom, ![X1, X2]:(r1_tarski(k1_tarski(X1),X2)<=>r2_hidden(X1,X2)), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', l1_zfmisc_1)).
fof(t7_ordinal1, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', t7_ordinal1)).
fof(redefinition_r1_ordinal1, axiom, ![X1, X2]:((v3_ordinal1(X1)&v3_ordinal1(X2))=>(r1_ordinal1(X1,X2)<=>r1_tarski(X1,X2))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', redefinition_r1_ordinal1)).
fof(antisymmetry_r2_hidden, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', antisymmetry_r2_hidden)).
fof(t11_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k2_xboole_0(X1,X2),X3)=>r1_tarski(X1,X3)), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', t11_xboole_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', d10_xboole_0)).
fof(d1_enumset1, axiom, ![X1, X2, X3, X4]:(X4=k1_enumset1(X1,X2,X3)<=>![X5]:(r2_hidden(X5,X4)<=>~((~(X5=X1)&(~(X5=X2)&~(X5=X3)))))), file('data/MPTP/proofs/f488fe7f8a4d436fb498e4750a239ca9.E_input', d1_enumset1)).
fof(c_0_15, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:(v3_ordinal1(X2)=>(r2_hidden(X1,X2)<=>r1_ordinal1(k1_ordinal1(X1),X2))))), inference(assume_negation,[status(cth)],[t33_ordinal1])).
fof(c_0_16, plain, ![X33]:k1_ordinal1(X33)=k2_xboole_0(X33,k1_tarski(X33)), inference(variable_rename,[status(thm)],[d1_ordinal1])).
fof(c_0_17, plain, ![X49]:k2_tarski(X49,X49)=k1_tarski(X49), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_18, plain, ![X9, X10]:k3_tarski(k2_tarski(X9,X10))=k2_xboole_0(X9,X10), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_19, plain, ![X19, X20]:k1_enumset1(X19,X19,X20)=k2_tarski(X19,X20), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_20, negated_conjecture, (v3_ordinal1(esk1_0)&(v3_ordinal1(esk2_0)&((~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k1_ordinal1(esk1_0),esk2_0))&(r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k1_ordinal1(esk1_0),esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])).
cnf(c_0_21, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_25, plain, ![X21, X22, X23, X24, X25, X26, X27, X28]:(((~r2_hidden(X24,X23)|(r2_hidden(X24,X21)|r2_hidden(X24,X22))|X23!=k2_xboole_0(X21,X22))&((~r2_hidden(X25,X21)|r2_hidden(X25,X23)|X23!=k2_xboole_0(X21,X22))&(~r2_hidden(X25,X22)|r2_hidden(X25,X23)|X23!=k2_xboole_0(X21,X22))))&(((~r2_hidden(esk3_3(X26,X27,X28),X26)|~r2_hidden(esk3_3(X26,X27,X28),X28)|X28=k2_xboole_0(X26,X27))&(~r2_hidden(esk3_3(X26,X27,X28),X27)|~r2_hidden(esk3_3(X26,X27,X28),X28)|X28=k2_xboole_0(X26,X27)))&(r2_hidden(esk3_3(X26,X27,X28),X28)|(r2_hidden(esk3_3(X26,X27,X28),X26)|r2_hidden(esk3_3(X26,X27,X28),X27))|X28=k2_xboole_0(X26,X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_26, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (k1_ordinal1(X1)=k2_xboole_0(X1,k2_tarski(X1,X1))), inference(rw,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_28, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_23, c_0_24])).
fof(c_0_29, plain, ![X47, X48]:(~v3_ordinal1(X47)|(~v3_ordinal1(X48)|(r1_ordinal1(X47,X48)|r2_hidden(X48,X47)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t26_ordinal1])])])).
fof(c_0_30, plain, ![X8]:(~v3_ordinal1(X8)|v3_ordinal1(k1_ordinal1(X8))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t29_ordinal1])])).
fof(c_0_31, plain, ![X13, X14]:((~r1_tarski(k1_tarski(X13),X14)|r2_hidden(X13,X14))&(~r2_hidden(X13,X14)|r1_tarski(k1_tarski(X13),X14))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l1_zfmisc_1])])).
cnf(c_0_32, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X2!=k2_xboole_0(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)|~r1_ordinal1(k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_27]), c_0_24]), c_0_28])).
cnf(c_0_34, plain, (r1_ordinal1(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_35, negated_conjecture, (v3_ordinal1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_36, plain, (v3_ordinal1(k1_ordinal1(X1))|~v3_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_37, plain, ![X15, X16]:(~r2_hidden(X15,X16)|~r1_tarski(X16,X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_ordinal1])])).
cnf(c_0_38, plain, (r1_tarski(k1_tarski(X1),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
fof(c_0_39, plain, ![X11, X12]:((~r1_ordinal1(X11,X12)|r1_tarski(X11,X12)|(~v3_ordinal1(X11)|~v3_ordinal1(X12)))&(~r1_tarski(X11,X12)|r1_ordinal1(X11,X12)|(~v3_ordinal1(X11)|~v3_ordinal1(X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r1_ordinal1])])])).
cnf(c_0_40, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k3_tarski(k1_enumset1(X3,X3,X4))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_32, c_0_28])).
cnf(c_0_41, negated_conjecture, (r2_hidden(esk2_0,k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))))|~r2_hidden(esk1_0,esk2_0)|~v3_ordinal1(k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
cnf(c_0_42, plain, (v3_ordinal1(k3_tarski(k1_enumset1(X1,X1,k1_enumset1(X1,X1,X1))))|~v3_ordinal1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_27]), c_0_24]), c_0_28])).
cnf(c_0_43, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
fof(c_0_44, plain, ![X17, X18]:(~r2_hidden(X17,X18)|~r2_hidden(X18,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[antisymmetry_r2_hidden])])])).
cnf(c_0_45, plain, (~r2_hidden(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_46, plain, (r1_tarski(k1_enumset1(X1,X1,X1),X2)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_22]), c_0_24])).
cnf(c_0_47, plain, (r1_tarski(X1,X2)|~r1_ordinal1(X1,X2)|~v3_ordinal1(X1)|~v3_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_48, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k1_ordinal1(esk1_0),esk2_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_49, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k3_tarski(k1_enumset1(X3,X3,X2)))), inference(er,[status(thm)],[c_0_40])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk2_0,k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))))|~r2_hidden(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43])])).
cnf(c_0_51, plain, (~r2_hidden(X1,X2)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52, plain, (~r2_hidden(X1,k1_enumset1(X2,X2,X2))|~r2_hidden(X2,X1)), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_53, plain, (r1_tarski(X1,X2)|r2_hidden(X2,X1)|~v3_ordinal1(X2)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_47, c_0_34])).
cnf(c_0_54, negated_conjecture, (r2_hidden(esk1_0,esk2_0)|r1_ordinal1(k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))),esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_48, c_0_27]), c_0_24]), c_0_28])).
cnf(c_0_55, negated_conjecture, (~r2_hidden(esk1_0,esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51]), c_0_52])).
cnf(c_0_56, negated_conjecture, (r1_tarski(X1,esk1_0)|r2_hidden(esk1_0,X1)|~v3_ordinal1(X1)), inference(spm,[status(thm)],[c_0_53, c_0_43])).
fof(c_0_57, plain, ![X30, X31, X32]:(~r1_tarski(k2_xboole_0(X30,X31),X32)|r1_tarski(X30,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_xboole_1])])).
cnf(c_0_58, negated_conjecture, (r1_ordinal1(k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))),esk2_0)), inference(sr,[status(thm)],[c_0_54, c_0_55])).
fof(c_0_59, plain, ![X45, X46]:(((r1_tarski(X45,X46)|X45!=X46)&(r1_tarski(X46,X45)|X45!=X46))&(~r1_tarski(X45,X46)|~r1_tarski(X46,X45)|X45=X46)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_60, negated_conjecture, (r1_tarski(esk2_0,esk1_0)|r2_hidden(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_56, c_0_35])).
cnf(c_0_61, plain, (r1_tarski(X1,X3)|~r1_tarski(k2_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_62, negated_conjecture, (r1_tarski(k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))),esk2_0)|~v3_ordinal1(k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_58]), c_0_35])])).
cnf(c_0_63, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_64, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(sr,[status(thm)],[c_0_60, c_0_55])).
cnf(c_0_65, plain, (r1_tarski(X1,X3)|~r1_tarski(k3_tarski(k1_enumset1(X1,X1,X2)),X3)), inference(rw,[status(thm)],[c_0_61, c_0_28])).
cnf(c_0_66, negated_conjecture, (r1_tarski(k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_42]), c_0_43])])).
cnf(c_0_67, negated_conjecture, (esk2_0=esk1_0|~r1_tarski(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_68, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_69, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_70, plain, ![X34, X35, X36, X37, X38, X39, X40, X41, X42, X43]:(((~r2_hidden(X38,X37)|(X38=X34|(X38=X35|X38=X36))|X37!=k1_enumset1(X34,X35,X36))&((X39!=X34|r2_hidden(X39,X37)|X37!=k1_enumset1(X34,X35,X36))&((X39!=X35|r2_hidden(X39,X37)|X37!=k1_enumset1(X34,X35,X36))&(X39!=X36|r2_hidden(X39,X37)|X37!=k1_enumset1(X34,X35,X36)))))&(((esk4_4(X40,X41,X42,X43)!=X40|~r2_hidden(esk4_4(X40,X41,X42,X43),X43)|X43=k1_enumset1(X40,X41,X42))&((esk4_4(X40,X41,X42,X43)!=X41|~r2_hidden(esk4_4(X40,X41,X42,X43),X43)|X43=k1_enumset1(X40,X41,X42))&(esk4_4(X40,X41,X42,X43)!=X42|~r2_hidden(esk4_4(X40,X41,X42,X43),X43)|X43=k1_enumset1(X40,X41,X42))))&(r2_hidden(esk4_4(X40,X41,X42,X43),X43)|(esk4_4(X40,X41,X42,X43)=X40|(esk4_4(X40,X41,X42,X43)=X41|esk4_4(X40,X41,X42,X43)=X42))|X43=k1_enumset1(X40,X41,X42)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_enumset1])])])])])])])).
cnf(c_0_71, negated_conjecture, (~r2_hidden(esk2_0,k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))))), inference(spm,[status(thm)],[c_0_45, c_0_66])).
cnf(c_0_72, negated_conjecture, (esk2_0=esk1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68])])).
cnf(c_0_73, plain, (r2_hidden(X1,X3)|X3!=k3_tarski(k1_enumset1(X4,X4,X2))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_69, c_0_28])).
cnf(c_0_74, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_enumset1(X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_75, negated_conjecture, (~r2_hidden(esk1_0,k3_tarski(k1_enumset1(esk1_0,esk1_0,k1_enumset1(esk1_0,esk1_0,esk1_0))))), inference(rw,[status(thm)],[c_0_71, c_0_72])).
cnf(c_0_76, plain, (r2_hidden(X1,k3_tarski(k1_enumset1(X2,X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_73])).
cnf(c_0_77, plain, (r2_hidden(X1,k1_enumset1(X1,X2,X3))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_74])])).
cnf(c_0_78, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 79
# Proof object clause steps            : 48
# Proof object formula steps           : 31
# Proof object conjectures             : 24
# Proof object clause conjectures      : 21
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 37
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 34
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 30
# Processed clauses                    : 114
# ...of these trivial                  : 1
# ...subsumed                          : 15
# ...remaining for further processing  : 98
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 15
# Generated clauses                    : 102
# ...of the previous two non-trivial   : 99
# Contextual simplify-reflections      : 2
# Paramodulations                      : 89
# Factorizations                       : 0
# Equation resolutions                 : 12
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 36
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 42
# ...number of literals in the above   : 81
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 179
# Rec. Clause-clause subsumption calls : 131
# Non-unit clause-clause subsumptions  : 9
# Unit Clause-clause subsumption calls : 50
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 12
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3017

# -------------------------------------------------
# User time                : 0.009 s
# System time              : 0.003 s
# Total time               : 0.012 s
# Maximum resident set size: 28704 pages
