# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', d5_xboole_0)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t100_xboole_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t5_boole)).
fof(commutativity_k5_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', commutativity_k5_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', d4_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', commutativity_k3_xboole_0)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t48_xboole_1)).
fof(t7_xboole_0, axiom, ![X1]:~((~(X1=k1_xboole_0)&![X2]:~(r2_hidden(X2,X1)))), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t7_xboole_0)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', d1_zfmisc_1)).
fof(t25_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k4_subset_1(X1,X2,k3_subset_1(X1,X2))=k2_subset_1(X1)), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t25_subset_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t70_enumset1)).
fof(t91_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3)), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t91_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t28_xboole_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', fc1_subset_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', d5_subset_1)).
fof(t94_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t94_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t71_enumset1)).
fof(t72_enumset1, axiom, ![X1, X2, X3, X4]:k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t72_enumset1)).
fof(t73_enumset1, axiom, ![X1, X2, X3, X4, X5]:k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t73_enumset1)).
fof(t74_enumset1, axiom, ![X1, X2, X3, X4, X5, X6]:k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t74_enumset1)).
fof(t75_enumset1, axiom, ![X1, X2, X3, X4, X5, X6, X7]:k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t75_enumset1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', d4_subset_1)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', redefinition_k4_subset_1)).
fof(dt_k3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', dt_k3_subset_1)).
fof(t112_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2), file('data/MPTP/proofs/f4cd2cfc3e784ca091db4c4ca58b8a7a.E_input', t112_xboole_1)).
fof(c_0_27, plain, ![X8, X9, X10, X11, X12, X13, X14, X15]:((((r2_hidden(X11,X8)|~r2_hidden(X11,X10)|X10!=k4_xboole_0(X8,X9))&(~r2_hidden(X11,X9)|~r2_hidden(X11,X10)|X10!=k4_xboole_0(X8,X9)))&(~r2_hidden(X12,X8)|r2_hidden(X12,X9)|r2_hidden(X12,X10)|X10!=k4_xboole_0(X8,X9)))&((~r2_hidden(esk1_3(X13,X14,X15),X15)|(~r2_hidden(esk1_3(X13,X14,X15),X13)|r2_hidden(esk1_3(X13,X14,X15),X14))|X15=k4_xboole_0(X13,X14))&((r2_hidden(esk1_3(X13,X14,X15),X13)|r2_hidden(esk1_3(X13,X14,X15),X15)|X15=k4_xboole_0(X13,X14))&(~r2_hidden(esk1_3(X13,X14,X15),X14)|r2_hidden(esk1_3(X13,X14,X15),X15)|X15=k4_xboole_0(X13,X14))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
fof(c_0_28, plain, ![X17, X18]:k4_xboole_0(X17,X18)=k5_xboole_0(X17,k3_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
cnf(c_0_29, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_30, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_31, plain, ![X19]:k5_xboole_0(X19,k1_xboole_0)=X19, inference(variable_rename,[status(thm)],[t5_boole])).
fof(c_0_32, plain, ![X20, X21]:k5_xboole_0(X20,X21)=k5_xboole_0(X21,X20), inference(variable_rename,[status(thm)],[commutativity_k5_xboole_0])).
fof(c_0_33, plain, ![X22, X23, X24, X25, X26, X27, X28, X29]:((((r2_hidden(X25,X22)|~r2_hidden(X25,X24)|X24!=k3_xboole_0(X22,X23))&(r2_hidden(X25,X23)|~r2_hidden(X25,X24)|X24!=k3_xboole_0(X22,X23)))&(~r2_hidden(X26,X22)|~r2_hidden(X26,X23)|r2_hidden(X26,X24)|X24!=k3_xboole_0(X22,X23)))&((~r2_hidden(esk2_3(X27,X28,X29),X29)|(~r2_hidden(esk2_3(X27,X28,X29),X27)|~r2_hidden(esk2_3(X27,X28,X29),X28))|X29=k3_xboole_0(X27,X28))&((r2_hidden(esk2_3(X27,X28,X29),X27)|r2_hidden(esk2_3(X27,X28,X29),X29)|X29=k3_xboole_0(X27,X28))&(r2_hidden(esk2_3(X27,X28,X29),X28)|r2_hidden(esk2_3(X27,X28,X29),X29)|X29=k3_xboole_0(X27,X28))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_34, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_38, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_34])).
cnf(c_0_39, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(spm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_40, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X3,X2))), inference(er,[status(thm)],[c_0_37])).
fof(c_0_41, plain, ![X49, X50]:k3_xboole_0(X49,X50)=k3_xboole_0(X50,X49), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_42, plain, ![X47, X48]:k4_xboole_0(X47,k4_xboole_0(X47,X48))=k3_xboole_0(X47,X48), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_43, plain, (~r2_hidden(X1,k3_xboole_0(k1_xboole_0,X2))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_44, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_45, plain, ![X40]:(X40=k1_xboole_0|r2_hidden(esk6_1(X40),X40)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_xboole_0])])])])).
cnf(c_0_46, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_47, plain, (~r2_hidden(X1,k3_xboole_0(X2,k1_xboole_0))), inference(spm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_48, plain, (X1=k1_xboole_0|r2_hidden(esk6_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_49, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_30]), c_0_30])).
cnf(c_0_50, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_51, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X1))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_35])).
fof(c_0_52, plain, ![X31, X32, X33, X34, X35, X36]:(((~r2_hidden(X33,X32)|r1_tarski(X33,X31)|X32!=k1_zfmisc_1(X31))&(~r1_tarski(X34,X31)|r2_hidden(X34,X32)|X32!=k1_zfmisc_1(X31)))&((~r2_hidden(esk3_2(X35,X36),X36)|~r1_tarski(esk3_2(X35,X36),X35)|X36=k1_zfmisc_1(X35))&(r2_hidden(esk3_2(X35,X36),X36)|r1_tarski(esk3_2(X35,X36),X35)|X36=k1_zfmisc_1(X35)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_53, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k4_subset_1(X1,X2,k3_subset_1(X1,X2))=k2_subset_1(X1))), inference(assume_negation,[status(cth)],[t25_subset_1])).
fof(c_0_54, plain, ![X51, X52]:k3_tarski(k2_tarski(X51,X52))=k2_xboole_0(X51,X52), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_55, plain, ![X53, X54]:k1_enumset1(X53,X53,X54)=k2_tarski(X53,X54), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_56, plain, ![X93, X94, X95]:k5_xboole_0(k5_xboole_0(X93,X94),X95)=k5_xboole_0(X93,k5_xboole_0(X94,X95)), inference(variable_rename,[status(thm)],[t91_xboole_1])).
cnf(c_0_57, plain, (k3_xboole_0(X1,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_51]), c_0_50]), c_0_35])).
fof(c_0_58, plain, ![X42, X43]:(~r1_tarski(X42,X43)|k3_xboole_0(X42,X43)=X42), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_59, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_60, plain, ![X44, X45]:(((~m1_subset_1(X45,X44)|r2_hidden(X45,X44)|v1_xboole_0(X44))&(~r2_hidden(X45,X44)|m1_subset_1(X45,X44)|v1_xboole_0(X44)))&((~m1_subset_1(X45,X44)|v1_xboole_0(X45)|~v1_xboole_0(X44))&(~v1_xboole_0(X45)|m1_subset_1(X45,X44)|~v1_xboole_0(X44)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_61, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(esk4_0))&k4_subset_1(esk4_0,esk5_0,k3_subset_1(esk4_0,esk5_0))!=k2_subset_1(esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_53])])])).
fof(c_0_62, plain, ![X46]:~v1_xboole_0(k1_zfmisc_1(X46)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_63, plain, ![X58, X59]:(~m1_subset_1(X59,k1_zfmisc_1(X58))|k3_subset_1(X58,X59)=k4_xboole_0(X58,X59)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
fof(c_0_64, plain, ![X60, X61]:k2_xboole_0(X60,X61)=k5_xboole_0(k5_xboole_0(X60,X61),k3_xboole_0(X60,X61)), inference(variable_rename,[status(thm)],[t94_xboole_1])).
cnf(c_0_65, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_55])).
fof(c_0_67, plain, ![X62, X63, X64]:k2_enumset1(X62,X62,X63,X64)=k1_enumset1(X62,X63,X64), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_68, plain, ![X65, X66, X67, X68]:k3_enumset1(X65,X65,X66,X67,X68)=k2_enumset1(X65,X66,X67,X68), inference(variable_rename,[status(thm)],[t72_enumset1])).
fof(c_0_69, plain, ![X69, X70, X71, X72, X73]:k4_enumset1(X69,X69,X70,X71,X72,X73)=k3_enumset1(X69,X70,X71,X72,X73), inference(variable_rename,[status(thm)],[t73_enumset1])).
fof(c_0_70, plain, ![X74, X75, X76, X77, X78, X79]:k5_enumset1(X74,X74,X75,X76,X77,X78,X79)=k4_enumset1(X74,X75,X76,X77,X78,X79), inference(variable_rename,[status(thm)],[t74_enumset1])).
fof(c_0_71, plain, ![X80, X81, X82, X83, X84, X85, X86]:k6_enumset1(X80,X80,X81,X82,X83,X84,X85,X86)=k5_enumset1(X80,X81,X82,X83,X84,X85,X86), inference(variable_rename,[status(thm)],[t75_enumset1])).
cnf(c_0_72, plain, (k5_xboole_0(k5_xboole_0(X1,X2),X3)=k5_xboole_0(X1,k5_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_73, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_51, c_0_57])).
fof(c_0_74, plain, ![X87]:k2_subset_1(X87)=X87, inference(variable_rename,[status(thm)],[d4_subset_1])).
cnf(c_0_75, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_76, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_59])).
cnf(c_0_77, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_78, negated_conjecture, (m1_subset_1(esk5_0,k1_zfmisc_1(esk4_0))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_79, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_80, plain, ![X88, X89, X90]:(~m1_subset_1(X89,k1_zfmisc_1(X88))|~m1_subset_1(X90,k1_zfmisc_1(X88))|k4_subset_1(X88,X89,X90)=k2_xboole_0(X89,X90)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
fof(c_0_81, plain, ![X91, X92]:(~m1_subset_1(X92,k1_zfmisc_1(X91))|m1_subset_1(k3_subset_1(X91,X92),k1_zfmisc_1(X91))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_subset_1])])).
cnf(c_0_82, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_83, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_84, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_85, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_86, plain, (k3_enumset1(X1,X1,X2,X3,X4)=k2_enumset1(X1,X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_87, plain, (k4_enumset1(X1,X1,X2,X3,X4,X5)=k3_enumset1(X1,X2,X3,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_88, plain, (k5_enumset1(X1,X1,X2,X3,X4,X5,X6)=k4_enumset1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_89, plain, (k6_enumset1(X1,X1,X2,X3,X4,X5,X6,X7)=k5_enumset1(X1,X2,X3,X4,X5,X6,X7)), inference(split_conjunct,[status(thm)],[c_0_71])).
fof(c_0_90, plain, ![X55, X56, X57]:k5_xboole_0(k3_xboole_0(X55,X56),k3_xboole_0(X57,X56))=k3_xboole_0(k5_xboole_0(X55,X57),X56), inference(variable_rename,[status(thm)],[t112_xboole_1])).
cnf(c_0_91, plain, (k5_xboole_0(X1,k5_xboole_0(X1,X2))=X2), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_39])).
cnf(c_0_92, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2)))=k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_49, c_0_49])).
cnf(c_0_93, negated_conjecture, (k4_subset_1(esk4_0,esk5_0,k3_subset_1(esk4_0,esk5_0))!=k2_subset_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_94, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_95, plain, (k3_xboole_0(X1,X2)=X1|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_96, negated_conjecture, (r2_hidden(esk5_0,k1_zfmisc_1(esk4_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_79])).
cnf(c_0_97, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_98, plain, (m1_subset_1(k3_subset_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_99, plain, (k3_subset_1(X2,X1)=k5_xboole_0(X2,k3_xboole_0(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_82, c_0_30])).
cnf(c_0_100, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89])).
cnf(c_0_101, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_102, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_49])).
cnf(c_0_103, negated_conjecture, (k4_subset_1(esk4_0,esk5_0,k3_subset_1(esk4_0,esk5_0))!=esk4_0), inference(rw,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_104, negated_conjecture, (k3_xboole_0(esk4_0,esk5_0)=esk5_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_44])).
cnf(c_0_105, plain, (k4_subset_1(X2,X1,X3)=k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_97, c_0_84]), c_0_85]), c_0_86]), c_0_87]), c_0_88]), c_0_89])).
cnf(c_0_106, plain, (m1_subset_1(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_107, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,X2))=k5_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_100, c_0_72])).
cnf(c_0_108, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X2,X3))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_101, c_0_44])).
cnf(c_0_109, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_102, c_0_44])).
cnf(c_0_110, negated_conjecture, (k4_subset_1(esk4_0,esk5_0,k5_xboole_0(esk4_0,esk5_0))!=esk4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_99]), c_0_104]), c_0_78])])).
cnf(c_0_111, plain, (k4_subset_1(X1,X2,X3)=k4_subset_1(X4,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X4))|~m1_subset_1(X2,k1_zfmisc_1(X4))), inference(spm,[status(thm)],[c_0_105, c_0_105])).
cnf(c_0_112, negated_conjecture, (m1_subset_1(k5_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(esk4_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_104]), c_0_78])])).
cnf(c_0_113, plain, (k3_tarski(k6_enumset1(X1,X1,X1,X1,X1,X1,X1,k5_xboole_0(X2,X3)))=k5_xboole_0(X1,k5_xboole_0(X2,k5_xboole_0(X3,k3_xboole_0(X1,k5_xboole_0(X2,X3)))))), inference(spm,[status(thm)],[c_0_107, c_0_72])).
cnf(c_0_114, plain, (k3_xboole_0(X1,k5_xboole_0(X2,k3_xboole_0(X3,X1)))=k3_xboole_0(k5_xboole_0(X2,X3),X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109]), c_0_101]), c_0_44])).
cnf(c_0_115, negated_conjecture, (k4_subset_1(X1,esk5_0,k5_xboole_0(esk4_0,esk5_0))!=esk4_0|~m1_subset_1(k5_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(X1))|~m1_subset_1(esk5_0,k1_zfmisc_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112]), c_0_78])])).
cnf(c_0_116, plain, (k4_subset_1(X1,X2,k5_xboole_0(X3,X4))=k5_xboole_0(X2,k5_xboole_0(X3,k5_xboole_0(X4,k3_xboole_0(X2,k5_xboole_0(X3,X4)))))|~m1_subset_1(k5_xboole_0(X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_105, c_0_113])).
cnf(c_0_117, negated_conjecture, (k3_xboole_0(esk5_0,k5_xboole_0(X1,esk5_0))=k3_xboole_0(esk5_0,k5_xboole_0(X1,esk4_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_104]), c_0_44])).
cnf(c_0_118, plain, (k5_xboole_0(X1,k5_xboole_0(X2,X1))=X2), inference(spm,[status(thm)],[c_0_91, c_0_36])).
cnf(c_0_119, negated_conjecture, (~m1_subset_1(k5_xboole_0(esk4_0,esk5_0),k1_zfmisc_1(X1))|~m1_subset_1(esk5_0,k1_zfmisc_1(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117]), c_0_73]), c_0_50]), c_0_35]), c_0_118])])).
cnf(c_0_120, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_112]), c_0_78])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 121
# Proof object clause steps            : 66
# Proof object formula steps           : 55
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 27
# Proof object generating inferences   : 26
# Proof object simplifying inferences  : 50
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 27
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 35
# Processed clauses                    : 10200
# ...of these trivial                  : 243
# ...subsumed                          : 8688
# ...remaining for further processing  : 1269
# Other redundant clauses eliminated   : 8
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 47
# Backward-rewritten                   : 152
# Generated clauses                    : 205075
# ...of the previous two non-trivial   : 197900
# Contextual simplify-reflections      : 26
# Paramodulations                      : 204983
# Factorizations                       : 84
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1027
#    Positive orientable unit clauses  : 135
#    Positive unorientable unit clauses: 58
#    Negative unit clauses             : 21
#    Non-unit-clauses                  : 813
# Current number of unprocessed clauses: 186506
# ...number of literals in the above   : 457050
# Current number of archived formulas  : 0
# Current number of archived clauses   : 243
# Clause-clause subsumption calls (NU) : 112935
# Rec. Clause-clause subsumption calls : 65965
# Non-unit clause-clause subsumptions  : 5222
# Unit Clause-clause subsumption calls : 2482
# Rewrite failures with RHS unbound    : 60
# BW rewrite match attempts            : 3173
# BW rewrite match successes           : 538
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4504841

# -------------------------------------------------
# User time                : 2.841 s
# System time              : 0.047 s
# Total time               : 2.888 s
# Maximum resident set size: 186428 pages
