# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', d1_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', fc1_subset_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t28_xboole_1)).
fof(t25_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k4_subset_1(X1,X2,k3_subset_1(X1,X2))=k2_subset_1(X1)), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t25_subset_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', commutativity_k3_xboole_0)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', l51_zfmisc_1)).
fof(t70_enumset1, axiom, ![X1, X2]:k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t70_enumset1)).
fof(t112_xboole_1, axiom, ![X1, X2, X3]:k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t112_xboole_1)).
fof(d6_xboole_0, axiom, ![X1, X2]:k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1)), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', d6_xboole_0)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t100_xboole_1)).
fof(t71_enumset1, axiom, ![X1, X2, X3]:k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t71_enumset1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', commutativity_k2_tarski)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t22_xboole_1)).
fof(t46_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0, file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t46_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t48_xboole_1)).
fof(idempotence_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', idempotence_k3_xboole_0)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', d5_subset_1)).
fof(t94_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t94_xboole_1)).
fof(d4_subset_1, axiom, ![X1]:k2_subset_1(X1)=X1, file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', d4_subset_1)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', redefinition_k4_subset_1)).
fof(t5_boole, axiom, ![X1]:k5_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', t5_boole)).
fof(dt_k3_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>m1_subset_1(k3_subset_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/f5f0095180ce45b49f0cf46420daa5c2.E_input', dt_k3_subset_1)).
fof(c_0_23, plain, ![X4, X5, X6, X7, X8, X9]:(((~r2_hidden(X6,X5)|r1_tarski(X6,X4)|X5!=k1_zfmisc_1(X4))&(~r1_tarski(X7,X4)|r2_hidden(X7,X5)|X5!=k1_zfmisc_1(X4)))&((~r2_hidden(esk1_2(X8,X9),X9)|~r1_tarski(esk1_2(X8,X9),X8)|X9=k1_zfmisc_1(X8))&(r2_hidden(esk1_2(X8,X9),X9)|r1_tarski(esk1_2(X8,X9),X8)|X9=k1_zfmisc_1(X8)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
cnf(c_0_24, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_25, plain, ![X11, X12]:(((~m1_subset_1(X12,X11)|r2_hidden(X12,X11)|v1_xboole_0(X11))&(~r2_hidden(X12,X11)|m1_subset_1(X12,X11)|v1_xboole_0(X11)))&((~m1_subset_1(X12,X11)|v1_xboole_0(X12)|~v1_xboole_0(X11))&(~v1_xboole_0(X12)|m1_subset_1(X12,X11)|~v1_xboole_0(X11)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_26, plain, ![X13]:~v1_xboole_0(k1_zfmisc_1(X13)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_27, plain, ![X14, X15]:(~r1_tarski(X14,X15)|k3_xboole_0(X14,X15)=X14), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_28, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k4_subset_1(X1,X2,k3_subset_1(X1,X2))=k2_subset_1(X1))), inference(assume_negation,[status(cth)],[t25_subset_1])).
cnf(c_0_32, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_33, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29]), c_0_30])).
fof(c_0_34, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk2_0))&k4_subset_1(esk2_0,esk3_0,k3_subset_1(esk2_0,esk3_0))!=k2_subset_1(esk2_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_31])])])).
fof(c_0_35, plain, ![X18, X19]:k3_xboole_0(X18,X19)=k3_xboole_0(X19,X18), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_36, plain, ![X20, X21]:k3_tarski(k2_tarski(X20,X21))=k2_xboole_0(X20,X21), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
fof(c_0_37, plain, ![X22, X23]:k1_enumset1(X22,X22,X23)=k2_tarski(X22,X23), inference(variable_rename,[status(thm)],[t70_enumset1])).
fof(c_0_38, plain, ![X24, X25, X26]:k5_xboole_0(k3_xboole_0(X24,X25),k3_xboole_0(X26,X25))=k3_xboole_0(k5_xboole_0(X24,X26),X25), inference(variable_rename,[status(thm)],[t112_xboole_1])).
cnf(c_0_39, plain, (k3_xboole_0(X1,X2)=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_40, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_42, plain, ![X40, X41]:k5_xboole_0(X40,X41)=k2_xboole_0(k4_xboole_0(X40,X41),k4_xboole_0(X41,X40)), inference(variable_rename,[status(thm)],[d6_xboole_0])).
cnf(c_0_43, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_44, plain, (k1_enumset1(X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
fof(c_0_45, plain, ![X34, X35]:k4_xboole_0(X34,X35)=k5_xboole_0(X34,k3_xboole_0(X34,X35)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_46, plain, ![X29, X30, X31]:k2_enumset1(X29,X29,X30,X31)=k1_enumset1(X29,X30,X31), inference(variable_rename,[status(thm)],[t71_enumset1])).
fof(c_0_47, plain, ![X42, X43]:k2_tarski(X42,X43)=k2_tarski(X43,X42), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_48, plain, ![X27, X28]:k2_xboole_0(X27,k3_xboole_0(X27,X28))=X27, inference(variable_rename,[status(thm)],[t22_xboole_1])).
cnf(c_0_49, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X3,X2))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_50, negated_conjecture, (k3_xboole_0(esk2_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40]), c_0_41])).
cnf(c_0_51, plain, (k5_xboole_0(X1,X2)=k2_xboole_0(k4_xboole_0(X1,X2),k4_xboole_0(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_52, plain, (k3_tarski(k1_enumset1(X1,X1,X2))=k2_xboole_0(X1,X2)), inference(rw,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_53, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, plain, (k2_enumset1(X1,X1,X2,X3)=k1_enumset1(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_56, plain, ![X32, X33]:k4_xboole_0(X32,k2_xboole_0(X32,X33))=k1_xboole_0, inference(variable_rename,[status(thm)],[t46_xboole_1])).
cnf(c_0_57, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, negated_conjecture, (k5_xboole_0(k3_xboole_0(X1,esk3_0),esk3_0)=k3_xboole_0(esk3_0,k5_xboole_0(X1,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_41])).
fof(c_0_59, plain, ![X36, X37]:k4_xboole_0(X36,k4_xboole_0(X36,X37))=k3_xboole_0(X36,X37), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_60, plain, ![X44]:k3_xboole_0(X44,X44)=X44, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k3_xboole_0])])).
cnf(c_0_61, plain, (k5_xboole_0(X1,X2)=k3_tarski(k2_enumset1(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X1,k3_xboole_0(X1,X2)),k5_xboole_0(X2,k3_xboole_0(X2,X1))))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52]), c_0_53]), c_0_53]), c_0_53]), c_0_54])).
cnf(c_0_62, plain, (k2_enumset1(X1,X1,X1,X2)=k2_enumset1(X2,X2,X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_44]), c_0_44]), c_0_54]), c_0_54])).
cnf(c_0_63, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (k3_tarski(k2_enumset1(X1,X1,X1,k3_xboole_0(X1,X2)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_52]), c_0_54])).
cnf(c_0_65, negated_conjecture, (k3_xboole_0(esk3_0,k5_xboole_0(esk2_0,esk2_0))=k5_xboole_0(esk3_0,esk3_0)), inference(spm,[status(thm)],[c_0_58, c_0_50])).
cnf(c_0_66, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_67, plain, ![X38, X39]:(~m1_subset_1(X39,k1_zfmisc_1(X38))|k3_subset_1(X38,X39)=k4_xboole_0(X38,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_68, plain, (k3_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_69, plain, (k5_xboole_0(X1,X2)=k5_xboole_0(X2,X1)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_62]), c_0_61])).
fof(c_0_70, plain, ![X51, X52]:k2_xboole_0(X51,X52)=k5_xboole_0(k5_xboole_0(X51,X52),k3_xboole_0(X51,X52)), inference(variable_rename,[status(thm)],[t94_xboole_1])).
cnf(c_0_71, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k3_tarski(k2_enumset1(X1,X1,X1,X2))))=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_52]), c_0_53]), c_0_54])).
cnf(c_0_72, negated_conjecture, (k3_tarski(k2_enumset1(esk3_0,esk3_0,esk3_0,k5_xboole_0(esk3_0,esk3_0)))=esk3_0), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_73, negated_conjecture, (k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,esk3_0))=esk2_0), inference(spm,[status(thm)],[c_0_64, c_0_50])).
cnf(c_0_74, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_53]), c_0_53])).
cnf(c_0_75, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76, plain, (k5_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X1,k5_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_68]), c_0_69]), c_0_41])).
cnf(c_0_77, plain, (k5_xboole_0(k3_xboole_0(X1,X2),k3_xboole_0(X2,X3))=k3_xboole_0(k5_xboole_0(X1,X3),X2)), inference(spm,[status(thm)],[c_0_49, c_0_41])).
fof(c_0_78, plain, ![X45]:k2_subset_1(X45)=X45, inference(variable_rename,[status(thm)],[d4_subset_1])).
fof(c_0_79, plain, ![X46, X47, X48]:(~m1_subset_1(X47,k1_zfmisc_1(X46))|~m1_subset_1(X48,k1_zfmisc_1(X46))|k4_subset_1(X46,X47,X48)=k2_xboole_0(X47,X48)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
cnf(c_0_80, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_81, negated_conjecture, (k5_xboole_0(esk3_0,esk3_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_68])).
cnf(c_0_82, negated_conjecture, (k5_xboole_0(esk2_0,esk2_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_73]), c_0_68])).
fof(c_0_83, plain, ![X53]:k5_xboole_0(X53,k1_xboole_0)=X53, inference(variable_rename,[status(thm)],[t5_boole])).
cnf(c_0_84, negated_conjecture, (k5_xboole_0(esk2_0,k3_xboole_0(esk2_0,k5_xboole_0(esk2_0,esk3_0)))=esk3_0), inference(spm,[status(thm)],[c_0_74, c_0_50])).
cnf(c_0_85, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_75, c_0_53])).
cnf(c_0_86, plain, (k5_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,k5_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_76, c_0_41])).
cnf(c_0_87, negated_conjecture, (k5_xboole_0(esk3_0,k3_xboole_0(X1,esk2_0))=k3_xboole_0(esk2_0,k5_xboole_0(X1,esk3_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_50]), c_0_69]), c_0_41])).
cnf(c_0_88, negated_conjecture, (k4_subset_1(esk2_0,esk3_0,k3_subset_1(esk2_0,esk3_0))!=k2_subset_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_89, plain, (k2_subset_1(X1)=X1), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_90, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_91, plain, (k3_tarski(k2_enumset1(X1,X1,X1,X2))=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_52]), c_0_54])).
cnf(c_0_92, negated_conjecture, (k3_xboole_0(esk3_0,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_81]), c_0_82])).
cnf(c_0_93, plain, (k5_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_94, plain, (k5_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_64]), c_0_68])).
cnf(c_0_95, negated_conjecture, (k3_xboole_0(esk2_0,k5_xboole_0(esk2_0,esk3_0))=k5_xboole_0(esk2_0,esk3_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_84]), c_0_50])).
cnf(c_0_96, plain, (k3_xboole_0(X1,k5_xboole_0(X2,X1))=k3_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(rw,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_97, negated_conjecture, (k5_xboole_0(esk3_0,k3_xboole_0(esk2_0,X1))=k3_xboole_0(esk2_0,k5_xboole_0(X1,esk3_0))), inference(spm,[status(thm)],[c_0_87, c_0_41])).
cnf(c_0_98, negated_conjecture, (k4_subset_1(esk2_0,esk3_0,k3_subset_1(esk2_0,esk3_0))!=esk2_0), inference(rw,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_99, plain, (k4_subset_1(X2,X1,X3)=k3_tarski(k2_enumset1(X1,X1,X1,X3))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_52]), c_0_54])).
cnf(c_0_100, plain, (k3_tarski(k2_enumset1(X1,X1,X1,X2))=k5_xboole_0(k3_xboole_0(X1,X2),k5_xboole_0(X1,X2))), inference(rw,[status(thm)],[c_0_91, c_0_69])).
cnf(c_0_101, negated_conjecture, (k5_xboole_0(esk3_0,k3_xboole_0(X1,esk3_0))=k3_xboole_0(esk3_0,k5_xboole_0(esk2_0,X1))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_41])).
cnf(c_0_102, negated_conjecture, (k3_xboole_0(k1_xboole_0,k5_xboole_0(X1,esk3_0))=k3_xboole_0(X1,k1_xboole_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_92]), c_0_93]), c_0_41])).
cnf(c_0_103, plain, (k3_xboole_0(k1_xboole_0,X1)=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_94]), c_0_94])).
cnf(c_0_104, negated_conjecture, (k3_tarski(k2_enumset1(esk2_0,esk2_0,esk2_0,k5_xboole_0(esk2_0,esk3_0)))=esk2_0), inference(spm,[status(thm)],[c_0_64, c_0_95])).
cnf(c_0_105, negated_conjecture, (k5_xboole_0(esk2_0,k5_xboole_0(esk2_0,esk3_0))=esk3_0), inference(rw,[status(thm)],[c_0_84, c_0_95])).
cnf(c_0_106, plain, (k3_tarski(k2_enumset1(X1,X1,X1,X1))=X1), inference(spm,[status(thm)],[c_0_64, c_0_68])).
fof(c_0_107, plain, ![X49, X50]:(~m1_subset_1(X50,k1_zfmisc_1(X49))|m1_subset_1(k3_subset_1(X49,X50),k1_zfmisc_1(X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_subset_1])])).
cnf(c_0_108, negated_conjecture, (k3_xboole_0(k3_xboole_0(esk2_0,X1),k3_xboole_0(esk2_0,k5_xboole_0(X1,esk3_0)))=k3_subset_1(k3_xboole_0(esk2_0,X1),esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k3_xboole_0(esk2_0,X1)))), inference(spm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_109, negated_conjecture, (k4_subset_1(esk2_0,esk3_0,k5_xboole_0(esk2_0,esk3_0))!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_85]), c_0_50]), c_0_40])])).
cnf(c_0_110, plain, (k4_subset_1(X1,X2,X3)=k5_xboole_0(k3_xboole_0(X2,X3),k5_xboole_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_111, negated_conjecture, (k3_xboole_0(esk3_0,k5_xboole_0(X1,esk3_0))=k3_xboole_0(esk3_0,k5_xboole_0(esk2_0,X1))), inference(rw,[status(thm)],[c_0_101, c_0_76])).
cnf(c_0_112, negated_conjecture, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(rw,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_113, negated_conjecture, (k5_xboole_0(esk3_0,k5_xboole_0(esk2_0,esk3_0))=esk2_0), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_91]), c_0_105]), c_0_95])).
cnf(c_0_114, plain, (k5_xboole_0(k1_xboole_0,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_106, c_0_91]), c_0_94]), c_0_68])).
cnf(c_0_115, plain, (m1_subset_1(k3_subset_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_116, negated_conjecture, (k3_subset_1(esk2_0,esk3_0)=k5_xboole_0(esk2_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_68]), c_0_95]), c_0_95]), c_0_40])])).
cnf(c_0_117, negated_conjecture, (~m1_subset_1(k5_xboole_0(esk2_0,esk3_0),k1_zfmisc_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_111]), c_0_94]), c_0_112]), c_0_113]), c_0_114]), c_0_40])])).
cnf(c_0_118, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_40])]), c_0_117]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 119
# Proof object clause steps            : 72
# Proof object formula steps           : 47
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 23
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 70
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 23
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 30
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 25
# Processed clauses                    : 4253
# ...of these trivial                  : 357
# ...subsumed                          : 3102
# ...remaining for further processing  : 794
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 37
# Backward-rewritten                   : 131
# Generated clauses                    : 48551
# ...of the previous two non-trivial   : 42595
# Contextual simplify-reflections      : 14
# Paramodulations                      : 48547
# Factorizations                       : 2
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 599
#    Positive orientable unit clauses  : 116
#    Positive unorientable unit clauses: 16
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 464
# Current number of unprocessed clauses: 37766
# ...number of literals in the above   : 89077
# Current number of archived formulas  : 0
# Current number of archived clauses   : 198
# Clause-clause subsumption calls (NU) : 129108
# Rec. Clause-clause subsumption calls : 94678
# Non-unit clause-clause subsumptions  : 3087
# Unit Clause-clause subsumption calls : 584
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1043
# BW rewrite match successes           : 275
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1198503

# -------------------------------------------------
# User time                : 0.503 s
# System time              : 0.010 s
# Total time               : 0.513 s
# Maximum resident set size: 38152 pages
