# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t31_relat_1, conjecture, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2))))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t31_relat_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t1_xboole_1)).
fof(t44_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(k4_xboole_0(X1,X2),X3)=>r1_tarski(X1,k2_xboole_0(X2,X3))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t44_xboole_1)).
fof(l51_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', l51_zfmisc_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t36_xboole_1)).
fof(t43_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,k2_xboole_0(X2,X3))=>r1_tarski(k4_xboole_0(X1,X2),X3)), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t43_xboole_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', commutativity_k2_tarski)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t3_xboole_0)).
fof(t7_tarski, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&![X3]:~((r2_hidden(X3,X2)&![X4]:~((r2_hidden(X4,X2)&r2_hidden(X4,X3))))))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t7_tarski)).
fof(d6_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', d6_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', d10_xboole_0)).
fof(t136_zfmisc_1, axiom, ![X1]:?[X2]:(r2_hidden(X1,X2)&(![X3, X4]:((r2_hidden(X3,X2)&r1_tarski(X4,X3))=>r2_hidden(X4,X2))&(![X3]:(r2_hidden(X3,X2)=>r2_hidden(k1_zfmisc_1(X3),X2))&![X3]:~((r1_tarski(X3,X2)&(~(r2_tarski(X3,X2))&~(r2_hidden(X3,X2)))))))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t136_zfmisc_1)).
fof(t8_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X3,X2))=>r1_tarski(k2_xboole_0(X1,X3),X2)), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t8_xboole_1)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t7_xboole_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t2_xboole_1)).
fof(t65_xboole_1, axiom, ![X1]:r1_xboole_0(X1,k1_xboole_0), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t65_xboole_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', d5_relat_1)).
fof(t10_xboole_1, axiom, ![X1, X2, X3]:(r1_tarski(X1,X2)=>r1_tarski(X1,k2_xboole_0(X3,X2))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t10_xboole_1)).
fof(t13_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>k1_relat_1(k2_xboole_0(X1,X2))=k2_xboole_0(k1_relat_1(X1),k1_relat_1(X2)))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t13_relat_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t39_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t1_boole)).
fof(t28_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2))))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', t28_relat_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', redefinition_k6_subset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/f6afd4f2cff6400ca49f33c7bb34c5d1.E_input', d3_tarski)).
fof(c_0_24, negated_conjecture, ~(![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>(r1_tarski(X1,X2)=>r1_tarski(k3_relat_1(X1),k3_relat_1(X2)))))), inference(assume_negation,[status(cth)],[t31_relat_1])).
fof(c_0_25, plain, ![X7, X8, X9]:(~r1_tarski(X7,X8)|~r1_tarski(X8,X9)|r1_tarski(X7,X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_26, negated_conjecture, (v1_relat_1(esk1_0)&(v1_relat_1(esk2_0)&(r1_tarski(esk1_0,esk2_0)&~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
fof(c_0_27, plain, ![X10, X11, X12]:(~r1_tarski(k4_xboole_0(X10,X11),X12)|r1_tarski(X10,k2_xboole_0(X11,X12))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t44_xboole_1])])).
fof(c_0_28, plain, ![X13, X14]:k3_tarski(k2_tarski(X13,X14))=k2_xboole_0(X13,X14), inference(variable_rename,[status(thm)],[l51_zfmisc_1])).
cnf(c_0_29, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (r1_tarski(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_31, plain, ![X15, X16]:r1_tarski(k4_xboole_0(X15,X16),X15), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_32, plain, ![X17, X18, X19]:(~r1_tarski(X17,k2_xboole_0(X18,X19))|r1_tarski(k4_xboole_0(X17,X18),X19)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_xboole_1])])).
cnf(c_0_33, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_34, plain, (k3_tarski(k2_tarski(X1,X2))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (r1_tarski(X1,esk2_0)|~r1_tarski(X1,esk1_0)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_38, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X2,X3)))|~r1_tarski(k4_xboole_0(X1,X2),X3)), inference(rw,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,X1),esk2_0)), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_40, plain, ![X31, X32]:k2_tarski(X31,X32)=k2_tarski(X32,X31), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
fof(c_0_41, plain, ![X20, X21, X23, X24, X25]:(((r2_hidden(esk3_2(X20,X21),X20)|r1_xboole_0(X20,X21))&(r2_hidden(esk3_2(X20,X21),X21)|r1_xboole_0(X20,X21)))&(~r2_hidden(X25,X23)|~r2_hidden(X25,X24)|~r1_xboole_0(X23,X24))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_42, plain, ![X26, X27, X29]:((r2_hidden(esk4_2(X26,X27),X27)|~r2_hidden(X26,X27))&(~r2_hidden(X29,X27)|~r2_hidden(X29,esk4_2(X26,X27))|~r2_hidden(X26,X27))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_tarski])])])])])).
fof(c_0_43, plain, ![X30]:(~v1_relat_1(X30)|k3_relat_1(X30)=k2_xboole_0(k1_relat_1(X30),k2_relat_1(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d6_relat_1])])).
cnf(c_0_44, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k3_tarski(k2_tarski(X2,X3)))), inference(rw,[status(thm)],[c_0_37, c_0_34])).
cnf(c_0_45, negated_conjecture, (r1_tarski(esk1_0,k3_tarski(k2_tarski(X1,esk2_0)))), inference(spm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_46, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_47, plain, ![X44, X45]:(((r1_tarski(X44,X45)|X44!=X45)&(r1_tarski(X45,X44)|X44!=X45))&(~r1_tarski(X44,X45)|~r1_tarski(X45,X44)|X44=X45)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_48, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_49, plain, (r2_hidden(esk4_2(X1,X2),X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_50, plain, ![X33, X35, X36, X37, X38]:(r2_hidden(X33,esk5_1(X33))&((~r2_hidden(X35,esk5_1(X33))|~r1_tarski(X36,X35)|r2_hidden(X36,esk5_1(X33)))&((~r2_hidden(X37,esk5_1(X33))|r2_hidden(k1_zfmisc_1(X37),esk5_1(X33)))&(~r1_tarski(X38,esk5_1(X33))|(r2_tarski(X38,esk5_1(X33))|r2_hidden(X38,esk5_1(X33))))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t136_zfmisc_1])])])])])).
fof(c_0_51, plain, ![X39, X40, X41]:(~r1_tarski(X39,X40)|~r1_tarski(X41,X40)|r1_tarski(k2_xboole_0(X39,X41),X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_xboole_1])])).
cnf(c_0_52, plain, (k3_relat_1(X1)=k2_xboole_0(k1_relat_1(X1),k2_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_53, plain, ![X42, X43]:r1_tarski(X42,k2_xboole_0(X42,X43)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_54, plain, ![X46]:r1_tarski(k1_xboole_0,X46), inference(variable_rename,[status(thm)],[t2_xboole_1])).
cnf(c_0_55, plain, (r1_tarski(X1,X2)|~r1_tarski(X3,k3_tarski(k2_tarski(X4,X2)))|~r1_tarski(X1,k4_xboole_0(X3,X4))), inference(spm,[status(thm)],[c_0_29, c_0_44])).
cnf(c_0_56, negated_conjecture, (r1_tarski(esk1_0,k3_tarski(k2_tarski(esk2_0,X1)))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_57, plain, (r1_tarski(X1,X2)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (~r2_hidden(esk4_2(X1,X2),X3)|~r2_hidden(X1,X2)|~r1_xboole_0(X3,X2)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_59, plain, (r2_hidden(X1,esk5_1(X1))), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_60, plain, ![X47]:r1_xboole_0(X47,k1_xboole_0), inference(variable_rename,[status(thm)],[t65_xboole_1])).
fof(c_0_61, plain, ![X48, X49, X50, X52, X53, X54, X55, X57]:(((~r2_hidden(X50,X49)|r2_hidden(k4_tarski(esk6_3(X48,X49,X50),X50),X48)|X49!=k2_relat_1(X48))&(~r2_hidden(k4_tarski(X53,X52),X48)|r2_hidden(X52,X49)|X49!=k2_relat_1(X48)))&((~r2_hidden(esk7_2(X54,X55),X55)|~r2_hidden(k4_tarski(X57,esk7_2(X54,X55)),X54)|X55=k2_relat_1(X54))&(r2_hidden(esk7_2(X54,X55),X55)|r2_hidden(k4_tarski(esk8_2(X54,X55),esk7_2(X54,X55)),X54)|X55=k2_relat_1(X54)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
cnf(c_0_62, plain, (r1_tarski(k2_xboole_0(X1,X3),X2)|~r1_tarski(X1,X2)|~r1_tarski(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_63, plain, (k3_relat_1(X1)=k3_tarski(k2_tarski(k1_relat_1(X1),k2_relat_1(X1)))|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_52, c_0_34])).
fof(c_0_64, plain, ![X59, X60, X61]:(~r1_tarski(X59,X60)|r1_tarski(X59,k2_xboole_0(X61,X60))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_xboole_1])])).
cnf(c_0_65, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_66, plain, ![X62, X63]:(~v1_relat_1(X62)|(~v1_relat_1(X63)|k1_relat_1(k2_xboole_0(X62,X63))=k2_xboole_0(k1_relat_1(X62),k1_relat_1(X63)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_relat_1])])])).
fof(c_0_67, plain, ![X64, X65]:k2_xboole_0(X64,k4_xboole_0(X65,X64))=k2_xboole_0(X64,X65), inference(variable_rename,[status(thm)],[t39_xboole_1])).
cnf(c_0_68, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_69, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_70, negated_conjecture, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
cnf(c_0_71, plain, (r1_tarski(X1,X1)), inference(er,[status(thm)],[c_0_57])).
fof(c_0_72, plain, ![X66]:k2_xboole_0(X66,k1_xboole_0)=X66, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_73, plain, (~r2_hidden(X1,X2)|~r1_xboole_0(esk5_1(esk4_2(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_74, plain, (r1_xboole_0(X1,k1_xboole_0)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_75, plain, (r2_hidden(k4_tarski(esk6_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_76, plain, (r1_tarski(k3_tarski(k2_tarski(X1,X3)),X2)|~r1_tarski(X3,X2)|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_62, c_0_34])).
cnf(c_0_77, plain, (k3_tarski(k2_tarski(k2_relat_1(X1),k1_relat_1(X1)))=k3_relat_1(X1)|~v1_relat_1(X1)), inference(rw,[status(thm)],[c_0_63, c_0_46])).
cnf(c_0_78, plain, (r1_tarski(X1,k2_xboole_0(X3,X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_79, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_65, c_0_34])).
cnf(c_0_80, plain, (k1_relat_1(k2_xboole_0(X1,X2))=k2_xboole_0(k1_relat_1(X1),k1_relat_1(X2))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_81, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_82, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k4_xboole_0(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_70, c_0_71])).
cnf(c_0_84, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_85, plain, ![X67, X68]:(~v1_relat_1(X67)|(~v1_relat_1(X68)|r1_tarski(k6_subset_1(k2_relat_1(X67),k2_relat_1(X68)),k2_relat_1(k6_subset_1(X67,X68))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_relat_1])])])).
fof(c_0_86, plain, ![X69, X70]:k6_subset_1(X69,X70)=k4_xboole_0(X69,X70), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
cnf(c_0_87, plain, (~r2_hidden(X1,k1_xboole_0)), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_88, plain, (r2_hidden(k4_tarski(esk6_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_75])).
fof(c_0_89, plain, ![X71, X72, X73, X74, X75]:((~r1_tarski(X71,X72)|(~r2_hidden(X73,X71)|r2_hidden(X73,X72)))&((r2_hidden(esk9_2(X74,X75),X74)|r1_tarski(X74,X75))&(~r2_hidden(esk9_2(X74,X75),X75)|r1_tarski(X74,X75)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_90, negated_conjecture, (~r1_tarski(k3_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_91, plain, (r1_tarski(k3_relat_1(X1),X2)|~r1_tarski(k1_relat_1(X1),X2)|~r1_tarski(k2_relat_1(X1),X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_76, c_0_77])).
cnf(c_0_92, negated_conjecture, (v1_relat_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_93, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X3,X2)))|~r1_tarski(X1,X2)), inference(rw,[status(thm)],[c_0_78, c_0_34])).
cnf(c_0_94, plain, (r1_tarski(X1,k3_tarski(k2_tarski(X2,X1)))), inference(spm,[status(thm)],[c_0_79, c_0_46])).
cnf(c_0_95, plain, (k1_relat_1(k3_tarski(k2_tarski(X1,X2)))=k3_tarski(k2_tarski(k1_relat_1(X1),k1_relat_1(X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_34]), c_0_34])).
cnf(c_0_96, plain, (k3_tarski(k2_tarski(X1,k4_xboole_0(X2,X1)))=k3_tarski(k2_tarski(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_34]), c_0_34])).
cnf(c_0_97, negated_conjecture, (k4_xboole_0(esk1_0,esk2_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_98, plain, (k3_tarski(k2_tarski(X1,k1_xboole_0))=X1), inference(rw,[status(thm)],[c_0_84, c_0_34])).
cnf(c_0_99, plain, (r1_tarski(k6_subset_1(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k6_subset_1(X1,X2)))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_100, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_101, plain, (~r2_hidden(X1,k2_relat_1(k1_xboole_0))), inference(spm,[status(thm)],[c_0_87, c_0_88])).
cnf(c_0_102, plain, (r2_hidden(esk9_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_103, negated_conjecture, (~r1_tarski(k1_relat_1(esk1_0),k3_relat_1(esk2_0))|~r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92])])).
cnf(c_0_104, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(X1,k1_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_93, c_0_77])).
cnf(c_0_105, negated_conjecture, (v1_relat_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_106, plain, (r1_tarski(k1_relat_1(X1),k1_relat_1(k3_tarski(k2_tarski(X2,X1))))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_107, negated_conjecture, (k3_tarski(k2_tarski(esk2_0,esk1_0))=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97]), c_0_98])).
cnf(c_0_108, plain, (r1_tarski(k4_xboole_0(k2_relat_1(X1),k2_relat_1(X2)),k2_relat_1(k4_xboole_0(X1,X2)))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_99, c_0_100]), c_0_100])).
cnf(c_0_109, plain, (r1_tarski(k2_relat_1(k1_xboole_0),X1)), inference(spm,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_110, negated_conjecture, (~r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))|~r1_tarski(k1_relat_1(esk1_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_105])])).
cnf(c_0_111, negated_conjecture, (r1_tarski(k1_relat_1(esk1_0),k1_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_92]), c_0_105])])).
cnf(c_0_112, plain, (r1_tarski(k2_relat_1(X1),k3_relat_1(X1))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_79, c_0_77])).
cnf(c_0_113, plain, (r1_tarski(k2_relat_1(X1),k3_tarski(k2_tarski(k2_relat_1(X2),k2_relat_1(k4_xboole_0(X1,X2)))))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_38, c_0_108])).
cnf(c_0_114, plain, (k2_relat_1(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_82, c_0_109])).
cnf(c_0_115, negated_conjecture, (~r1_tarski(k2_relat_1(esk1_0),k3_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_116, plain, (r1_tarski(X1,k3_relat_1(X2))|~r1_tarski(X1,k2_relat_1(X2))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_29, c_0_112])).
cnf(c_0_117, negated_conjecture, (r1_tarski(k2_relat_1(esk1_0),k2_relat_1(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_97]), c_0_114]), c_0_98]), c_0_105]), c_0_92])])).
cnf(c_0_118, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_116]), c_0_117]), c_0_105])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 119
# Proof object clause steps            : 70
# Proof object formula steps           : 49
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 28
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 24
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 38
# Processed clauses                    : 1156
# ...of these trivial                  : 5
# ...subsumed                          : 722
# ...remaining for further processing  : 429
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 19
# Generated clauses                    : 3733
# ...of the previous two non-trivial   : 3293
# Contextual simplify-reflections      : 9
# Paramodulations                      : 3729
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 361
#    Positive orientable unit clauses  : 32
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 315
# Current number of unprocessed clauses: 2172
# ...number of literals in the above   : 7095
# Current number of archived formulas  : 0
# Current number of archived clauses   : 66
# Clause-clause subsumption calls (NU) : 17650
# Rec. Clause-clause subsumption calls : 11659
# Non-unit clause-clause subsumptions  : 609
# Unit Clause-clause subsumption calls : 983
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 86
# BW rewrite match successes           : 19
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 43744

# -------------------------------------------------
# User time                : 0.044 s
# System time              : 0.000 s
# Total time               : 0.044 s
# Maximum resident set size: 28856 pages
