# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.009 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t52_subset_1, conjecture, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(~(r2_hidden(X4,X2))<=>r2_hidden(X4,X3)))=>X2=k3_subset_1(X1,X3)))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', t52_subset_1)).
fof(d1_zfmisc_1, axiom, ![X1, X2]:(X2=k1_zfmisc_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>r1_tarski(X3,X1))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', d1_zfmisc_1)).
fof(d2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))=>(m1_subset_1(X2,X1)<=>r2_hidden(X2,X1)))&(v1_xboole_0(X1)=>(m1_subset_1(X2,X1)<=>v1_xboole_0(X2)))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', d2_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', fc1_subset_1)).
fof(dt_k6_subset_1, axiom, ![X1, X2]:m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1)), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', dt_k6_subset_1)).
fof(redefinition_k6_subset_1, axiom, ![X1, X2]:k6_subset_1(X1,X2)=k4_xboole_0(X1,X2), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', redefinition_k6_subset_1)).
fof(t100_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2)), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', t100_xboole_1)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', t28_xboole_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', commutativity_k3_xboole_0)).
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', d5_xboole_0)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', d1_xboole_0)).
fof(d4_xboole_0, axiom, ![X1, X2, X3]:(X3=k3_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&r2_hidden(X4,X2)))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', d4_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', d3_tarski)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', t48_xboole_1)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', t47_xboole_1)).
fof(d5_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k3_subset_1(X1,X2)=k4_xboole_0(X1,X2)), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', d5_subset_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/f8256cbea82f42aebc4dfe77502b9ab4.E_input', d10_xboole_0)).
fof(c_0_17, negated_conjecture, ~(![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(X1))=>(![X4]:(m1_subset_1(X4,X1)=>(~(r2_hidden(X4,X2))<=>r2_hidden(X4,X3)))=>X2=k3_subset_1(X1,X3))))), inference(assume_negation,[status(cth)],[t52_subset_1])).
fof(c_0_18, plain, ![X22, X23, X24, X25, X26, X27]:(((~r2_hidden(X24,X23)|r1_tarski(X24,X22)|X23!=k1_zfmisc_1(X22))&(~r1_tarski(X25,X22)|r2_hidden(X25,X23)|X23!=k1_zfmisc_1(X22)))&((~r2_hidden(esk5_2(X26,X27),X27)|~r1_tarski(esk5_2(X26,X27),X26)|X27=k1_zfmisc_1(X26))&(r2_hidden(esk5_2(X26,X27),X27)|r1_tarski(esk5_2(X26,X27),X26)|X27=k1_zfmisc_1(X26)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_zfmisc_1])])])])])])).
fof(c_0_19, plain, ![X44, X45]:(((~m1_subset_1(X45,X44)|r2_hidden(X45,X44)|v1_xboole_0(X44))&(~r2_hidden(X45,X44)|m1_subset_1(X45,X44)|v1_xboole_0(X44)))&((~m1_subset_1(X45,X44)|v1_xboole_0(X45)|~v1_xboole_0(X44))&(~v1_xboole_0(X45)|m1_subset_1(X45,X44)|~v1_xboole_0(X44)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_subset_1])])])])).
fof(c_0_20, negated_conjecture, ![X8]:(m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))&(((r2_hidden(X8,esk2_0)|r2_hidden(X8,esk3_0)|~m1_subset_1(X8,esk1_0))&(~r2_hidden(X8,esk3_0)|~r2_hidden(X8,esk2_0)|~m1_subset_1(X8,esk1_0)))&esk2_0!=k3_subset_1(esk1_0,esk3_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_17])])])])])])).
fof(c_0_21, plain, ![X43]:~v1_xboole_0(k1_zfmisc_1(X43)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_22, plain, (r1_tarski(X1,X3)|~r2_hidden(X1,X2)|X2!=k1_zfmisc_1(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_23, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_25, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_26, plain, ![X59, X60]:m1_subset_1(k6_subset_1(X59,X60),k1_zfmisc_1(X59)), inference(variable_rename,[status(thm)],[dt_k6_subset_1])).
fof(c_0_27, plain, ![X63, X64]:k6_subset_1(X63,X64)=k4_xboole_0(X63,X64), inference(variable_rename,[status(thm)],[redefinition_k6_subset_1])).
fof(c_0_28, plain, ![X31, X32]:k4_xboole_0(X31,X32)=k5_xboole_0(X31,k3_xboole_0(X31,X32)), inference(variable_rename,[status(thm)],[t100_xboole_1])).
fof(c_0_29, plain, ![X48, X49]:(~r1_tarski(X48,X49)|k3_xboole_0(X48,X49)=X48), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
cnf(c_0_30, plain, (r1_tarski(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X2))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_31, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
fof(c_0_32, plain, ![X18, X19]:k3_xboole_0(X18,X19)=k3_xboole_0(X19,X18), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_33, plain, ![X9, X10, X11, X12, X13, X14, X15, X16]:((((r2_hidden(X12,X9)|~r2_hidden(X12,X11)|X11!=k4_xboole_0(X9,X10))&(~r2_hidden(X12,X10)|~r2_hidden(X12,X11)|X11!=k4_xboole_0(X9,X10)))&(~r2_hidden(X13,X9)|r2_hidden(X13,X10)|r2_hidden(X13,X11)|X11!=k4_xboole_0(X9,X10)))&((~r2_hidden(esk4_3(X14,X15,X16),X16)|(~r2_hidden(esk4_3(X14,X15,X16),X14)|r2_hidden(esk4_3(X14,X15,X16),X15))|X16=k4_xboole_0(X14,X15))&((r2_hidden(esk4_3(X14,X15,X16),X14)|r2_hidden(esk4_3(X14,X15,X16),X16)|X16=k4_xboole_0(X14,X15))&(~r2_hidden(esk4_3(X14,X15,X16),X15)|r2_hidden(esk4_3(X14,X15,X16),X16)|X16=k4_xboole_0(X14,X15))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_34, plain, (m1_subset_1(k6_subset_1(X1,X2),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_35, plain, (k6_subset_1(X1,X2)=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (k4_xboole_0(X1,X2)=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_38, negated_conjecture, (r1_tarski(esk2_0,esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_39, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_40, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_41, plain, ![X33, X34, X35]:((~v1_xboole_0(X33)|~r2_hidden(X34,X33))&(r2_hidden(esk6_1(X35),X35)|v1_xboole_0(X35))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
cnf(c_0_42, plain, (m1_subset_1(k5_xboole_0(X1,k3_xboole_0(X1,X2)),k1_zfmisc_1(X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35]), c_0_36])).
cnf(c_0_43, negated_conjecture, (k3_xboole_0(esk1_0,esk2_0)=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_45, plain, (X3!=k5_xboole_0(X4,k3_xboole_0(X4,X2))|~r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_40, c_0_36])).
cnf(c_0_46, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_47, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(k5_xboole_0(esk1_0,esk2_0),k1_zfmisc_1(esk1_0))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
fof(c_0_49, plain, ![X50, X51, X52, X53, X54, X55, X56, X57]:((((r2_hidden(X53,X50)|~r2_hidden(X53,X52)|X52!=k3_xboole_0(X50,X51))&(r2_hidden(X53,X51)|~r2_hidden(X53,X52)|X52!=k3_xboole_0(X50,X51)))&(~r2_hidden(X54,X50)|~r2_hidden(X54,X51)|r2_hidden(X54,X52)|X52!=k3_xboole_0(X50,X51)))&((~r2_hidden(esk8_3(X55,X56,X57),X57)|(~r2_hidden(esk8_3(X55,X56,X57),X55)|~r2_hidden(esk8_3(X55,X56,X57),X56))|X57=k3_xboole_0(X55,X56))&((r2_hidden(esk8_3(X55,X56,X57),X55)|r2_hidden(esk8_3(X55,X56,X57),X57)|X57=k3_xboole_0(X55,X56))&(r2_hidden(esk8_3(X55,X56,X57),X56)|r2_hidden(esk8_3(X55,X56,X57),X57)|X57=k3_xboole_0(X55,X56))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_xboole_0])])])])])])).
cnf(c_0_50, negated_conjecture, (r2_hidden(esk3_0,k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_44]), c_0_25])).
cnf(c_0_51, plain, (~r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_45])).
fof(c_0_52, plain, ![X37, X38, X39, X40, X41]:((~r1_tarski(X37,X38)|(~r2_hidden(X39,X37)|r2_hidden(X39,X38)))&((r2_hidden(esk7_2(X40,X41),X40)|r1_tarski(X40,X41))&(~r2_hidden(esk7_2(X40,X41),X41)|r1_tarski(X40,X41)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_53, negated_conjecture, (r2_hidden(X1,esk2_0)|r2_hidden(X1,esk3_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_54, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(csr,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_55, negated_conjecture, (r2_hidden(k5_xboole_0(esk1_0,esk2_0),k1_zfmisc_1(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_48]), c_0_25])).
cnf(c_0_56, plain, (r2_hidden(X1,X3)|r2_hidden(X1,X4)|~r2_hidden(X1,X2)|X4!=k4_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_57, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k3_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_58, negated_conjecture, (r1_tarski(esk3_0,esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_50])).
fof(c_0_59, plain, ![X61, X62]:k4_xboole_0(X61,k4_xboole_0(X61,X62))=k3_xboole_0(X61,X62), inference(variable_rename,[status(thm)],[t48_xboole_1])).
cnf(c_0_60, negated_conjecture, (~r2_hidden(X1,k5_xboole_0(esk1_0,esk2_0))|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_51, c_0_43])).
cnf(c_0_61, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_63, negated_conjecture, (r2_hidden(X1,esk3_0)|r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_64, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_65, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),esk1_0)), inference(spm,[status(thm)],[c_0_30, c_0_55])).
cnf(c_0_66, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X4!=k5_xboole_0(X2,k3_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_56, c_0_36])).
cnf(c_0_67, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_xboole_0(X2,X3))), inference(er,[status(thm)],[c_0_57])).
cnf(c_0_68, negated_conjecture, (k3_xboole_0(esk1_0,esk3_0)=esk3_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_58]), c_0_39])).
fof(c_0_69, plain, ![X46, X47]:k4_xboole_0(X46,k3_xboole_0(X46,X47))=k4_xboole_0(X46,X47), inference(variable_rename,[status(thm)],[t47_xboole_1])).
fof(c_0_70, plain, ![X29, X30]:(~m1_subset_1(X30,k1_zfmisc_1(X29))|k3_subset_1(X29,X30)=k4_xboole_0(X29,X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_subset_1])])).
cnf(c_0_71, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_72, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),X1)|~r2_hidden(esk7_2(k5_xboole_0(esk1_0,esk2_0),X1),esk2_0)), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_73, negated_conjecture, (r1_tarski(X1,esk3_0)|r2_hidden(esk7_2(X1,esk3_0),esk2_0)|~r2_hidden(esk7_2(X1,esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_62, c_0_63])).
cnf(c_0_74, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,k5_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_75, negated_conjecture, (~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_76, plain, (r2_hidden(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|r2_hidden(X1,X3)|~r2_hidden(X1,X2)), inference(er,[status(thm)],[c_0_66])).
cnf(c_0_77, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
cnf(c_0_78, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_79, plain, (k3_subset_1(X2,X1)=k4_xboole_0(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_80, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k5_xboole_0(X1,k3_xboole_0(X1,X2))))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_36]), c_0_36])).
fof(c_0_81, plain, ![X20, X21]:(((r1_tarski(X20,X21)|X20!=X21)&(r1_tarski(X21,X20)|X20!=X21))&(~r1_tarski(X20,X21)|~r1_tarski(X21,X20)|X20=X21)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_82, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),esk3_0)|~r2_hidden(esk7_2(k5_xboole_0(esk1_0,esk2_0),esk3_0),esk1_0)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_83, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),X1)|r2_hidden(esk7_2(k5_xboole_0(esk1_0,esk2_0),X1),esk1_0)), inference(spm,[status(thm)],[c_0_74, c_0_61])).
cnf(c_0_84, negated_conjecture, (~r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk3_0)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_75, c_0_54])).
cnf(c_0_85, negated_conjecture, (r2_hidden(X1,esk1_0)|~r2_hidden(X1,esk2_0)), inference(spm,[status(thm)],[c_0_64, c_0_38])).
cnf(c_0_86, plain, (r1_tarski(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3)))|r2_hidden(esk7_2(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3))),X3)|~r2_hidden(esk7_2(X1,k5_xboole_0(X2,k3_xboole_0(X2,X3))),X2)), inference(spm,[status(thm)],[c_0_62, c_0_76])).
cnf(c_0_87, negated_conjecture, (r1_tarski(esk3_0,X1)|r2_hidden(esk7_2(esk3_0,X1),esk1_0)), inference(spm,[status(thm)],[c_0_77, c_0_61])).
cnf(c_0_88, plain, (k5_xboole_0(X1,k3_xboole_0(X1,k3_xboole_0(X1,X2)))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_36]), c_0_36])).
cnf(c_0_89, plain, (k5_xboole_0(X2,k3_xboole_0(X2,X1))=k3_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_79, c_0_36])).
cnf(c_0_90, plain, (m1_subset_1(k3_xboole_0(X1,X2),k1_zfmisc_1(X1))), inference(spm,[status(thm)],[c_0_42, c_0_80])).
cnf(c_0_91, negated_conjecture, (k3_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk2_0))=k5_xboole_0(esk1_0,esk2_0)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_65]), c_0_39])).
cnf(c_0_92, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_93, negated_conjecture, (r1_tarski(k5_xboole_0(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_94, negated_conjecture, (r1_tarski(esk3_0,X1)|~r2_hidden(esk7_2(esk3_0,X1),esk2_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_61]), c_0_85])).
cnf(c_0_95, negated_conjecture, (r1_tarski(esk3_0,k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,X1)))|r2_hidden(esk7_2(esk3_0,k5_xboole_0(esk1_0,k3_xboole_0(esk1_0,X1))),X1)), inference(spm,[status(thm)],[c_0_86, c_0_87])).
cnf(c_0_96, plain, (k3_subset_1(X1,k3_xboole_0(X1,X2))=k5_xboole_0(X1,k3_xboole_0(X1,X2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_90])])).
cnf(c_0_97, negated_conjecture, (k5_xboole_0(esk1_0,k5_xboole_0(esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_43]), c_0_91])).
cnf(c_0_98, negated_conjecture, (k5_xboole_0(esk1_0,esk2_0)=esk3_0|~r1_tarski(esk3_0,k5_xboole_0(esk1_0,esk2_0))), inference(spm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_99, negated_conjecture, (r1_tarski(esk3_0,k5_xboole_0(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_43])).
cnf(c_0_100, negated_conjecture, (esk2_0!=k3_subset_1(esk1_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_101, negated_conjecture, (k3_subset_1(esk1_0,k5_xboole_0(esk1_0,esk2_0))=esk2_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_91]), c_0_97])).
cnf(c_0_102, negated_conjecture, (k5_xboole_0(esk1_0,esk2_0)=esk3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_99])])).
cnf(c_0_103, negated_conjecture, (k3_subset_1(esk1_0,esk3_0)=k5_xboole_0(esk1_0,esk3_0)), inference(spm,[status(thm)],[c_0_96, c_0_68])).
cnf(c_0_104, negated_conjecture, (k5_xboole_0(esk1_0,esk3_0)!=esk2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_89]), c_0_68]), c_0_44])])).
cnf(c_0_105, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102]), c_0_103]), c_0_104]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 106
# Proof object clause steps            : 71
# Proof object formula steps           : 35
# Proof object conjectures             : 40
# Proof object clause conjectures      : 37
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 25
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 33
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 42
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 40
# Processed clauses                    : 18493
# ...of these trivial                  : 268
# ...subsumed                          : 15867
# ...remaining for further processing  : 2358
# Other redundant clauses eliminated   : 10
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 148
# Backward-rewritten                   : 376
# Generated clauses                    : 253059
# ...of the previous two non-trivial   : 235259
# Contextual simplify-reflections      : 58
# Paramodulations                      : 252887
# Factorizations                       : 160
# Equation resolutions                 : 10
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1783
#    Positive orientable unit clauses  : 194
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 23
#    Non-unit-clauses                  : 1565
# Current number of unprocessed clauses: 215791
# ...number of literals in the above   : 855317
# Current number of archived formulas  : 0
# Current number of archived clauses   : 567
# Clause-clause subsumption calls (NU) : 450837
# Rec. Clause-clause subsumption calls : 261435
# Non-unit clause-clause subsumptions  : 10198
# Unit Clause-clause subsumption calls : 28876
# Rewrite failures with RHS unbound    : 56
# BW rewrite match attempts            : 381
# BW rewrite match successes           : 69
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5175535

# -------------------------------------------------
# User time                : 3.192 s
# System time              : 0.070 s
# Total time               : 3.261 s
# Maximum resident set size: 178380 pages
