# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t43_zfmisc_1, axiom, ![X1, X2, X3]:~((k1_tarski(X1)=k2_xboole_0(X2,X3)&(~((X2=k1_tarski(X1)&X3=k1_tarski(X1)))&(~((X2=k1_xboole_0&X3=k1_tarski(X1)))&~((X2=k1_tarski(X1)&X3=k1_xboole_0)))))), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t43_zfmisc_1)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t22_xboole_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', d1_tarski)).
fof(t18_zfmisc_1, axiom, ![X1, X2]:(k3_xboole_0(k1_tarski(X1),k1_tarski(X2))=k1_tarski(X1)=>X1=X2), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t18_zfmisc_1)).
fof(t167_relat_1, axiom, ![X1, X2]:(v1_relat_1(X2)=>r1_tarski(k10_relat_1(X2,X1),k1_relat_1(X2))), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t167_relat_1)).
fof(t10_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>~((~(X2=k1_xboole_0)&![X3]:(m1_subset_1(X3,X1)=>~(r2_hidden(X3,X2)))))), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t10_subset_1)).
fof(t172_relat_1, conjecture, ![X1]:k10_relat_1(k1_xboole_0,X1)=k1_xboole_0, file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t172_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', d10_xboole_0)).
fof(t60_relat_1, axiom, (k1_relat_1(k1_xboole_0)=k1_xboole_0&k2_relat_1(k1_xboole_0)=k1_xboole_0), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t60_relat_1)).
fof(t2_xboole_1, axiom, ![X1]:r1_tarski(k1_xboole_0,X1), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t2_xboole_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', cc2_relat_1)).
fof(t61_setfam_1, axiom, ![X1]:m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(X1))), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', t61_setfam_1)).
fof(fc7_relat_1, axiom, ![X1, X2, X3, X4]:v1_relat_1(k2_tarski(k4_tarski(X1,X2),k4_tarski(X3,X4))), file('data/MPTP/proofs/f8afbb52815f401a86cbeb42f75efda5.E_input', fc7_relat_1)).
fof(c_0_13, plain, ![X6, X7, X8]:((((X7=k1_tarski(X6)|X7=k1_xboole_0|X7=k1_tarski(X6)|k1_tarski(X6)!=k2_xboole_0(X7,X8))&(X8=k1_xboole_0|X7=k1_xboole_0|X7=k1_tarski(X6)|k1_tarski(X6)!=k2_xboole_0(X7,X8)))&((X7=k1_tarski(X6)|X8=k1_tarski(X6)|X7=k1_tarski(X6)|k1_tarski(X6)!=k2_xboole_0(X7,X8))&(X8=k1_xboole_0|X8=k1_tarski(X6)|X7=k1_tarski(X6)|k1_tarski(X6)!=k2_xboole_0(X7,X8))))&(((X7=k1_tarski(X6)|X7=k1_xboole_0|X8=k1_tarski(X6)|k1_tarski(X6)!=k2_xboole_0(X7,X8))&(X8=k1_xboole_0|X7=k1_xboole_0|X8=k1_tarski(X6)|k1_tarski(X6)!=k2_xboole_0(X7,X8)))&((X7=k1_tarski(X6)|X8=k1_tarski(X6)|X8=k1_tarski(X6)|k1_tarski(X6)!=k2_xboole_0(X7,X8))&(X8=k1_xboole_0|X8=k1_tarski(X6)|X8=k1_tarski(X6)|k1_tarski(X6)!=k2_xboole_0(X7,X8))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t43_zfmisc_1])])])).
cnf(c_0_14, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|X1=k1_tarski(X2)|k1_tarski(X2)!=k2_xboole_0(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_15, plain, ![X33, X34]:k2_xboole_0(X33,k3_xboole_0(X33,X34))=X33, inference(variable_rename,[status(thm)],[t22_xboole_1])).
cnf(c_0_16, plain, (X1=k1_xboole_0|X1=k1_tarski(X2)|k1_tarski(X2)!=k2_xboole_0(X3,X1)), inference(cn,[status(thm)],[c_0_14])).
cnf(c_0_17, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_18, plain, ![X20, X21, X22, X23, X24, X25]:(((~r2_hidden(X22,X21)|X22=X20|X21!=k1_tarski(X20))&(X23!=X20|r2_hidden(X23,X21)|X21!=k1_tarski(X20)))&((~r2_hidden(esk3_2(X24,X25),X25)|esk3_2(X24,X25)!=X24|X25=k1_tarski(X24))&(r2_hidden(esk3_2(X24,X25),X25)|esk3_2(X24,X25)=X24|X25=k1_tarski(X24)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_19, plain, ![X9, X10]:(k3_xboole_0(k1_tarski(X9),k1_tarski(X10))!=k1_tarski(X9)|X9=X10), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_zfmisc_1])])).
cnf(c_0_20, plain, (k3_xboole_0(k1_tarski(X1),X2)=k1_tarski(X1)|k3_xboole_0(k1_tarski(X1),X2)=k1_xboole_0), inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17])])).
fof(c_0_21, plain, ![X27, X28]:(~v1_relat_1(X28)|r1_tarski(k10_relat_1(X28,X27),k1_relat_1(X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t167_relat_1])])).
cnf(c_0_22, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_18])).
fof(c_0_23, plain, ![X17, X18]:((m1_subset_1(esk2_2(X17,X18),X17)|X18=k1_xboole_0|~m1_subset_1(X18,k1_zfmisc_1(X17)))&(r2_hidden(esk2_2(X17,X18),X18)|X18=k1_xboole_0|~m1_subset_1(X18,k1_zfmisc_1(X17)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_subset_1])])])])])).
cnf(c_0_24, plain, (X1=X2|k3_xboole_0(k1_tarski(X1),k1_tarski(X2))!=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (k3_xboole_0(k1_tarski(X1),X2)=k1_xboole_0|k1_tarski(X1)!=k1_xboole_0), inference(ef,[status(thm)],[c_0_20])).
fof(c_0_26, negated_conjecture, ~(![X1]:k10_relat_1(k1_xboole_0,X1)=k1_xboole_0), inference(assume_negation,[status(cth)],[t172_relat_1])).
fof(c_0_27, plain, ![X14, X15]:(((r1_tarski(X14,X15)|X14!=X15)&(r1_tarski(X15,X14)|X14!=X15))&(~r1_tarski(X14,X15)|~r1_tarski(X15,X14)|X14=X15)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_28, plain, (r1_tarski(k10_relat_1(X1,X2),k1_relat_1(X1))|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (k1_relat_1(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t60_relat_1])).
fof(c_0_30, plain, ![X11]:r1_tarski(k1_xboole_0,X11), inference(variable_rename,[status(thm)],[t2_xboole_1])).
fof(c_0_31, plain, ![X12, X13]:(~v1_relat_1(X12)|(~m1_subset_1(X13,k1_zfmisc_1(X12))|v1_relat_1(X13))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
cnf(c_0_32, plain, (X1=X2|~r2_hidden(X1,k1_tarski(X2))), inference(er,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (r2_hidden(esk2_2(X1,X2),X2)|X2=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (X1=X2|k1_tarski(X1)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_24, c_0_25])).
fof(c_0_35, plain, ![X16]:m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(X16))), inference(variable_rename,[status(thm)],[t61_setfam_1])).
fof(c_0_36, negated_conjecture, k10_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0, inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_26])])])).
cnf(c_0_37, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_38, plain, (r1_tarski(k10_relat_1(k1_xboole_0,X1),k1_xboole_0)|~v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_39, plain, (r1_tarski(k1_xboole_0,X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_40, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_41, plain, (m1_subset_1(esk2_2(X1,X2),X1)|X2=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_42, plain, (esk2_2(X1,k1_tarski(X2))=X2|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_33]), c_0_34])).
cnf(c_0_43, plain, (m1_subset_1(k1_tarski(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_44, negated_conjecture, (k10_relat_1(k1_xboole_0,esk1_0)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_45, plain, (k10_relat_1(k1_xboole_0,X1)=k1_xboole_0|~v1_relat_1(k1_xboole_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38]), c_0_39])])).
cnf(c_0_46, plain, (X1=k1_xboole_0|v1_relat_1(esk2_2(k1_zfmisc_1(X2),X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v1_relat_1(X2)), inference(spm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_47, plain, (esk2_2(k1_zfmisc_1(X1),k1_tarski(k1_xboole_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_48, negated_conjecture, (~v1_relat_1(k1_xboole_0)), inference(spm,[status(thm)],[c_0_44, c_0_45])).
fof(c_0_49, plain, ![X29, X30, X31, X32]:v1_relat_1(k2_tarski(k4_tarski(X29,X30),k4_tarski(X31,X32))), inference(variable_rename,[status(thm)],[fc7_relat_1])).
cnf(c_0_50, plain, (k1_tarski(k1_xboole_0)=k1_xboole_0|~v1_relat_1(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47]), c_0_43])]), c_0_48])).
cnf(c_0_51, plain, (v1_relat_1(k2_tarski(k4_tarski(X1,X2),k4_tarski(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_52, plain, (k1_tarski(k1_xboole_0)=k1_xboole_0), inference(spm,[status(thm)],[c_0_50, c_0_51])).
cnf(c_0_53, plain, (k1_xboole_0=X1), inference(spm,[status(thm)],[c_0_34, c_0_52])).
cnf(c_0_54, plain, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_53]), c_0_48]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 55
# Proof object clause steps            : 29
# Proof object formula steps           : 26
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 14
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 77
# ...of these trivial                  : 1
# ...subsumed                          : 9
# ...remaining for further processing  : 67
# Other redundant clauses eliminated   : 33
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 36
# Generated clauses                    : 213
# ...of the previous two non-trivial   : 176
# Contextual simplify-reflections      : 1
# Paramodulations                      : 179
# Factorizations                       : 2
# Equation resolutions                 : 33
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 5
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 1
# Current number of unprocessed clauses: 137
# ...number of literals in the above   : 362
# Current number of archived formulas  : 0
# Current number of archived clauses   : 58
# Clause-clause subsumption calls (NU) : 41
# Rec. Clause-clause subsumption calls : 36
# Non-unit clause-clause subsumptions  : 7
# Unit Clause-clause subsumption calls : 16
# Rewrite failures with RHS unbound    : 5
# BW rewrite match attempts            : 36
# BW rewrite match successes           : 33
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2979

# -------------------------------------------------
# User time                : 0.006 s
# System time              : 0.004 s
# Total time               : 0.010 s
# Maximum resident set size: 26216 pages
