# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', d1_tarski)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', t69_enumset1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', redefinition_k6_domain_1)).
fof(t35_orders_2, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', t35_orders_2)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', dt_l1_orders_2)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', d2_tarski)).
fof(d7_relat_2, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(r7_relat_2(X1,X2)<=>![X3, X4]:~((r2_hidden(X3,X2)&(r2_hidden(X4,X2)&(~(r2_hidden(k4_tarski(X3,X4),X1))&~(r2_hidden(k4_tarski(X4,X3),X1)))))))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', d7_relat_2)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', cc2_relat_1)).
fof(dt_u1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', dt_u1_orders_2)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', fc6_relat_1)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', dt_k6_domain_1)).
fof(d11_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v6_orders_2(X2,X1)<=>r7_relat_2(u1_orders_2(X1),X2)))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', d11_orders_2)).
fof(d9_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_orders_2(X1,X2,X3)<=>r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1)))))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', d9_orders_2)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/fb3f7305310b4f9f96c49ede1fc8dfe5.E_input', t24_orders_2)).
fof(c_0_15, plain, ![X55, X56, X57, X58, X59, X60]:(((~r2_hidden(X57,X56)|X57=X55|X56!=k1_tarski(X55))&(X58!=X55|r2_hidden(X58,X56)|X56!=k1_tarski(X55)))&((~r2_hidden(esk9_2(X59,X60),X60)|esk9_2(X59,X60)!=X59|X60=k1_tarski(X59))&(r2_hidden(esk9_2(X59,X60),X60)|esk9_2(X59,X60)=X59|X60=k1_tarski(X59)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
fof(c_0_16, plain, ![X10]:k2_tarski(X10,X10)=k1_tarski(X10), inference(variable_rename,[status(thm)],[t69_enumset1])).
cnf(c_0_17, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_19, plain, ![X64, X65]:(v1_xboole_0(X64)|~m1_subset_1(X65,X64)|k6_domain_1(X64,X65)=k1_tarski(X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_20, plain, (X1=X3|X2!=k2_tarski(X3,X3)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_21, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_22, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v6_orders_2(k6_domain_1(u1_struct_0(X1),X2),X1)&m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1))))))), inference(assume_negation,[status(cth)],[t35_orders_2])).
fof(c_0_23, plain, ![X35]:(v2_struct_0(X35)|~l1_struct_0(X35)|~v1_xboole_0(u1_struct_0(X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_24, plain, ![X29]:(~l1_orders_2(X29)|l1_struct_0(X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_25, plain, (X1=X2|~r2_hidden(X1,k2_tarski(X2,X2))), inference(er,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (k2_tarski(X2,X2)=k6_domain_1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_21, c_0_18])).
cnf(c_0_27, plain, (r2_hidden(esk9_2(X1,X2),X2)|esk9_2(X1,X2)=X1|X2=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_28, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&l1_orders_2(esk1_0)))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~v6_orders_2(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_22])])])])).
cnf(c_0_29, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_30, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_31, plain, (X1=X2|v1_xboole_0(X3)|~r2_hidden(X1,k6_domain_1(X3,X2))|~m1_subset_1(X2,X3)), inference(spm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_32, plain, (X2=k2_tarski(X1,X1)|esk9_2(X1,X2)=X1|r2_hidden(esk9_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_27, c_0_18])).
cnf(c_0_33, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_34, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_35, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_36, plain, (esk9_2(X1,k6_domain_1(X2,X3))=X1|k6_domain_1(X2,X3)=k2_tarski(X1,X1)|esk9_2(X1,k6_domain_1(X2,X3))=X3|v1_xboole_0(X2)|~m1_subset_1(X3,X2)), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_37, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_34]), c_0_35])])).
fof(c_0_39, plain, ![X16, X17, X18, X19, X20, X21, X22, X23]:(((~r2_hidden(X19,X18)|(X19=X16|X19=X17)|X18!=k2_tarski(X16,X17))&((X20!=X16|r2_hidden(X20,X18)|X18!=k2_tarski(X16,X17))&(X20!=X17|r2_hidden(X20,X18)|X18!=k2_tarski(X16,X17))))&(((esk4_3(X21,X22,X23)!=X21|~r2_hidden(esk4_3(X21,X22,X23),X23)|X23=k2_tarski(X21,X22))&(esk4_3(X21,X22,X23)!=X22|~r2_hidden(esk4_3(X21,X22,X23),X23)|X23=k2_tarski(X21,X22)))&(r2_hidden(esk4_3(X21,X22,X23),X23)|(esk4_3(X21,X22,X23)=X21|esk4_3(X21,X22,X23)=X22)|X23=k2_tarski(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_40, plain, ![X36, X37, X38, X39, X40]:((~r7_relat_2(X36,X37)|(~r2_hidden(X38,X37)|(~r2_hidden(X39,X37)|(r2_hidden(k4_tarski(X38,X39),X36)|r2_hidden(k4_tarski(X39,X38),X36))))|~v1_relat_1(X36))&((r2_hidden(esk6_2(X36,X40),X40)|r7_relat_2(X36,X40)|~v1_relat_1(X36))&((r2_hidden(esk7_2(X36,X40),X40)|r7_relat_2(X36,X40)|~v1_relat_1(X36))&((~r2_hidden(k4_tarski(esk6_2(X36,X40),esk7_2(X36,X40)),X36)|r7_relat_2(X36,X40)|~v1_relat_1(X36))&(~r2_hidden(k4_tarski(esk7_2(X36,X40),esk6_2(X36,X40)),X36)|r7_relat_2(X36,X40)|~v1_relat_1(X36)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_relat_2])])])])])])])).
cnf(c_0_41, plain, (X2=k1_tarski(X1)|~r2_hidden(esk9_2(X1,X2),X2)|esk9_2(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_42, negated_conjecture, (esk9_2(X1,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(X1,X1)|esk9_2(X1,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=X1), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37]), c_0_38])).
cnf(c_0_43, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44, plain, (r2_hidden(esk6_2(X1,X2),X2)|r7_relat_2(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_45, plain, (X2=k2_tarski(X1,X1)|esk9_2(X1,X2)!=X1|~r2_hidden(esk9_2(X1,X2),X2)), inference(rw,[status(thm)],[c_0_41, c_0_18])).
cnf(c_0_46, negated_conjecture, (esk9_2(esk2_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))=esk2_0|k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(er,[status(thm)],[inference(ef,[status(thm)],[c_0_42])])).
cnf(c_0_47, plain, (r2_hidden(X1,k2_tarski(X1,X2))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_43])])).
cnf(c_0_48, plain, (r7_relat_2(X1,X2)|~r2_hidden(k4_tarski(esk6_2(X1,X2),esk7_2(X1,X2)),X1)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_49, plain, (esk6_2(X1,k2_tarski(X2,X2))=X2|r7_relat_2(X1,k2_tarski(X2,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_44])).
cnf(c_0_50, plain, (r2_hidden(esk7_2(X1,X2),X2)|r7_relat_2(X1,X2)|~v1_relat_1(X1)), inference(split_conjunct,[status(thm)],[c_0_40])).
fof(c_0_51, plain, ![X43, X44]:(~v1_relat_1(X43)|(~m1_subset_1(X44,k1_zfmisc_1(X43))|v1_relat_1(X44))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_52, plain, ![X47]:(~l1_orders_2(X47)|m1_subset_1(u1_orders_2(X47),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X47),u1_struct_0(X47))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_orders_2])])).
fof(c_0_53, plain, ![X27, X28]:v1_relat_1(k2_zfmisc_1(X27,X28)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_54, plain, ![X50, X51]:(v1_xboole_0(X50)|~m1_subset_1(X51,X50)|m1_subset_1(k6_domain_1(X50,X51),k1_zfmisc_1(X50))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_55, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)|~r2_hidden(esk2_0,k6_domain_1(u1_struct_0(esk1_0),esk2_0))), inference(spm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_56, plain, (v1_xboole_0(X1)|r2_hidden(X2,k6_domain_1(X1,X2))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_47, c_0_26])).
fof(c_0_57, plain, ![X48, X49]:((~v6_orders_2(X49,X48)|r7_relat_2(u1_orders_2(X48),X49)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X48)))|~l1_orders_2(X48))&(~r7_relat_2(u1_orders_2(X48),X49)|v6_orders_2(X49,X48)|~m1_subset_1(X49,k1_zfmisc_1(u1_struct_0(X48)))|~l1_orders_2(X48))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d11_orders_2])])])])).
cnf(c_0_58, plain, (r7_relat_2(X1,k2_tarski(X2,X2))|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,esk7_2(X1,k2_tarski(X2,X2))),X1)), inference(spm,[status(thm)],[c_0_48, c_0_49])).
cnf(c_0_59, plain, (esk7_2(X1,k2_tarski(X2,X2))=X2|r7_relat_2(X1,k2_tarski(X2,X2))|~v1_relat_1(X1)), inference(spm,[status(thm)],[c_0_25, c_0_50])).
cnf(c_0_60, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_61, plain, (m1_subset_1(u1_orders_2(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_62, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_63, negated_conjecture, (~v6_orders_2(k6_domain_1(u1_struct_0(esk1_0),esk2_0),esk1_0)|~m1_subset_1(k6_domain_1(u1_struct_0(esk1_0),esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_64, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_65, negated_conjecture, (k6_domain_1(u1_struct_0(esk1_0),esk2_0)=k2_tarski(esk2_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_37])]), c_0_38])).
cnf(c_0_66, plain, (v6_orders_2(X2,X1)|~r7_relat_2(u1_orders_2(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_67, plain, (r7_relat_2(X1,k2_tarski(X2,X2))|~v1_relat_1(X1)|~r2_hidden(k4_tarski(X2,X2),X1)), inference(spm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_68, plain, (v1_relat_1(u1_orders_2(X1))|~l1_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_62])])).
fof(c_0_69, plain, ![X7, X8, X9]:((~r1_orders_2(X7,X8,X9)|r2_hidden(k4_tarski(X8,X9),u1_orders_2(X7))|~m1_subset_1(X9,u1_struct_0(X7))|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7))&(~r2_hidden(k4_tarski(X8,X9),u1_orders_2(X7))|r1_orders_2(X7,X8,X9)|~m1_subset_1(X9,u1_struct_0(X7))|~m1_subset_1(X8,u1_struct_0(X7))|~l1_orders_2(X7))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_orders_2])])])])).
cnf(c_0_70, negated_conjecture, (~v6_orders_2(k2_tarski(esk2_0,esk2_0),esk1_0)|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_26]), c_0_37])]), c_0_38])).
cnf(c_0_71, negated_conjecture, (m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_37])]), c_0_38])).
cnf(c_0_72, plain, (v6_orders_2(k2_tarski(X1,X1),X2)|~r2_hidden(k4_tarski(X1,X1),u1_orders_2(X2))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_68])).
cnf(c_0_73, plain, (r2_hidden(k4_tarski(X2,X3),u1_orders_2(X1))|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_74, negated_conjecture, (~v6_orders_2(k2_tarski(esk2_0,esk2_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_70, c_0_71])])).
cnf(c_0_75, plain, (v6_orders_2(k2_tarski(X1,X1),X2)|~r1_orders_2(X2,X1,X1)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(spm,[status(thm)],[c_0_72, c_0_73])).
fof(c_0_76, plain, ![X45, X46]:(v2_struct_0(X45)|(~v3_orders_2(X45)|~l1_orders_2(X45))|(~m1_subset_1(X46,u1_struct_0(X45))|r1_orders_2(X45,X46,X46))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
cnf(c_0_77, negated_conjecture, (~r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75]), c_0_71]), c_0_37]), c_0_35])])).
cnf(c_0_78, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_79, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78]), c_0_37]), c_0_35]), c_0_79])]), c_0_33]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 50
# Proof object formula steps           : 31
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 23
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 20
# Proof object simplifying inferences  : 34
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 43
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 42
# Processed clauses                    : 2406
# ...of these trivial                  : 31
# ...subsumed                          : 1009
# ...remaining for further processing  : 1366
# Other redundant clauses eliminated   : 131
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 67
# Backward-rewritten                   : 113
# Generated clauses                    : 38944
# ...of the previous two non-trivial   : 36196
# Contextual simplify-reflections      : 52
# Paramodulations                      : 38759
# Factorizations                       : 57
# Equation resolutions                 : 131
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1140
#    Positive orientable unit clauses  : 22
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 1113
# Current number of unprocessed clauses: 33543
# ...number of literals in the above   : 231139
# Current number of archived formulas  : 0
# Current number of archived clauses   : 222
# Clause-clause subsumption calls (NU) : 199363
# Rec. Clause-clause subsumption calls : 54377
# Non-unit clause-clause subsumptions  : 1030
# Unit Clause-clause subsumption calls : 442
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 135
# BW rewrite match successes           : 20
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 982362

# -------------------------------------------------
# User time                : 0.580 s
# System time              : 0.003 s
# Total time               : 0.584 s
# Maximum resident set size: 35692 pages
