# No SInE strategy applied
# Auto-Mode selected heuristic G_E___207_C01_F1_SE_CS_SP_PI_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(cc5_funct_1, axiom, ![X1]:(v1_zfmisc_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_zfmisc_1(X2))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', cc5_funct_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', t3_subset)).
fof(d1_xboole_0, axiom, ![X1]:(v1_xboole_0(X1)<=>![X2]:~(r2_hidden(X2,X1))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', d1_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', d3_tarski)).
fof(rc1_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&(~(v1_xboole_0(X2))&v1_zfmisc_1(X2)))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', rc1_tex_2)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', d2_tarski)).
fof(rc1_connsp_1, axiom, ![X1]:(l1_pre_topc(X1)=>?[X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))&v1_xboole_0(X2))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', rc1_connsp_1)).
fof(t20_tex_2, conjecture, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', t20_tex_2)).
fof(t13_tex_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(X3=u1_struct_0(X2)=>(v1_subset_1(X3,u1_struct_0(X1))<=>v1_tex_2(X2,X1)))))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', t13_tex_2)).
fof(t1_tsep_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1))))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', t1_tsep_1)).
fof(d7_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(v1_subset_1(X2,X1)<=>~(X2=X1))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', d7_subset_1)).
fof(t7_tex_2, axiom, ![X1]:((~(v1_xboole_0(X1))&~(v1_zfmisc_1(X1)))=>![X2]:(m1_subset_1(X2,X1)=>v1_subset_1(k6_domain_1(X1,X2),X1))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', t7_tex_2)).
fof(dt_k1_tex_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v1_pre_topc(k1_tex_2(X1,X2))&m1_pre_topc(k1_tex_2(X1,X2),X1)))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', dt_k1_tex_2)).
fof(t8_tex_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>~((v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))&v7_struct_0(X1))))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', t8_tex_2)).
fof(fc2_tex_2, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_pre_topc(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>(~(v2_struct_0(k1_tex_2(X1,X2)))&(v7_struct_0(k1_tex_2(X1,X2))&v1_pre_topc(k1_tex_2(X1,X2))))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', fc2_tex_2)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', fc1_struct_0)).
fof(cc2_tex_2, axiom, ![X1]:((~(v1_xboole_0(X1))&v1_zfmisc_1(X1))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>(~(v1_xboole_0(X2))=>~(v1_subset_1(X2,X1))))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', cc2_tex_2)).
fof(cc4_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>~(v1_subset_1(X2,X1)))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', cc4_subset_1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', dt_l1_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', dt_m1_pre_topc)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', fc2_struct_0)).
fof(t6_tex_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>~((v1_subset_1(k6_domain_1(X1,X2),X1)&v1_zfmisc_1(X1))))), file('data/MPTP/proofs/fb50bc18728d41d5ba0da68e108ac37b.E_input', t6_tex_2)).
fof(c_0_22, plain, ![X5, X6]:(~v1_zfmisc_1(X5)|(~m1_subset_1(X6,k1_zfmisc_1(X5))|v1_zfmisc_1(X6))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_funct_1])])])).
fof(c_0_23, plain, ![X7, X8]:((~m1_subset_1(X7,k1_zfmisc_1(X8))|r1_tarski(X7,X8))&(~r1_tarski(X7,X8)|m1_subset_1(X7,k1_zfmisc_1(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_24, plain, ![X9, X10, X11]:((~v1_xboole_0(X9)|~r2_hidden(X10,X9))&(r2_hidden(esk1_1(X11),X11)|v1_xboole_0(X11))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_xboole_0])])])])])])).
fof(c_0_25, plain, ![X13, X14, X15, X16, X17]:((~r1_tarski(X13,X14)|(~r2_hidden(X15,X13)|r2_hidden(X15,X14)))&((r2_hidden(esk2_2(X16,X17),X16)|r1_tarski(X16,X17))&(~r2_hidden(esk2_2(X16,X17),X17)|r1_tarski(X16,X17)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_26, plain, (v1_zfmisc_1(X2)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_27, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_28, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, plain, (v1_zfmisc_1(X1)|~r1_tarski(X1,X2)|~v1_zfmisc_1(X2)), inference(spm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_28, c_0_29])).
fof(c_0_32, plain, ![X19]:((m1_subset_1(esk3_1(X19),k1_zfmisc_1(X19))|v1_xboole_0(X19))&((~v1_xboole_0(esk3_1(X19))|v1_xboole_0(X19))&(v1_zfmisc_1(esk3_1(X19))|v1_xboole_0(X19)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc1_tex_2])])])])])).
fof(c_0_33, plain, ![X21, X22, X23, X24, X25, X26, X27, X28]:(((~r2_hidden(X24,X23)|(X24=X21|X24=X22)|X23!=k2_tarski(X21,X22))&((X25!=X21|r2_hidden(X25,X23)|X23!=k2_tarski(X21,X22))&(X25!=X22|r2_hidden(X25,X23)|X23!=k2_tarski(X21,X22))))&(((esk4_3(X26,X27,X28)!=X26|~r2_hidden(esk4_3(X26,X27,X28),X28)|X28=k2_tarski(X26,X27))&(esk4_3(X26,X27,X28)!=X27|~r2_hidden(esk4_3(X26,X27,X28),X28)|X28=k2_tarski(X26,X27)))&(r2_hidden(esk4_3(X26,X27,X28),X28)|(esk4_3(X26,X27,X28)=X26|esk4_3(X26,X27,X28)=X27)|X28=k2_tarski(X26,X27)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_34, plain, (v1_zfmisc_1(X1)|~v1_xboole_0(X1)|~v1_zfmisc_1(X2)), inference(spm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_35, plain, (v1_zfmisc_1(esk3_1(X1))|v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_32])).
fof(c_0_36, plain, ![X30]:((m1_subset_1(esk5_1(X30),k1_zfmisc_1(u1_struct_0(X30)))|~l1_pre_topc(X30))&(v1_xboole_0(esk5_1(X30))|~l1_pre_topc(X30))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[rc1_connsp_1])])])])).
fof(c_0_37, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v1_tex_2(k1_tex_2(X1,X2),X1)<=>v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1)))))), inference(assume_negation,[status(cth)],[t20_tex_2])).
fof(c_0_38, plain, ![X34, X35, X36]:((~v1_subset_1(X36,u1_struct_0(X34))|v1_tex_2(X35,X34)|X36!=u1_struct_0(X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X34)))|~m1_pre_topc(X35,X34)|~l1_pre_topc(X34))&(~v1_tex_2(X35,X34)|v1_subset_1(X36,u1_struct_0(X34))|X36!=u1_struct_0(X35)|~m1_subset_1(X36,k1_zfmisc_1(u1_struct_0(X34)))|~m1_pre_topc(X35,X34)|~l1_pre_topc(X34))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_tex_2])])])])).
fof(c_0_39, plain, ![X37, X38]:(~l1_pre_topc(X37)|(~m1_pre_topc(X38,X37)|m1_subset_1(u1_struct_0(X38),k1_zfmisc_1(u1_struct_0(X37))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_tsep_1])])])).
fof(c_0_40, plain, ![X39, X40]:((~v1_subset_1(X40,X39)|X40!=X39|~m1_subset_1(X40,k1_zfmisc_1(X39)))&(X40=X39|v1_subset_1(X40,X39)|~m1_subset_1(X40,k1_zfmisc_1(X39)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_subset_1])])])])).
cnf(c_0_41, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_42, plain, (v1_xboole_0(X1)|v1_zfmisc_1(X2)|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_43, plain, (v1_xboole_0(esk5_1(X1))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_44, negated_conjecture, ((~v2_struct_0(esk6_0)&l1_pre_topc(esk6_0))&(m1_subset_1(esk7_0,u1_struct_0(esk6_0))&((~v1_tex_2(k1_tex_2(esk6_0,esk7_0),esk6_0)|~v1_subset_1(k6_domain_1(u1_struct_0(esk6_0),esk7_0),u1_struct_0(esk6_0)))&(v1_tex_2(k1_tex_2(esk6_0,esk7_0),esk6_0)|v1_subset_1(k6_domain_1(u1_struct_0(esk6_0),esk7_0),u1_struct_0(esk6_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_37])])])])).
cnf(c_0_45, plain, (v1_tex_2(X3,X2)|~v1_subset_1(X1,u1_struct_0(X2))|X1!=u1_struct_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X3,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_46, plain, (m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_47, plain, (X1=X2|v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_48, plain, (r2_hidden(X1,k2_tarski(X2,X1))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_41])])).
cnf(c_0_49, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk5_1(X2))|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_50, negated_conjecture, (l1_pre_topc(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_51, plain, ![X41, X42]:(v1_xboole_0(X41)|v1_zfmisc_1(X41)|(~m1_subset_1(X42,X41)|v1_subset_1(k6_domain_1(X41,X42),X41))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_tex_2])])])])).
cnf(c_0_52, plain, (v1_tex_2(X1,X2)|~m1_pre_topc(X1,X2)|~v1_subset_1(u1_struct_0(X1),u1_struct_0(X2))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_45]), c_0_46])).
cnf(c_0_53, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v1_subset_1(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_47, c_0_46])).
fof(c_0_54, plain, ![X43, X44]:((~v2_struct_0(k1_tex_2(X43,X44))|(v2_struct_0(X43)|~l1_pre_topc(X43)|~m1_subset_1(X44,u1_struct_0(X43))))&((v1_pre_topc(k1_tex_2(X43,X44))|(v2_struct_0(X43)|~l1_pre_topc(X43)|~m1_subset_1(X44,u1_struct_0(X43))))&(m1_pre_topc(k1_tex_2(X43,X44),X43)|(v2_struct_0(X43)|~l1_pre_topc(X43)|~m1_subset_1(X44,u1_struct_0(X43)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_tex_2])])])])).
cnf(c_0_55, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(spm,[status(thm)],[c_0_28, c_0_48])).
cnf(c_0_56, negated_conjecture, (v1_xboole_0(X1)|v1_zfmisc_1(esk5_1(esk6_0))), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_57, negated_conjecture, (~v1_tex_2(k1_tex_2(esk6_0,esk7_0),esk6_0)|~v1_subset_1(k6_domain_1(u1_struct_0(esk6_0),esk7_0),u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_58, plain, (v1_xboole_0(X1)|v1_zfmisc_1(X1)|v1_subset_1(k6_domain_1(X1,X2),X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_60, plain, (u1_struct_0(X1)=u1_struct_0(X2)|v1_tex_2(X1,X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_61, plain, (m1_pre_topc(k1_tex_2(X1,X2),X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_62, negated_conjecture, (v1_zfmisc_1(esk5_1(esk6_0))), inference(spm,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_63, plain, ![X45, X46]:(v2_struct_0(X45)|~l1_struct_0(X45)|(~m1_subset_1(X46,u1_struct_0(X45))|(~v1_subset_1(k6_domain_1(u1_struct_0(X45),X46),u1_struct_0(X45))|~v7_struct_0(X45)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_tex_2])])])])).
fof(c_0_64, plain, ![X48, X49]:((~v2_struct_0(k1_tex_2(X48,X49))|(v2_struct_0(X48)|~l1_pre_topc(X48)|~m1_subset_1(X49,u1_struct_0(X48))))&((v7_struct_0(k1_tex_2(X48,X49))|(v2_struct_0(X48)|~l1_pre_topc(X48)|~m1_subset_1(X49,u1_struct_0(X48))))&(v1_pre_topc(k1_tex_2(X48,X49))|(v2_struct_0(X48)|~l1_pre_topc(X48)|~m1_subset_1(X49,u1_struct_0(X48)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_tex_2])])])])).
fof(c_0_65, plain, ![X47]:(~v2_struct_0(X47)|~l1_struct_0(X47)|v1_xboole_0(u1_struct_0(X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_66, negated_conjecture, (v1_xboole_0(u1_struct_0(esk6_0))|v1_zfmisc_1(u1_struct_0(esk6_0))|~v1_tex_2(k1_tex_2(esk6_0,esk7_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58]), c_0_59])])).
cnf(c_0_67, plain, (u1_struct_0(k1_tex_2(X1,X2))=u1_struct_0(X1)|v1_tex_2(k1_tex_2(X1,X2),X1)|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_60, c_0_61])).
cnf(c_0_68, negated_conjecture, (~v2_struct_0(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_69, negated_conjecture, (v1_zfmisc_1(X1)|~v1_xboole_0(X1)), inference(spm,[status(thm)],[c_0_34, c_0_62])).
cnf(c_0_70, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_subset_1(k6_domain_1(u1_struct_0(X1),X2),u1_struct_0(X1))|~v7_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_71, plain, (v7_struct_0(k1_tex_2(X1,X2))|v2_struct_0(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_72, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_73, negated_conjecture, (u1_struct_0(k1_tex_2(esk6_0,esk7_0))=u1_struct_0(esk6_0)|v1_zfmisc_1(u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_50]), c_0_59])]), c_0_68]), c_0_69])).
fof(c_0_74, plain, ![X50, X51]:(v1_xboole_0(X50)|~v1_zfmisc_1(X50)|(~m1_subset_1(X51,k1_zfmisc_1(X50))|(v1_xboole_0(X51)|~v1_subset_1(X51,X50)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_tex_2])])])])).
fof(c_0_75, plain, ![X52, X53]:(~v1_xboole_0(X52)|(~m1_subset_1(X53,k1_zfmisc_1(X52))|~v1_subset_1(X53,X52))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc4_subset_1])])])])).
cnf(c_0_76, plain, (v2_struct_0(X1)|v1_zfmisc_1(u1_struct_0(X1))|~v7_struct_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_58]), c_0_69])).
cnf(c_0_77, negated_conjecture, (v7_struct_0(k1_tex_2(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_59]), c_0_50])]), c_0_68])).
cnf(c_0_78, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk6_0))|~l1_struct_0(k1_tex_2(esk6_0,esk7_0))|~v2_struct_0(k1_tex_2(esk6_0,esk7_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_69])).
fof(c_0_79, plain, ![X54]:(~l1_pre_topc(X54)|l1_struct_0(X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_80, plain, ![X55, X56]:(~l1_pre_topc(X55)|(~m1_pre_topc(X56,X55)|l1_pre_topc(X56))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
cnf(c_0_81, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_82, plain, (~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_83, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk6_0))|~l1_struct_0(k1_tex_2(esk6_0,esk7_0))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_73]), c_0_77])]), c_0_78])).
cnf(c_0_84, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_85, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_86, plain, (v1_xboole_0(X1)|~v1_subset_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~v1_zfmisc_1(X2)), inference(csr,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_87, plain, (v1_subset_1(X3,u1_struct_0(X2))|~v1_tex_2(X1,X2)|X3!=u1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_88, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk6_0))|~l1_pre_topc(k1_tex_2(esk6_0,esk7_0))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_89, plain, (v2_struct_0(X1)|l1_pre_topc(k1_tex_2(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_85, c_0_61])).
cnf(c_0_90, plain, (v1_xboole_0(u1_struct_0(X1))|~m1_pre_topc(X1,X2)|~v1_subset_1(u1_struct_0(X1),u1_struct_0(X2))|~l1_pre_topc(X2)|~v1_zfmisc_1(u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_86, c_0_46])).
cnf(c_0_91, plain, (v1_subset_1(u1_struct_0(X1),u1_struct_0(X2))|~m1_pre_topc(X1,X2)|~v1_tex_2(X1,X2)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_87]), c_0_46])).
cnf(c_0_92, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk6_0))|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89]), c_0_50]), c_0_59])]), c_0_68])).
fof(c_0_93, plain, ![X57]:(v2_struct_0(X57)|~l1_struct_0(X57)|~v1_xboole_0(u1_struct_0(X57))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_94, plain, (v1_xboole_0(u1_struct_0(X1))|~m1_pre_topc(X1,X2)|~v1_tex_2(X1,X2)|~l1_pre_topc(X2)|~v1_zfmisc_1(u1_struct_0(X2))), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_95, negated_conjecture, (v1_zfmisc_1(u1_struct_0(esk6_0))), inference(spm,[status(thm)],[c_0_92, c_0_59])).
fof(c_0_96, plain, ![X58, X59]:(v1_xboole_0(X58)|(~m1_subset_1(X59,X58)|(~v1_subset_1(k6_domain_1(X58,X59),X58)|~v1_zfmisc_1(X58)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_tex_2])])])])).
cnf(c_0_97, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_98, negated_conjecture, (v1_xboole_0(u1_struct_0(X1))|~m1_pre_topc(X1,esk6_0)|~v1_tex_2(X1,esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_50])])).
cnf(c_0_99, plain, (v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~v1_subset_1(k6_domain_1(X1,X2),X1)|~v1_zfmisc_1(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_100, negated_conjecture, (v1_tex_2(k1_tex_2(esk6_0,esk7_0),esk6_0)|v1_subset_1(k6_domain_1(u1_struct_0(esk6_0),esk7_0),u1_struct_0(esk6_0))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_101, negated_conjecture, (v2_struct_0(X1)|~l1_struct_0(X1)|~m1_pre_topc(X1,esk6_0)|~v1_tex_2(X1,esk6_0)), inference(spm,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_102, negated_conjecture, (v1_tex_2(k1_tex_2(esk6_0,esk7_0),esk6_0)|v1_xboole_0(u1_struct_0(esk6_0))|~v1_zfmisc_1(u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_59])])).
cnf(c_0_103, negated_conjecture, (v2_struct_0(k1_tex_2(esk6_0,X1))|~l1_struct_0(k1_tex_2(esk6_0,X1))|~v1_tex_2(k1_tex_2(esk6_0,X1),esk6_0)|~m1_subset_1(X1,u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_61]), c_0_50])]), c_0_68])).
cnf(c_0_104, negated_conjecture, (v1_tex_2(k1_tex_2(esk6_0,esk7_0),esk6_0)|v1_xboole_0(u1_struct_0(esk6_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_102, c_0_95])])).
cnf(c_0_105, plain, (v2_struct_0(X1)|~v2_struct_0(k1_tex_2(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_106, negated_conjecture, (v2_struct_0(k1_tex_2(esk6_0,esk7_0))|v1_xboole_0(u1_struct_0(esk6_0))|~l1_struct_0(k1_tex_2(esk6_0,esk7_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_59])])).
cnf(c_0_107, negated_conjecture, (v1_xboole_0(u1_struct_0(esk6_0))|~l1_struct_0(k1_tex_2(esk6_0,esk7_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106]), c_0_50]), c_0_59])]), c_0_68])).
cnf(c_0_108, negated_conjecture, (v1_xboole_0(u1_struct_0(esk6_0))|~l1_pre_topc(k1_tex_2(esk6_0,esk7_0))), inference(spm,[status(thm)],[c_0_107, c_0_84])).
cnf(c_0_109, negated_conjecture, (v1_xboole_0(u1_struct_0(esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_89]), c_0_50]), c_0_59])]), c_0_68])).
cnf(c_0_110, negated_conjecture, (~l1_struct_0(esk6_0)), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_109]), c_0_68])).
cnf(c_0_111, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_84]), c_0_50])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 112
# Proof object clause steps            : 67
# Proof object formula steps           : 45
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 22
# Proof object generating inferences   : 34
# Proof object simplifying inferences  : 48
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 22
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 44
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 44
# Processed clauses                    : 253
# ...of these trivial                  : 1
# ...subsumed                          : 91
# ...remaining for further processing  : 161
# Other redundant clauses eliminated   : 30
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 14
# Generated clauses                    : 404
# ...of the previous two non-trivial   : 329
# Contextual simplify-reflections      : 13
# Paramodulations                      : 360
# Factorizations                       : 16
# Equation resolutions                 : 30
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 118
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 103
# Current number of unprocessed clauses: 96
# ...number of literals in the above   : 404
# Current number of archived formulas  : 0
# Current number of archived clauses   : 37
# Clause-clause subsumption calls (NU) : 2934
# Rec. Clause-clause subsumption calls : 1132
# Non-unit clause-clause subsumptions  : 113
# Unit Clause-clause subsumption calls : 82
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 7
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8636

# -------------------------------------------------
# User time                : 0.015 s
# System time              : 0.000 s
# Total time               : 0.015 s
# Maximum resident set size: 28652 pages
