# No SInE strategy applied
# Auto-Mode selected heuristic G_E___208_C18_F1_SE_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t94_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2)), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t94_xboole_1)).
fof(t48_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t48_xboole_1)).
fof(t2_boole, axiom, ![X1]:k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0, file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t2_boole)).
fof(t34_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),k1_tarski(X4)))<=>(X1=X3&X2=X4)), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t34_zfmisc_1)).
fof(t69_enumset1, axiom, ![X1]:k2_tarski(X1,X1)=k1_tarski(X1), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t69_enumset1)).
fof(t77_enumset1, axiom, ![X1, X2]:k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t77_enumset1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', idempotence_k2_xboole_0)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t3_boole)).
fof(t23_relat_1, conjecture, ![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2)))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t23_relat_1)).
fof(t36_zfmisc_1, axiom, ![X1, X2, X3]:(k2_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))&k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t36_zfmisc_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t41_xboole_1)).
fof(symmetry_r1_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)=>r1_xboole_0(X2,X1)), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', symmetry_r1_xboole_0)).
fof(t83_xboole_1, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k4_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t83_xboole_1)).
fof(t65_zfmisc_1, axiom, ![X1, X2]:(k4_xboole_0(X1,k1_tarski(X2))=X1<=>~(r2_hidden(X2,X1))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t65_zfmisc_1)).
fof(d4_relat_1, axiom, ![X1, X2]:(X2=k1_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X3,X4),X1))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', d4_relat_1)).
fof(t128_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k1_tarski(X3),X4))<=>(X1=X3&r2_hidden(X2,X4))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t128_zfmisc_1)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t3_xboole_0)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', l32_xboole_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t36_xboole_1)).
fof(d5_relat_1, axiom, ![X1, X2]:(X2=k2_relat_1(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:r2_hidden(k4_tarski(X4,X3),X1))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', d5_relat_1)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', commutativity_k3_xboole_0)).
fof(t129_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,k1_tarski(X4)))<=>(r2_hidden(X1,X3)&X2=X4)), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', t129_zfmisc_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', d10_xboole_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('data/MPTP/proofs/fc06bd4bf0cd407c9049186136aa5c3d.E_input', d3_tarski)).
fof(c_0_24, plain, ![X5, X6]:k2_xboole_0(X5,X6)=k5_xboole_0(k5_xboole_0(X5,X6),k3_xboole_0(X5,X6)), inference(variable_rename,[status(thm)],[t94_xboole_1])).
fof(c_0_25, plain, ![X7, X8]:k4_xboole_0(X7,k4_xboole_0(X7,X8))=k3_xboole_0(X7,X8), inference(variable_rename,[status(thm)],[t48_xboole_1])).
fof(c_0_26, plain, ![X9]:k3_xboole_0(X9,k1_xboole_0)=k1_xboole_0, inference(variable_rename,[status(thm)],[t2_boole])).
fof(c_0_27, plain, ![X38, X39, X40, X41]:(((X38=X40|~r2_hidden(k4_tarski(X38,X39),k2_zfmisc_1(k1_tarski(X40),k1_tarski(X41))))&(X39=X41|~r2_hidden(k4_tarski(X38,X39),k2_zfmisc_1(k1_tarski(X40),k1_tarski(X41)))))&(X38!=X40|X39!=X41|r2_hidden(k4_tarski(X38,X39),k2_zfmisc_1(k1_tarski(X40),k1_tarski(X41))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t34_zfmisc_1])])])).
fof(c_0_28, plain, ![X24]:k2_tarski(X24,X24)=k1_tarski(X24), inference(variable_rename,[status(thm)],[t69_enumset1])).
fof(c_0_29, plain, ![X25, X26]:k2_enumset1(X25,X25,X25,X26)=k2_tarski(X25,X26), inference(variable_rename,[status(thm)],[t77_enumset1])).
fof(c_0_30, plain, ![X10]:k2_xboole_0(X10,X10)=X10, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_31, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k3_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_32, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_33, plain, (k3_xboole_0(X1,k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_34, plain, ![X11]:k4_xboole_0(X11,k1_xboole_0)=X11, inference(variable_rename,[status(thm)],[t3_boole])).
fof(c_0_35, negated_conjecture, ~(![X1, X2, X3]:(v1_relat_1(X3)=>(X3=k1_tarski(k4_tarski(X1,X2))=>(k1_relat_1(X3)=k1_tarski(X1)&k2_relat_1(X3)=k1_tarski(X2))))), inference(assume_negation,[status(cth)],[t23_relat_1])).
cnf(c_0_36, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),k1_tarski(X4)))|X1!=X2|X3!=X4), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_37, plain, (k2_tarski(X1,X1)=k1_tarski(X1)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_38, plain, (k2_enumset1(X1,X1,X1,X2)=k2_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_39, plain, ![X31, X32, X33]:(k2_zfmisc_1(k1_tarski(X31),k2_tarski(X32,X33))=k2_tarski(k4_tarski(X31,X32),k4_tarski(X31,X33))&k2_zfmisc_1(k2_tarski(X31,X32),k1_tarski(X33))=k2_tarski(k4_tarski(X31,X33),k4_tarski(X32,X33))), inference(variable_rename,[status(thm)],[t36_zfmisc_1])).
fof(c_0_40, plain, ![X15, X16, X17]:k4_xboole_0(k4_xboole_0(X15,X16),X17)=k4_xboole_0(X15,k2_xboole_0(X16,X17)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
cnf(c_0_41, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_42, plain, (k2_xboole_0(X1,X2)=k5_xboole_0(k5_xboole_0(X1,X2),k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(rw,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_43, plain, (k4_xboole_0(X1,k4_xboole_0(X1,k1_xboole_0))=k1_xboole_0), inference(rw,[status(thm)],[c_0_33, c_0_32])).
cnf(c_0_44, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_45, plain, ![X18, X19]:(~r1_xboole_0(X18,X19)|r1_xboole_0(X19,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_r1_xboole_0])])).
fof(c_0_46, plain, ![X20, X21]:((~r1_xboole_0(X20,X21)|k4_xboole_0(X20,X21)=X20)&(k4_xboole_0(X20,X21)!=X20|r1_xboole_0(X20,X21))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t83_xboole_1])])).
fof(c_0_47, plain, ![X22, X23]:((k4_xboole_0(X22,k1_tarski(X23))!=X22|~r2_hidden(X23,X22))&(r2_hidden(X23,X22)|k4_xboole_0(X22,k1_tarski(X23))=X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t65_zfmisc_1])])])).
fof(c_0_48, negated_conjecture, (v1_relat_1(esk3_0)&(esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))&(k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_35])])])).
cnf(c_0_49, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),k2_enumset1(X4,X4,X4,X4)))|X3!=X4|X1!=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_37]), c_0_37]), c_0_38]), c_0_38])).
cnf(c_0_50, plain, (k2_zfmisc_1(k2_tarski(X1,X2),k1_tarski(X3))=k2_tarski(k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, plain, (k5_xboole_0(k5_xboole_0(X1,X1),k4_xboole_0(X1,k4_xboole_0(X1,X1)))=X1), inference(rw,[status(thm)],[c_0_41, c_0_42])).
cnf(c_0_53, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(rw,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_54, plain, (r1_xboole_0(X2,X1)|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_55, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X1,X2)!=X1), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_56, plain, (r2_hidden(X1,X2)|k4_xboole_0(X2,k1_tarski(X1))=X2), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, negated_conjecture, (esk3_0=k1_tarski(k4_tarski(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_58, plain, ![X42, X43, X44, X46, X47, X48, X49, X51]:(((~r2_hidden(X44,X43)|r2_hidden(k4_tarski(X44,esk4_3(X42,X43,X44)),X42)|X43!=k1_relat_1(X42))&(~r2_hidden(k4_tarski(X46,X47),X42)|r2_hidden(X46,X43)|X43!=k1_relat_1(X42)))&((~r2_hidden(esk5_2(X48,X49),X49)|~r2_hidden(k4_tarski(esk5_2(X48,X49),X51),X48)|X49=k1_relat_1(X48))&(r2_hidden(esk5_2(X48,X49),X49)|r2_hidden(k4_tarski(esk5_2(X48,X49),esk6_2(X48,X49)),X48)|X49=k1_relat_1(X48)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_relat_1])])])])])])).
cnf(c_0_59, plain, (r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X2,X2)))), inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_49])])).
cnf(c_0_60, plain, (k2_zfmisc_1(k2_enumset1(X1,X1,X1,X2),k2_enumset1(X3,X3,X3,X3))=k2_enumset1(k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X1,X3),k4_tarski(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_37]), c_0_38]), c_0_38]), c_0_38])).
fof(c_0_61, plain, ![X27, X28, X29, X30]:(((X27=X29|~r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(k1_tarski(X29),X30)))&(r2_hidden(X28,X30)|~r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(k1_tarski(X29),X30))))&(X27!=X29|~r2_hidden(X28,X30)|r2_hidden(k4_tarski(X27,X28),k2_zfmisc_1(k1_tarski(X29),X30)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t128_zfmisc_1])])])).
cnf(c_0_62, plain, (k4_xboole_0(X1,k1_tarski(X2))!=X1|~r2_hidden(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_63, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k5_xboole_0(k5_xboole_0(X2,X3),k4_xboole_0(X2,k4_xboole_0(X2,X3))))), inference(rw,[status(thm)],[c_0_51, c_0_42])).
cnf(c_0_64, plain, (k5_xboole_0(k5_xboole_0(X1,X1),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53]), c_0_44])).
cnf(c_0_65, plain, (k4_xboole_0(X1,X2)=X1|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_66, plain, (r1_xboole_0(X1,X2)|k4_xboole_0(X2,X1)!=X2), inference(spm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_67, plain, (k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X1))=X2|r2_hidden(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_37]), c_0_38])).
cnf(c_0_68, negated_conjecture, (esk3_0=k2_enumset1(k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0),k4_tarski(esk1_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_37]), c_0_38])).
cnf(c_0_69, plain, (r2_hidden(X1,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_70, plain, (r2_hidden(k4_tarski(X1,X2),k2_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2)))), inference(rw,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_71, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k1_tarski(X2),X4))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_72, plain, (k2_zfmisc_1(k1_tarski(X1),k2_tarski(X2,X3))=k2_tarski(k4_tarski(X1,X2),k4_tarski(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_73, plain, (k4_xboole_0(X1,k2_enumset1(X2,X2,X2,X2))!=X1|~r2_hidden(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_37]), c_0_38])).
cnf(c_0_74, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X2)=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_53]), c_0_44]), c_0_64])).
cnf(c_0_75, plain, (k4_xboole_0(X1,X2)=X1|k4_xboole_0(X2,X1)!=X2), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_76, negated_conjecture, (k4_xboole_0(X1,esk3_0)=X1|r2_hidden(k4_tarski(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_67, c_0_68])).
fof(c_0_77, plain, ![X64, X65, X67, X68, X69]:(((r2_hidden(esk10_2(X64,X65),X64)|r1_xboole_0(X64,X65))&(r2_hidden(esk10_2(X64,X65),X65)|r1_xboole_0(X64,X65)))&(~r2_hidden(X69,X67)|~r2_hidden(X69,X68)|~r1_xboole_0(X67,X68))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
cnf(c_0_78, plain, (r2_hidden(X1,k1_relat_1(X2))|~r2_hidden(k4_tarski(X1,X3),X2)), inference(er,[status(thm)],[c_0_69])).
cnf(c_0_79, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),esk3_0)), inference(spm,[status(thm)],[c_0_70, c_0_68])).
cnf(c_0_80, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(k2_enumset1(X2,X2,X2,X2),X4))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_37]), c_0_38])).
cnf(c_0_81, plain, (k2_zfmisc_1(k2_enumset1(X1,X1,X1,X1),k2_enumset1(X2,X2,X2,X3))=k2_enumset1(k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X2),k4_tarski(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_72, c_0_37]), c_0_38]), c_0_38]), c_0_38])).
cnf(c_0_82, plain, (~r2_hidden(X1,k4_xboole_0(X2,k2_enumset1(X1,X1,X1,X1)))), inference(spm,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_83, negated_conjecture, (k4_xboole_0(esk3_0,X1)=esk3_0|r2_hidden(k4_tarski(esk1_0,esk2_0),X1)), inference(spm,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_84, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r1_xboole_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_85, negated_conjecture, (r2_hidden(esk1_0,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_78, c_0_79])).
fof(c_0_86, plain, ![X80, X81]:((k4_xboole_0(X80,X81)!=k1_xboole_0|r1_tarski(X80,X81))&(~r1_tarski(X80,X81)|k4_xboole_0(X80,X81)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
fof(c_0_87, plain, ![X78, X79]:r1_tarski(k4_xboole_0(X78,X79),X78), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_88, plain, ![X53, X54, X55, X57, X58, X59, X60, X62]:(((~r2_hidden(X55,X54)|r2_hidden(k4_tarski(esk7_3(X53,X54,X55),X55),X53)|X54!=k2_relat_1(X53))&(~r2_hidden(k4_tarski(X58,X57),X53)|r2_hidden(X57,X54)|X54!=k2_relat_1(X53)))&((~r2_hidden(esk8_2(X59,X60),X60)|~r2_hidden(k4_tarski(X62,esk8_2(X59,X60)),X59)|X60=k2_relat_1(X59))&(r2_hidden(esk8_2(X59,X60),X60)|r2_hidden(k4_tarski(esk9_2(X59,X60),esk8_2(X59,X60)),X59)|X60=k2_relat_1(X59)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_relat_1])])])])])])).
fof(c_0_89, plain, ![X70, X71]:k3_xboole_0(X70,X71)=k3_xboole_0(X71,X70), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
cnf(c_0_90, plain, (X1=X2|~r2_hidden(k4_tarski(X1,X3),k2_enumset1(k4_tarski(X2,X4),k4_tarski(X2,X4),k4_tarski(X2,X4),k4_tarski(X2,X5)))), inference(spm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_91, negated_conjecture, (r2_hidden(k4_tarski(esk1_0,esk2_0),k2_enumset1(X1,X1,X1,X1))|~r2_hidden(X1,esk3_0)), inference(spm,[status(thm)],[c_0_82, c_0_83])).
cnf(c_0_92, plain, (r2_hidden(k4_tarski(X1,esk4_3(X3,X2,X1)),X3)|~r2_hidden(X1,X2)|X2!=k1_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_93, negated_conjecture, (~r2_hidden(esk1_0,X1)|~r1_xboole_0(X1,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_84, c_0_85])).
cnf(c_0_94, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_95, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_96, plain, (r2_hidden(X2,X4)|~r2_hidden(k4_tarski(X1,X2),X3)|X4!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_97, plain, ![X34, X35, X36, X37]:(((r2_hidden(X34,X36)|~r2_hidden(k4_tarski(X34,X35),k2_zfmisc_1(X36,k1_tarski(X37))))&(X35=X37|~r2_hidden(k4_tarski(X34,X35),k2_zfmisc_1(X36,k1_tarski(X37)))))&(~r2_hidden(X34,X36)|X35!=X37|r2_hidden(k4_tarski(X34,X35),k2_zfmisc_1(X36,k1_tarski(X37))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t129_zfmisc_1])])])).
cnf(c_0_98, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_99, negated_conjecture, (esk1_0=X1|~r2_hidden(k4_tarski(X1,X2),esk3_0)), inference(spm,[status(thm)],[c_0_90, c_0_91])).
cnf(c_0_100, plain, (r2_hidden(k4_tarski(X1,esk4_3(X2,k1_relat_1(X2),X1)),X2)|~r2_hidden(X1,k1_relat_1(X2))), inference(er,[status(thm)],[c_0_92])).
cnf(c_0_101, negated_conjecture, (k4_xboole_0(X1,k1_relat_1(esk3_0))!=X1|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_93, c_0_55])).
cnf(c_0_102, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_94, c_0_95])).
cnf(c_0_103, plain, (r2_hidden(X1,k2_relat_1(X2))|~r2_hidden(k4_tarski(X3,X1),X2)), inference(er,[status(thm)],[c_0_96])).
cnf(c_0_104, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k1_tarski(X2)))), inference(split_conjunct,[status(thm)],[c_0_97])).
fof(c_0_105, plain, ![X83, X84]:(((r1_tarski(X83,X84)|X83!=X84)&(r1_tarski(X84,X83)|X83!=X84))&(~r1_tarski(X83,X84)|~r1_tarski(X84,X83)|X83=X84)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_106, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k4_xboole_0(X2,k4_xboole_0(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_32]), c_0_32])).
fof(c_0_107, plain, ![X72, X73, X74, X75, X76]:((~r1_tarski(X72,X73)|(~r2_hidden(X74,X72)|r2_hidden(X74,X73)))&((r2_hidden(esk11_2(X75,X76),X75)|r1_tarski(X75,X76))&(~r2_hidden(esk11_2(X75,X76),X76)|r1_tarski(X75,X76)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_108, negated_conjecture, (esk1_0=X1|~r2_hidden(X1,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_109, plain, (r2_hidden(esk10_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_110, negated_conjecture, (~r2_hidden(esk1_0,k4_xboole_0(X1,k1_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_101, c_0_74])).
cnf(c_0_111, plain, (k4_xboole_0(k2_enumset1(X1,X1,X1,X1),X2)=k1_xboole_0|r2_hidden(X1,k4_xboole_0(k2_enumset1(X1,X1,X1,X1),X2))), inference(spm,[status(thm)],[c_0_102, c_0_67])).
cnf(c_0_112, negated_conjecture, (r2_hidden(esk2_0,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_103, c_0_79])).
cnf(c_0_113, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,k2_enumset1(X2,X2,X2,X2)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_37]), c_0_38])).
cnf(c_0_114, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_115, plain, (r1_tarski(k4_xboole_0(X1,k4_xboole_0(X1,X2)),X2)), inference(spm,[status(thm)],[c_0_95, c_0_106])).
cnf(c_0_116, plain, (r2_hidden(esk11_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_117, plain, (r2_hidden(esk10_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_118, negated_conjecture, (esk10_2(X1,k1_relat_1(esk3_0))=esk1_0|r1_xboole_0(X1,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_108, c_0_109])).
cnf(c_0_119, negated_conjecture, (k4_xboole_0(k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0),k1_relat_1(esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_110, c_0_111])).
cnf(c_0_120, negated_conjecture, (~r2_hidden(esk2_0,X1)|~r1_xboole_0(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_84, c_0_112])).
cnf(c_0_121, plain, (X1=X2|~r2_hidden(k4_tarski(X3,X1),k2_enumset1(k4_tarski(X4,X2),k4_tarski(X4,X2),k4_tarski(X4,X2),k4_tarski(X4,X2)))), inference(spm,[status(thm)],[c_0_113, c_0_81])).
cnf(c_0_122, plain, (r2_hidden(k4_tarski(esk7_3(X3,X2,X1),X1),X3)|~r2_hidden(X1,X2)|X2!=k2_relat_1(X3)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_123, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=X2|~r1_tarski(X2,k4_xboole_0(X1,k4_xboole_0(X1,X2)))), inference(spm,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_124, plain, (r1_tarski(X1,X2)|~r2_hidden(esk11_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_125, negated_conjecture, (esk11_2(k1_relat_1(esk3_0),X1)=esk1_0|r1_tarski(k1_relat_1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_108, c_0_116])).
cnf(c_0_126, negated_conjecture, (r2_hidden(esk1_0,X1)|r1_xboole_0(X1,k1_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_117, c_0_118])).
cnf(c_0_127, negated_conjecture, (k4_xboole_0(k1_relat_1(esk3_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))=k1_relat_1(esk3_0)|k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_119])).
cnf(c_0_128, negated_conjecture, (k4_xboole_0(X1,k2_relat_1(esk3_0))!=X1|~r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_120, c_0_55])).
cnf(c_0_129, negated_conjecture, (X1=esk2_0|~r2_hidden(k4_tarski(X2,X1),esk3_0)), inference(spm,[status(thm)],[c_0_121, c_0_68])).
cnf(c_0_130, plain, (r2_hidden(k4_tarski(esk7_3(X1,k2_relat_1(X1),X2),X2),X1)|~r2_hidden(X2,k2_relat_1(X1))), inference(er,[status(thm)],[c_0_122])).
cnf(c_0_131, negated_conjecture, (k1_relat_1(esk3_0)=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|~r1_tarski(k1_relat_1(esk3_0),k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_119]), c_0_44]), c_0_44])).
cnf(c_0_132, negated_conjecture, (r1_tarski(k1_relat_1(esk3_0),X1)|~r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_133, negated_conjecture, (k4_xboole_0(X1,k1_relat_1(esk3_0))=X1|r2_hidden(esk1_0,X1)), inference(spm,[status(thm)],[c_0_65, c_0_126])).
cnf(c_0_134, negated_conjecture, (k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_127]), c_0_85])])).
cnf(c_0_135, negated_conjecture, (~r2_hidden(esk2_0,k4_xboole_0(X1,k2_relat_1(esk3_0)))), inference(spm,[status(thm)],[c_0_128, c_0_74])).
cnf(c_0_136, negated_conjecture, (X1=esk2_0|~r2_hidden(X1,k2_relat_1(esk3_0))), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_137, negated_conjecture, (k1_relat_1(esk3_0)!=k1_tarski(esk1_0)|k2_relat_1(esk3_0)!=k1_tarski(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_138, negated_conjecture, (k1_relat_1(esk3_0)=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|~r2_hidden(esk1_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_139, negated_conjecture, (r2_hidden(esk1_0,k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_133]), c_0_134])).
cnf(c_0_140, negated_conjecture, (k4_xboole_0(k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0),k2_relat_1(esk3_0))=k1_xboole_0), inference(spm,[status(thm)],[c_0_135, c_0_111])).
cnf(c_0_141, negated_conjecture, (esk11_2(k2_relat_1(esk3_0),X1)=esk2_0|r1_tarski(k2_relat_1(esk3_0),X1)), inference(spm,[status(thm)],[c_0_136, c_0_116])).
cnf(c_0_142, negated_conjecture, (k1_relat_1(esk3_0)!=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)|k2_relat_1(esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_137, c_0_37]), c_0_37]), c_0_38]), c_0_38])).
cnf(c_0_143, negated_conjecture, (k1_relat_1(esk3_0)=k2_enumset1(esk1_0,esk1_0,esk1_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_138, c_0_139])])).
cnf(c_0_144, negated_conjecture, (k2_relat_1(esk3_0)=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)|~r1_tarski(k2_relat_1(esk3_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_140]), c_0_44]), c_0_44])).
cnf(c_0_145, negated_conjecture, (r1_tarski(k2_relat_1(esk3_0),X1)|~r2_hidden(esk2_0,X1)), inference(spm,[status(thm)],[c_0_124, c_0_141])).
cnf(c_0_146, negated_conjecture, (k2_relat_1(esk3_0)!=k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_142, c_0_143])])).
cnf(c_0_147, negated_conjecture, (k4_xboole_0(k2_relat_1(esk3_0),k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))=k2_relat_1(esk3_0)|k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_75, c_0_140])).
cnf(c_0_148, negated_conjecture, (~r2_hidden(esk2_0,k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_146])).
cnf(c_0_149, plain, (k2_enumset1(X1,X1,X1,X1)=k1_xboole_0|r2_hidden(X1,k2_enumset1(X1,X1,X1,X1))), inference(spm,[status(thm)],[c_0_53, c_0_67])).
cnf(c_0_150, negated_conjecture, (k2_enumset1(esk2_0,esk2_0,esk2_0,esk2_0)!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_147]), c_0_112])])).
cnf(c_0_151, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_150]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 152
# Proof object clause steps            : 103
# Proof object formula steps           : 49
# Proof object conjectures             : 44
# Proof object clause conjectures      : 41
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 59
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 53
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 49
# Processed clauses                    : 2440
# ...of these trivial                  : 29
# ...subsumed                          : 1759
# ...remaining for further processing  : 652
# Other redundant clauses eliminated   : 122
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 14
# Backward-rewritten                   : 120
# Generated clauses                    : 12469
# ...of the previous two non-trivial   : 10006
# Contextual simplify-reflections      : 11
# Paramodulations                      : 12344
# Factorizations                       : 4
# Equation resolutions                 : 122
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 464
#    Positive orientable unit clauses  : 65
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 23
#    Non-unit-clauses                  : 374
# Current number of unprocessed clauses: 7581
# ...number of literals in the above   : 17509
# Current number of archived formulas  : 0
# Current number of archived clauses   : 182
# Clause-clause subsumption calls (NU) : 43784
# Rec. Clause-clause subsumption calls : 31406
# Non-unit clause-clause subsumptions  : 774
# Unit Clause-clause subsumption calls : 2670
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 63
# BW rewrite match successes           : 15
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 169780

# -------------------------------------------------
# User time                : 0.105 s
# System time              : 0.010 s
# Total time               : 0.115 s
# Maximum resident set size: 26244 pages
