# No SInE strategy applied
# Auto-Mode selected heuristic SAT001_MinMin_rr
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t39_waybel_0, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_yellow_0(X1,k6_waybel_0(X1,X2))&k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', t39_waybel_0)).
fof(t24_orders_2, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&l1_orders_2(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>r1_orders_2(X1,X2,X2))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', t24_orders_2)).
fof(t19_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', t19_yellow_0)).
fof(d8_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(r2_yellow_0(X1,X2)<=>?[X3]:(m1_subset_1(X3,u1_struct_0(X1))&(r1_lattice3(X1,X2,X3)&(![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>((r1_lattice3(X1,X2,X4)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X5)=>r1_orders_2(X1,X5,X4))))=>X4=X3))))))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', d8_yellow_0)).
fof(t7_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>((r1_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X2,X3))&((r1_orders_2(X1,X2,X3)=>r1_lattice3(X1,k1_tarski(X3),X2))&((r2_lattice3(X1,k1_tarski(X3),X2)=>r1_orders_2(X1,X3,X2))&(r1_orders_2(X1,X3,X2)=>r2_lattice3(X1,k1_tarski(X3),X2)))))))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', t7_yellow_0)).
fof(t31_yellow_0, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(((X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3))=>(r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2)))))&((r1_lattice3(X1,X3,X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X3,X4)=>r1_orders_2(X1,X4,X2))))=>(X2=k2_yellow_0(X1,X3)&r2_yellow_0(X1,X3)))))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', t31_yellow_0)).
fof(d10_yellow_0, axiom, ![X1]:(l1_orders_2(X1)=>![X2, X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_yellow_0(X1,X2)=>(X3=k2_yellow_0(X1,X2)<=>(r1_lattice3(X1,X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X4)=>r1_orders_2(X1,X4,X3)))))))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', d10_yellow_0)).
fof(t36_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&l1_orders_2(X1))))=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r1_lattice3(X1,X2,X3)<=>r1_lattice3(X1,k4_waybel_0(X1,X2),X3))))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', t36_waybel_0)).
fof(t5_subset, axiom, ![X1, X2, X3]:~((r2_hidden(X1,X2)&(m1_subset_1(X2,k1_zfmisc_1(X3))&v1_xboole_0(X3)))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', t5_subset)).
fof(dt_k6_waybel_0, axiom, ![X1, X2]:(((~(v2_struct_0(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', dt_k6_waybel_0)).
fof(t18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>(r2_hidden(X3,k6_waybel_0(X1,X2))<=>r1_orders_2(X1,X2,X3))))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', t18_waybel_0)).
fof(d18_waybel_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>k6_waybel_0(X1,X2)=k4_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2)))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', d18_waybel_0)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', redefinition_k6_domain_1)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('data/MPTP/proofs/fe1d7d3b95a641d7b2666d73b1465c3c.E_input', dt_k6_domain_1)).
fof(c_0_14, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), introduced(definition)).
fof(c_0_15, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>(((X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3)))=>(r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4))))))&((r1_orders_2(X1,X4,X2)&(r1_orders_2(X1,X4,X3)&![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>((r1_orders_2(X1,X5,X2)&r1_orders_2(X1,X5,X3))=>r1_orders_2(X1,X5,X4)))))=>(X4=k11_lattice3(X1,X2,X3)&r2_yellow_0(X1,k2_tarski(X2,X3))))))), inference(split_equiv,[status(thm)],[c_0_14])).
fof(c_0_16, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_orders_2(X1)&(v4_orders_2(X1)&(v5_orders_2(X1)&l1_orders_2(X1)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(r2_yellow_0(X1,k6_waybel_0(X1,X2))&k2_yellow_0(X1,k6_waybel_0(X1,X2))=X2)))), inference(assume_negation,[status(cth)],[t39_waybel_0])).
fof(c_0_17, plain, ![X61, X62, X63, X64, X65]:(((r1_orders_2(X63,X64,X62)|(X64!=k11_lattice3(X63,X62,X61)|~r2_yellow_0(X63,k2_tarski(X62,X61)))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61))&((r1_orders_2(X63,X64,X61)|(X64!=k11_lattice3(X63,X62,X61)|~r2_yellow_0(X63,k2_tarski(X62,X61)))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61))&(~m1_subset_1(X65,u1_struct_0(X63))|(~r1_orders_2(X63,X65,X62)|~r1_orders_2(X63,X65,X61)|r1_orders_2(X63,X65,X64))|(X64!=k11_lattice3(X63,X62,X61)|~r2_yellow_0(X63,k2_tarski(X62,X61)))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61))))&(((X64=k11_lattice3(X63,X62,X61)|(m1_subset_1(esk9_4(X61,X62,X63,X64),u1_struct_0(X63))|~r1_orders_2(X63,X64,X61)|~r1_orders_2(X63,X64,X62))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61))&(r2_yellow_0(X63,k2_tarski(X62,X61))|(m1_subset_1(esk9_4(X61,X62,X63,X64),u1_struct_0(X63))|~r1_orders_2(X63,X64,X61)|~r1_orders_2(X63,X64,X62))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61)))&((((X64=k11_lattice3(X63,X62,X61)|(r1_orders_2(X63,esk9_4(X61,X62,X63,X64),X62)|~r1_orders_2(X63,X64,X61)|~r1_orders_2(X63,X64,X62))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61))&(r2_yellow_0(X63,k2_tarski(X62,X61))|(r1_orders_2(X63,esk9_4(X61,X62,X63,X64),X62)|~r1_orders_2(X63,X64,X61)|~r1_orders_2(X63,X64,X62))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61)))&((X64=k11_lattice3(X63,X62,X61)|(r1_orders_2(X63,esk9_4(X61,X62,X63,X64),X61)|~r1_orders_2(X63,X64,X61)|~r1_orders_2(X63,X64,X62))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61))&(r2_yellow_0(X63,k2_tarski(X62,X61))|(r1_orders_2(X63,esk9_4(X61,X62,X63,X64),X61)|~r1_orders_2(X63,X64,X61)|~r1_orders_2(X63,X64,X62))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61))))&((X64=k11_lattice3(X63,X62,X61)|(~r1_orders_2(X63,esk9_4(X61,X62,X63,X64),X64)|~r1_orders_2(X63,X64,X61)|~r1_orders_2(X63,X64,X62))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61))&(r2_yellow_0(X63,k2_tarski(X62,X61))|(~r1_orders_2(X63,esk9_4(X61,X62,X63,X64),X64)|~r1_orders_2(X63,X64,X61)|~r1_orders_2(X63,X64,X62))|~m1_subset_1(X64,u1_struct_0(X63))|~epred1_3(X63,X62,X61)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_15])])])])])).
fof(c_0_18, plain, ![X57, X58]:(v2_struct_0(X57)|(~v3_orders_2(X57)|~l1_orders_2(X57))|(~m1_subset_1(X58,u1_struct_0(X57))|r1_orders_2(X57,X58,X58))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_orders_2])])])])).
fof(c_0_19, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_orders_2(esk1_0)&(v4_orders_2(esk1_0)&(v5_orders_2(esk1_0)&l1_orders_2(esk1_0)))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(~r2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))|k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))!=esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_16])])])])).
fof(c_0_20, axiom, ![X1]:((v5_orders_2(X1)&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[t19_yellow_0, c_0_14])).
cnf(c_0_21, plain, (X1=k11_lattice3(X2,X3,X4)|~r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X1)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X4)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_24, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_26, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_27, plain, (X1=k11_lattice3(X2,X3,X4)|r1_orders_2(X2,esk9_4(X4,X3,X2,X1),X3)|~r1_orders_2(X2,X1,X4)|~r1_orders_2(X2,X1,X3)|~m1_subset_1(X1,u1_struct_0(X2))|~epred1_3(X2,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_28, plain, ![X20, X21, X22]:(~v5_orders_2(X20)|~l1_orders_2(X20)|(~m1_subset_1(X21,u1_struct_0(X20))|(~m1_subset_1(X22,u1_struct_0(X20))|epred1_3(X20,X21,X22)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_20])])])).
cnf(c_0_29, plain, (r2_yellow_0(X1,k2_tarski(X2,X3))|r1_orders_2(X1,esk9_4(X3,X2,X1,X4),X3)|~r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~epred1_3(X1,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_30, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_31, plain, (k11_lattice3(X1,X2,X3)=X3|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X3,X3)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_32, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), c_0_26])).
cnf(c_0_33, plain, (k11_lattice3(X1,X2,X3)=X2|~epred1_3(X1,X2,X3)|~r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))), inference(spm,[status(thm)],[c_0_21, c_0_27])).
cnf(c_0_34, plain, (epred1_3(X1,X2,X3)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_35, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_36, plain, (r1_orders_2(X1,X2,X3)|X2!=k11_lattice3(X1,X4,X3)|~r2_yellow_0(X1,k2_tarski(X4,X3))|~m1_subset_1(X2,u1_struct_0(X1))|~epred1_3(X1,X4,X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_37, negated_conjecture, (r1_orders_2(esk1_0,esk9_4(X1,X2,esk1_0,esk2_0),X1)|r2_yellow_0(esk1_0,k2_tarski(X2,X1))|~epred1_3(esk1_0,X2,X1)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,esk2_0,X2)), inference(spm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_38, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X2|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_39, negated_conjecture, (k11_lattice3(esk1_0,X1,X2)=X1|~epred1_3(esk1_0,X1,X2)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_33, c_0_32])).
cnf(c_0_40, negated_conjecture, (epred1_3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35]), c_0_25])])).
cnf(c_0_41, plain, (r1_orders_2(X1,k11_lattice3(X1,X2,X3),X3)|~epred1_3(X1,X2,X3)|~r2_yellow_0(X1,k2_tarski(X2,X3))|~m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))), inference(er,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (r1_orders_2(esk1_0,esk9_4(esk2_0,X1,esk1_0,esk2_0),esk2_0)|r2_yellow_0(esk1_0,k2_tarski(X1,esk2_0))|~epred1_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_32]), c_0_30])])).
cnf(c_0_43, negated_conjecture, (X1=X2|~r1_orders_2(esk1_0,X2,X1)|~r1_orders_2(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39]), c_0_40])).
cnf(c_0_44, plain, (r1_orders_2(esk1_0,esk9_4(esk2_0,X1,esk1_0,esk2_0),esk2_0)|r1_orders_2(esk1_0,k11_lattice3(esk1_0,X1,esk2_0),esk2_0)|~epred1_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(k11_lattice3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_41, c_0_42])).
fof(c_0_45, plain, ![X43, X44, X46, X47, X49, X50, X53]:(((m1_subset_1(esk5_2(X43,X44),u1_struct_0(X43))|~r2_yellow_0(X43,X44)|~l1_orders_2(X43))&((r1_lattice3(X43,X44,esk5_2(X43,X44))|~r2_yellow_0(X43,X44)|~l1_orders_2(X43))&((~m1_subset_1(X46,u1_struct_0(X43))|(~r1_lattice3(X43,X44,X46)|r1_orders_2(X43,X46,esk5_2(X43,X44)))|~r2_yellow_0(X43,X44)|~l1_orders_2(X43))&((m1_subset_1(esk6_3(X43,X44,X47),u1_struct_0(X43))|~r1_lattice3(X43,X44,X47)|X47=esk5_2(X43,X44)|~m1_subset_1(X47,u1_struct_0(X43))|~r2_yellow_0(X43,X44)|~l1_orders_2(X43))&((r1_lattice3(X43,X44,esk6_3(X43,X44,X47))|~r1_lattice3(X43,X44,X47)|X47=esk5_2(X43,X44)|~m1_subset_1(X47,u1_struct_0(X43))|~r2_yellow_0(X43,X44)|~l1_orders_2(X43))&(~r1_orders_2(X43,esk6_3(X43,X44,X47),X47)|~r1_lattice3(X43,X44,X47)|X47=esk5_2(X43,X44)|~m1_subset_1(X47,u1_struct_0(X43))|~r2_yellow_0(X43,X44)|~l1_orders_2(X43)))))))&(((m1_subset_1(esk8_3(X43,X49,X50),u1_struct_0(X43))|m1_subset_1(esk7_3(X43,X49,X50),u1_struct_0(X43))|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))&(((r1_lattice3(X43,X49,esk8_3(X43,X49,X50))|m1_subset_1(esk7_3(X43,X49,X50),u1_struct_0(X43))|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))&(~m1_subset_1(X53,u1_struct_0(X43))|(~r1_lattice3(X43,X49,X53)|r1_orders_2(X43,X53,esk8_3(X43,X49,X50)))|m1_subset_1(esk7_3(X43,X49,X50),u1_struct_0(X43))|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43)))&(esk8_3(X43,X49,X50)!=X50|m1_subset_1(esk7_3(X43,X49,X50),u1_struct_0(X43))|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))))&(((m1_subset_1(esk8_3(X43,X49,X50),u1_struct_0(X43))|r1_lattice3(X43,X49,esk7_3(X43,X49,X50))|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))&(((r1_lattice3(X43,X49,esk8_3(X43,X49,X50))|r1_lattice3(X43,X49,esk7_3(X43,X49,X50))|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))&(~m1_subset_1(X53,u1_struct_0(X43))|(~r1_lattice3(X43,X49,X53)|r1_orders_2(X43,X53,esk8_3(X43,X49,X50)))|r1_lattice3(X43,X49,esk7_3(X43,X49,X50))|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43)))&(esk8_3(X43,X49,X50)!=X50|r1_lattice3(X43,X49,esk7_3(X43,X49,X50))|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))))&((m1_subset_1(esk8_3(X43,X49,X50),u1_struct_0(X43))|~r1_orders_2(X43,esk7_3(X43,X49,X50),X50)|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))&(((r1_lattice3(X43,X49,esk8_3(X43,X49,X50))|~r1_orders_2(X43,esk7_3(X43,X49,X50),X50)|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))&(~m1_subset_1(X53,u1_struct_0(X43))|(~r1_lattice3(X43,X49,X53)|r1_orders_2(X43,X53,esk8_3(X43,X49,X50)))|~r1_orders_2(X43,esk7_3(X43,X49,X50),X50)|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43)))&(esk8_3(X43,X49,X50)!=X50|~r1_orders_2(X43,esk7_3(X43,X49,X50),X50)|~r1_lattice3(X43,X49,X50)|~m1_subset_1(X50,u1_struct_0(X43))|r2_yellow_0(X43,X49)|~l1_orders_2(X43))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_yellow_0])])])])])])).
fof(c_0_46, plain, ![X23, X24, X25]:((~r1_lattice3(X23,k1_tarski(X25),X24)|r1_orders_2(X23,X24,X25)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|~l1_orders_2(X23))&((~r1_orders_2(X23,X24,X25)|r1_lattice3(X23,k1_tarski(X25),X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|~l1_orders_2(X23))&((~r2_lattice3(X23,k1_tarski(X25),X24)|r1_orders_2(X23,X25,X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|~l1_orders_2(X23))&(~r1_orders_2(X23,X25,X24)|r2_lattice3(X23,k1_tarski(X25),X24)|~m1_subset_1(X25,u1_struct_0(X23))|~m1_subset_1(X24,u1_struct_0(X23))|~l1_orders_2(X23))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_yellow_0])])])])).
cnf(c_0_47, negated_conjecture, (k11_lattice3(esk1_0,X1,esk2_0)=esk2_0|r1_orders_2(esk1_0,esk9_4(esk2_0,X1,esk1_0,esk2_0),esk2_0)|~epred1_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,k11_lattice3(esk1_0,X1,esk2_0))|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(k11_lattice3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_44]), c_0_30])])).
cnf(c_0_48, plain, (r1_orders_2(X2,X1,esk5_2(X2,X3))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~r2_yellow_0(X2,X3)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_49, plain, (r1_orders_2(X1,X3,X2)|~r1_lattice3(X1,k1_tarski(X2),X3)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_50, plain, (r1_lattice3(X1,X2,esk5_2(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_51, plain, (m1_subset_1(esk5_2(X1,X2),u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_52, plain, (k11_lattice3(esk1_0,X1,esk2_0)=esk2_0|~epred1_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,k11_lattice3(esk1_0,X1,esk2_0))|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(k11_lattice3(esk1_0,X1,esk2_0),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_47]), c_0_30])])).
cnf(c_0_53, negated_conjecture, (esk5_2(esk1_0,X1)=X2|~r1_orders_2(esk1_0,esk5_2(esk1_0,X1),X2)|~r1_lattice3(esk1_0,X1,X2)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_43, c_0_48]), c_0_25])])).
cnf(c_0_54, plain, (r1_orders_2(X1,esk5_2(X1,k1_tarski(X2)),X2)|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])).
cnf(c_0_55, negated_conjecture, (X1=esk2_0|~epred1_3(esk1_0,X1,esk2_0)|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_52, c_0_39])).
fof(c_0_56, plain, ![X12, X13, X14, X15, X16]:(((r1_lattice3(X12,X14,X13)|(X13!=k2_yellow_0(X12,X14)|~r2_yellow_0(X12,X14))|~m1_subset_1(X13,u1_struct_0(X12))|(~v5_orders_2(X12)|~l1_orders_2(X12)))&(~m1_subset_1(X15,u1_struct_0(X12))|(~r1_lattice3(X12,X14,X15)|r1_orders_2(X12,X15,X13))|(X13!=k2_yellow_0(X12,X14)|~r2_yellow_0(X12,X14))|~m1_subset_1(X13,u1_struct_0(X12))|(~v5_orders_2(X12)|~l1_orders_2(X12))))&(((X13=k2_yellow_0(X12,X16)|(m1_subset_1(esk3_3(X12,X13,X16),u1_struct_0(X12))|~r1_lattice3(X12,X16,X13))|~m1_subset_1(X13,u1_struct_0(X12))|(~v5_orders_2(X12)|~l1_orders_2(X12)))&(r2_yellow_0(X12,X16)|(m1_subset_1(esk3_3(X12,X13,X16),u1_struct_0(X12))|~r1_lattice3(X12,X16,X13))|~m1_subset_1(X13,u1_struct_0(X12))|(~v5_orders_2(X12)|~l1_orders_2(X12))))&(((X13=k2_yellow_0(X12,X16)|(r1_lattice3(X12,X16,esk3_3(X12,X13,X16))|~r1_lattice3(X12,X16,X13))|~m1_subset_1(X13,u1_struct_0(X12))|(~v5_orders_2(X12)|~l1_orders_2(X12)))&(r2_yellow_0(X12,X16)|(r1_lattice3(X12,X16,esk3_3(X12,X13,X16))|~r1_lattice3(X12,X16,X13))|~m1_subset_1(X13,u1_struct_0(X12))|(~v5_orders_2(X12)|~l1_orders_2(X12))))&((X13=k2_yellow_0(X12,X16)|(~r1_orders_2(X12,esk3_3(X12,X13,X16),X13)|~r1_lattice3(X12,X16,X13))|~m1_subset_1(X13,u1_struct_0(X12))|(~v5_orders_2(X12)|~l1_orders_2(X12)))&(r2_yellow_0(X12,X16)|(~r1_orders_2(X12,esk3_3(X12,X13,X16),X13)|~r1_lattice3(X12,X16,X13))|~m1_subset_1(X13,u1_struct_0(X12))|(~v5_orders_2(X12)|~l1_orders_2(X12))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[t31_yellow_0])])])])])])).
cnf(c_0_57, negated_conjecture, (esk5_2(esk1_0,k1_tarski(X1))=X1|~r1_lattice3(esk1_0,k1_tarski(X1),X1)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(esk5_2(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_25])])).
cnf(c_0_58, negated_conjecture, (X1=esk2_0|~r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_40]), c_0_30])])).
cnf(c_0_59, plain, (X1=k2_yellow_0(X2,X3)|~r1_orders_2(X2,esk3_3(X2,X1,X3),X1)|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_60, negated_conjecture, (esk5_2(esk1_0,k1_tarski(X1))=X1|~r1_lattice3(esk1_0,k1_tarski(X1),X1)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_51]), c_0_25])])).
cnf(c_0_61, plain, (r1_lattice3(X1,k1_tarski(X3),X2)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_62, plain, (r2_yellow_0(X1,X2)|r1_lattice3(X1,X2,esk3_3(X1,X3,X2))|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (r2_yellow_0(X1,X2)|m1_subset_1(esk3_3(X1,X3,X2),u1_struct_0(X1))|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_64, negated_conjecture, (X1=esk2_0|~r1_orders_2(esk1_0,esk2_0,X1)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_32]), c_0_30])])).
fof(c_0_65, plain, ![X26, X27, X28, X29]:(((r1_lattice3(X26,X27,X28)|X28!=k2_yellow_0(X26,X27)|~r2_yellow_0(X26,X27)|~m1_subset_1(X28,u1_struct_0(X26))|~l1_orders_2(X26))&(~m1_subset_1(X29,u1_struct_0(X26))|(~r1_lattice3(X26,X27,X29)|r1_orders_2(X26,X29,X28))|X28!=k2_yellow_0(X26,X27)|~r2_yellow_0(X26,X27)|~m1_subset_1(X28,u1_struct_0(X26))|~l1_orders_2(X26)))&((m1_subset_1(esk4_3(X26,X27,X28),u1_struct_0(X26))|~r1_lattice3(X26,X27,X28)|X28=k2_yellow_0(X26,X27)|~r2_yellow_0(X26,X27)|~m1_subset_1(X28,u1_struct_0(X26))|~l1_orders_2(X26))&((r1_lattice3(X26,X27,esk4_3(X26,X27,X28))|~r1_lattice3(X26,X27,X28)|X28=k2_yellow_0(X26,X27)|~r2_yellow_0(X26,X27)|~m1_subset_1(X28,u1_struct_0(X26))|~l1_orders_2(X26))&(~r1_orders_2(X26,esk4_3(X26,X27,X28),X28)|~r1_lattice3(X26,X27,X28)|X28=k2_yellow_0(X26,X27)|~r2_yellow_0(X26,X27)|~m1_subset_1(X28,u1_struct_0(X26))|~l1_orders_2(X26))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_yellow_0])])])])])).
cnf(c_0_66, plain, (esk5_2(X1,X2)=k2_yellow_0(X1,X3)|~r1_lattice3(X1,X2,esk3_3(X1,esk5_2(X1,X2),X3))|~r1_lattice3(X1,X3,esk5_2(X1,X2))|~r2_yellow_0(X1,X2)|~m1_subset_1(esk3_3(X1,esk5_2(X1,X2),X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_48]), c_0_51])).
cnf(c_0_67, negated_conjecture, (esk5_2(esk1_0,k1_tarski(X1))=X1|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_61]), c_0_25])]), c_0_32])).
cnf(c_0_68, plain, (r2_yellow_0(X1,X2)|~r1_orders_2(X1,esk3_3(X1,X3,X2),X3)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_69, plain, (r1_orders_2(X1,esk3_3(X1,X2,k1_tarski(X3)),X3)|r2_yellow_0(X1,k1_tarski(X3))|~r1_lattice3(X1,k1_tarski(X3),X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_62]), c_0_63])).
cnf(c_0_70, negated_conjecture, (esk5_2(esk1_0,X1)=esk2_0|~r1_orders_2(esk1_0,esk5_2(esk1_0,X1),esk2_0)|~r1_lattice3(esk1_0,X1,esk2_0)|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(esk5_2(esk1_0,X1),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_48]), c_0_30]), c_0_25])])).
cnf(c_0_71, plain, (r1_orders_2(X2,X1,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|X4!=k2_yellow_0(X2,X3)|~r2_yellow_0(X2,X3)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_72, negated_conjecture, (X1=k2_yellow_0(esk1_0,X2)|~r1_lattice3(esk1_0,k1_tarski(X1),esk3_3(esk1_0,X1,X2))|~r1_lattice3(esk1_0,X2,X1)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(esk3_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_25]), c_0_35])])).
cnf(c_0_73, plain, (X1=k2_yellow_0(X2,X3)|r1_lattice3(X2,X3,esk3_3(X2,X1,X3))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_74, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(X1),X1)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_67]), c_0_25])])).
cnf(c_0_75, plain, (r2_yellow_0(X1,k1_tarski(X2))|~r1_lattice3(X1,k1_tarski(X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_76, negated_conjecture, (esk5_2(esk1_0,k1_tarski(esk2_0))=esk2_0|~r1_lattice3(esk1_0,k1_tarski(esk2_0),esk2_0)|~r2_yellow_0(esk1_0,k1_tarski(esk2_0))|~m1_subset_1(esk5_2(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_54]), c_0_30]), c_0_25])])).
cnf(c_0_77, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,X3))|~r1_lattice3(X1,X3,X2)|~r2_yellow_0(X1,X3)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(er,[status(thm)],[c_0_71])).
cnf(c_0_78, plain, (X1=k2_yellow_0(X2,X3)|m1_subset_1(esk3_3(X2,X1,X3),u1_struct_0(X2))|~r1_lattice3(X2,X3,X1)|~m1_subset_1(X1,u1_struct_0(X2))|~v5_orders_2(X2)|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_79, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(X1))=X1|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(esk3_3(esk1_0,X1,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73]), c_0_25]), c_0_35])]), c_0_74])).
cnf(c_0_80, plain, (r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(spm,[status(thm)],[c_0_75, c_0_61])).
fof(c_0_81, plain, ![X54, X55, X56]:((~r1_lattice3(X54,X55,X56)|r1_lattice3(X54,k4_waybel_0(X54,X55),X56)|~m1_subset_1(X56,u1_struct_0(X54))|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(X54)))|(v2_struct_0(X54)|(~v3_orders_2(X54)|(~v4_orders_2(X54)|~l1_orders_2(X54)))))&(~r1_lattice3(X54,k4_waybel_0(X54,X55),X56)|r1_lattice3(X54,X55,X56)|~m1_subset_1(X56,u1_struct_0(X54))|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(X54)))|(v2_struct_0(X54)|(~v3_orders_2(X54)|(~v4_orders_2(X54)|~l1_orders_2(X54)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_waybel_0])])])])])).
fof(c_0_82, plain, ![X38, X39, X40]:(~r2_hidden(X38,X39)|(~m1_subset_1(X39,k1_zfmisc_1(X40))|~v1_xboole_0(X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_83, plain, ![X8, X9]:(v2_struct_0(X8)|~l1_orders_2(X8)|~m1_subset_1(X9,u1_struct_0(X8))|m1_subset_1(k6_waybel_0(X8,X9),k1_zfmisc_1(u1_struct_0(X8)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_waybel_0])])])).
cnf(c_0_84, negated_conjecture, (esk5_2(esk1_0,k1_tarski(esk2_0))=esk2_0|~r1_lattice3(esk1_0,k1_tarski(esk2_0),esk2_0)|~r2_yellow_0(esk1_0,k1_tarski(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_51]), c_0_25])])).
cnf(c_0_85, plain, (k2_yellow_0(X1,X2)=k2_yellow_0(X1,X3)|~r1_lattice3(X1,X2,esk3_3(X1,k2_yellow_0(X1,X2),X3))|~r1_lattice3(X1,X3,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_77]), c_0_78])).
cnf(c_0_86, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(X1))=X1|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_78]), c_0_25]), c_0_35])]), c_0_74])).
cnf(c_0_87, negated_conjecture, (r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_32]), c_0_25]), c_0_35])])).
cnf(c_0_88, plain, (r1_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r1_lattice3(X1,k4_waybel_0(X1,X2),X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_89, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_90, plain, (v2_struct_0(X1)|m1_subset_1(k6_waybel_0(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_83])).
fof(c_0_91, plain, ![X33, X34, X35]:((~r2_hidden(X35,k6_waybel_0(X33,X34))|r1_orders_2(X33,X34,X35)|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(v2_struct_0(X33)|~l1_orders_2(X33)))&(~r1_orders_2(X33,X34,X35)|r2_hidden(X35,k6_waybel_0(X33,X34))|~m1_subset_1(X35,u1_struct_0(X33))|~m1_subset_1(X34,u1_struct_0(X33))|(v2_struct_0(X33)|~l1_orders_2(X33)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_waybel_0])])])])])).
cnf(c_0_92, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_lattice3(esk1_0,k1_tarski(esk2_0),esk2_0)|~r2_yellow_0(esk1_0,k1_tarski(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_84]), c_0_30]), c_0_25])])).
cnf(c_0_93, negated_conjecture, (X1=k2_yellow_0(esk1_0,X2)|~r1_lattice3(esk1_0,k1_tarski(X1),esk3_3(esk1_0,X1,X2))|~r1_lattice3(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86]), c_0_25]), c_0_35])]), c_0_87])).
cnf(c_0_94, plain, (X1=k2_yellow_0(X2,k4_waybel_0(X2,X3))|r1_lattice3(X2,X3,esk3_3(X2,X1,k4_waybel_0(X2,X3)))|v2_struct_0(X2)|~r1_lattice3(X2,k4_waybel_0(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_73]), c_0_78])).
cnf(c_0_95, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_96, plain, ![X36, X37]:(v2_struct_0(X36)|~l1_orders_2(X36)|(~m1_subset_1(X37,u1_struct_0(X36))|k6_waybel_0(X36,X37)=k4_waybel_0(X36,k6_domain_1(u1_struct_0(X36),X37)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d18_waybel_0])])])])).
fof(c_0_97, plain, ![X59, X60]:(v1_xboole_0(X59)|~m1_subset_1(X60,X59)|k6_domain_1(X59,X60)=k1_tarski(X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_98, plain, (v2_struct_0(X1)|~r2_hidden(X2,k6_waybel_0(X1,X3))|~v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_99, plain, (r2_hidden(X3,k6_waybel_0(X1,X2))|v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_100, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)|~r1_lattice3(esk1_0,k1_tarski(esk2_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_87]), c_0_30])])).
cnf(c_0_101, negated_conjecture, (r1_lattice3(esk1_0,k1_tarski(X1),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(spm,[status(thm)],[c_0_74, c_0_87])).
cnf(c_0_102, plain, (r2_yellow_0(X1,X2)|~r1_lattice3(X1,X3,esk3_3(X1,k2_yellow_0(X1,X3),X2))|~r1_lattice3(X1,X2,k2_yellow_0(X1,X3))|~r2_yellow_0(X1,X3)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_77]), c_0_63])).
cnf(c_0_103, negated_conjecture, (k2_yellow_0(esk1_0,k4_waybel_0(esk1_0,k1_tarski(X1)))=X1|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,k1_tarski(X1)),X1)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_25]), c_0_35]), c_0_95]), c_0_24])]), c_0_26])).
cnf(c_0_104, plain, (r1_lattice3(X1,k4_waybel_0(X1,X2),X3)|v2_struct_0(X1)|~r1_lattice3(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_105, plain, (v2_struct_0(X1)|k6_waybel_0(X1,X2)=k4_waybel_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_106, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_107, plain, (v2_struct_0(X1)|~r1_orders_2(X1,X2,X3)|~v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_98, c_0_99])).
cnf(c_0_108, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_30])])).
fof(c_0_109, plain, ![X41, X42]:(v1_xboole_0(X41)|~m1_subset_1(X42,X41)|m1_subset_1(k6_domain_1(X41,X42),k1_zfmisc_1(X41))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_110, negated_conjecture, (r2_yellow_0(esk1_0,X1)|~r1_lattice3(esk1_0,k1_tarski(X2),esk3_3(esk1_0,X2,X1))|~r1_lattice3(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_86]), c_0_25]), c_0_35])]), c_0_87])).
cnf(c_0_111, plain, (r1_lattice3(X1,X2,esk3_3(X1,X3,k4_waybel_0(X1,X2)))|r2_yellow_0(X1,k4_waybel_0(X1,X2))|v2_struct_0(X1)|~r1_lattice3(X1,k4_waybel_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_62]), c_0_63])).
cnf(c_0_112, negated_conjecture, (k2_yellow_0(esk1_0,k4_waybel_0(esk1_0,k1_tarski(X1)))=X1|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_104]), c_0_25]), c_0_95]), c_0_24])]), c_0_26]), c_0_101])).
cnf(c_0_113, plain, (k4_waybel_0(X1,k1_tarski(X2))=k6_waybel_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(spm,[status(thm)],[c_0_105, c_0_106])).
cnf(c_0_114, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_30]), c_0_25])]), c_0_26])).
cnf(c_0_115, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_116, negated_conjecture, (r2_yellow_0(esk1_0,k4_waybel_0(esk1_0,k1_tarski(X1)))|~r1_lattice3(esk1_0,k4_waybel_0(esk1_0,k1_tarski(X1)),X1)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_25]), c_0_35]), c_0_95]), c_0_24])]), c_0_26])).
cnf(c_0_117, negated_conjecture, (k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,X1))=X1|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_25])]), c_0_114]), c_0_26])).
cnf(c_0_118, plain, (v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(spm,[status(thm)],[c_0_115, c_0_106])).
cnf(c_0_119, negated_conjecture, (r2_yellow_0(esk1_0,k4_waybel_0(esk1_0,k1_tarski(X1)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_104]), c_0_25]), c_0_95]), c_0_24])]), c_0_26]), c_0_101])).
cnf(c_0_120, negated_conjecture, (~r2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))|k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))!=esk2_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_121, negated_conjecture, (k2_yellow_0(esk1_0,k6_waybel_0(esk1_0,X1))=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_118]), c_0_114])).
cnf(c_0_122, negated_conjecture, (r2_yellow_0(esk1_0,k6_waybel_0(esk1_0,X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_113]), c_0_25])]), c_0_114]), c_0_26])).
cnf(c_0_123, negated_conjecture, (~r2_yellow_0(esk1_0,k6_waybel_0(esk1_0,esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121]), c_0_30])])).
cnf(c_0_124, negated_conjecture, (r2_yellow_0(esk1_0,k6_waybel_0(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_118]), c_0_114])).
cnf(c_0_125, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124]), c_0_30])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 126
# Proof object clause steps            : 93
# Proof object formula steps           : 33
# Proof object conjectures             : 50
# Proof object clause conjectures      : 47
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 34
# Proof object initial formulas used   : 14
# Proof object generating inferences   : 57
# Proof object simplifying inferences  : 115
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 71
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 71
# Processed clauses                    : 705
# ...of these trivial                  : 3
# ...subsumed                          : 278
# ...remaining for further processing  : 424
# Other redundant clauses eliminated   : 7
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 65
# Backward-rewritten                   : 31
# Generated clauses                    : 1297
# ...of the previous two non-trivial   : 1105
# Contextual simplify-reflections      : 57
# Paramodulations                      : 1290
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 253
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 238
# Current number of unprocessed clauses: 483
# ...number of literals in the above   : 3909
# Current number of archived formulas  : 0
# Current number of archived clauses   : 164
# Clause-clause subsumption calls (NU) : 15038
# Rec. Clause-clause subsumption calls : 1894
# Non-unit clause-clause subsumptions  : 365
# Unit Clause-clause subsumption calls : 323
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 53044

# -------------------------------------------------
# User time                : 0.044 s
# System time              : 0.003 s
# Total time               : 0.047 s
# Maximum resident set size: 28792 pages
