# No SInE strategy applied
# Auto-Mode selected heuristic G_E___300_C01_F1_SE_CS_SP_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.008 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_xboole_0, axiom, ![X1, X2, X3]:(X3=k4_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)&~(r2_hidden(X4,X2))))), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', d5_xboole_0)).
fof(t3_xboole_0, axiom, ![X1, X2]:(~((~(r1_xboole_0(X1,X2))&![X3]:~((r2_hidden(X3,X1)&r2_hidden(X3,X2)))))&~((?[X3]:(r2_hidden(X3,X1)&r2_hidden(X3,X2))&r1_xboole_0(X1,X2)))), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t3_xboole_0)).
fof(t21_xboole_1, axiom, ![X1, X2]:k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t21_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t1_boole)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t36_xboole_1)).
fof(t3_boole, axiom, ![X1]:k4_xboole_0(X1,k1_xboole_0)=X1, file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t3_boole)).
fof(t16_xboole_1, axiom, ![X1, X2, X3]:k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3)), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t16_xboole_1)).
fof(t39_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t39_xboole_1)).
fof(t41_xboole_1, axiom, ![X1, X2, X3]:k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3)), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t41_xboole_1)).
fof(l32_xboole_1, axiom, ![X1, X2]:(k4_xboole_0(X1,X2)=k1_xboole_0<=>r1_tarski(X1,X2)), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', l32_xboole_1)).
fof(d7_xboole_0, axiom, ![X1, X2]:(r1_xboole_0(X1,X2)<=>k3_xboole_0(X1,X2)=k1_xboole_0), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', d7_xboole_0)).
fof(commutativity_k3_xboole_0, axiom, ![X1, X2]:k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', commutativity_k3_xboole_0)).
fof(t28_xboole_1, axiom, ![X1, X2]:(r1_tarski(X1,X2)=>k3_xboole_0(X1,X2)=X1), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t28_xboole_1)).
fof(t47_xboole_1, axiom, ![X1, X2]:k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t47_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', commutativity_k2_xboole_0)).
fof(t48_xboole_1, conjecture, ![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2), file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t48_xboole_1)).
fof(t22_xboole_1, axiom, ![X1, X2]:k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1, file('data/MPTP/proofs/ff7ceef429734d47808187ba8dec41d7.E_input', t22_xboole_1)).
fof(c_0_17, plain, ![X19, X20, X21, X22, X23, X24, X25, X26]:((((r2_hidden(X22,X19)|~r2_hidden(X22,X21)|X21!=k4_xboole_0(X19,X20))&(~r2_hidden(X22,X20)|~r2_hidden(X22,X21)|X21!=k4_xboole_0(X19,X20)))&(~r2_hidden(X23,X19)|r2_hidden(X23,X20)|r2_hidden(X23,X21)|X21!=k4_xboole_0(X19,X20)))&((~r2_hidden(esk3_3(X24,X25,X26),X26)|(~r2_hidden(esk3_3(X24,X25,X26),X24)|r2_hidden(esk3_3(X24,X25,X26),X25))|X26=k4_xboole_0(X24,X25))&((r2_hidden(esk3_3(X24,X25,X26),X24)|r2_hidden(esk3_3(X24,X25,X26),X26)|X26=k4_xboole_0(X24,X25))&(~r2_hidden(esk3_3(X24,X25,X26),X25)|r2_hidden(esk3_3(X24,X25,X26),X26)|X26=k4_xboole_0(X24,X25))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d5_xboole_0])])])])])])])).
cnf(c_0_18, plain, (~r2_hidden(X1,X2)|~r2_hidden(X1,X3)|X3!=k4_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_19, plain, ![X28, X29, X31, X32, X33]:(((r2_hidden(esk4_2(X28,X29),X28)|r1_xboole_0(X28,X29))&(r2_hidden(esk4_2(X28,X29),X29)|r1_xboole_0(X28,X29)))&(~r2_hidden(X33,X31)|~r2_hidden(X33,X32)|~r1_xboole_0(X31,X32))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t3_xboole_0])])])])])])])).
fof(c_0_20, plain, ![X14, X15]:k3_xboole_0(X14,k2_xboole_0(X14,X15))=X14, inference(variable_rename,[status(thm)],[t21_xboole_1])).
fof(c_0_21, plain, ![X7]:k2_xboole_0(X7,k1_xboole_0)=X7, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_22, plain, ![X50, X51]:r1_tarski(k4_xboole_0(X50,X51),X50), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_23, plain, ![X47]:k4_xboole_0(X47,k1_xboole_0)=X47, inference(variable_rename,[status(thm)],[t3_boole])).
cnf(c_0_24, plain, (~r2_hidden(X1,k4_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(er,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (r2_hidden(esk4_2(X1,X2),X2)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
fof(c_0_26, plain, ![X16, X17, X18]:k3_xboole_0(k3_xboole_0(X16,X17),X18)=k3_xboole_0(X16,k3_xboole_0(X17,X18)), inference(variable_rename,[status(thm)],[t16_xboole_1])).
cnf(c_0_27, plain, (k3_xboole_0(X1,k2_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_28, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_21])).
fof(c_0_29, plain, ![X10, X11]:k2_xboole_0(X10,k4_xboole_0(X11,X10))=k2_xboole_0(X10,X11), inference(variable_rename,[status(thm)],[t39_xboole_1])).
fof(c_0_30, plain, ![X40, X41, X42]:k4_xboole_0(k4_xboole_0(X40,X41),X42)=k4_xboole_0(X40,k2_xboole_0(X41,X42)), inference(variable_rename,[status(thm)],[t41_xboole_1])).
fof(c_0_31, plain, ![X43, X44]:((k4_xboole_0(X43,X44)!=k1_xboole_0|r1_tarski(X43,X44))&(~r1_tarski(X43,X44)|k4_xboole_0(X43,X44)=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l32_xboole_1])])).
cnf(c_0_32, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (k4_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_23])).
fof(c_0_34, plain, ![X34, X35]:((~r1_xboole_0(X34,X35)|k3_xboole_0(X34,X35)=k1_xboole_0)&(k3_xboole_0(X34,X35)!=k1_xboole_0|r1_xboole_0(X34,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_xboole_0])])).
cnf(c_0_35, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X3))|~r2_hidden(esk4_2(X1,k4_xboole_0(X2,X3)),X3)), inference(spm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_36, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_37, plain, (k3_xboole_0(k3_xboole_0(X1,X2),X3)=k3_xboole_0(X1,k3_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_38, plain, (k3_xboole_0(X1,X1)=X1), inference(spm,[status(thm)],[c_0_27, c_0_28])).
fof(c_0_39, plain, ![X48, X49]:k3_xboole_0(X48,X49)=k3_xboole_0(X49,X48), inference(variable_rename,[status(thm)],[commutativity_k3_xboole_0])).
fof(c_0_40, plain, ![X8, X9]:(~r1_tarski(X8,X9)|k3_xboole_0(X8,X9)=X8), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t28_xboole_1])])).
fof(c_0_41, plain, ![X12, X13]:k4_xboole_0(X12,k3_xboole_0(X12,X13))=k4_xboole_0(X12,X13), inference(variable_rename,[status(thm)],[t47_xboole_1])).
cnf(c_0_42, plain, (k2_xboole_0(X1,k4_xboole_0(X2,X1))=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_43, plain, (k4_xboole_0(k4_xboole_0(X1,X2),X3)=k4_xboole_0(X1,k2_xboole_0(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_44, plain, (k4_xboole_0(X1,X2)=k1_xboole_0|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_45, plain, (r1_tarski(X1,X1)), inference(spm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_46, plain, (k3_xboole_0(X1,X2)=k1_xboole_0|~r1_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_47, plain, (r1_xboole_0(X1,k4_xboole_0(X2,X1))), inference(spm,[status(thm)],[c_0_35, c_0_36])).
fof(c_0_48, plain, ![X38, X39]:k2_xboole_0(X38,X39)=k2_xboole_0(X39,X38), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_49, plain, (k3_xboole_0(X1,k3_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_50, plain, (k3_xboole_0(X1,X2)=k3_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (k3_xboole_0(X1,X2)=X1|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_40])).
cnf(c_0_52, plain, (r1_tarski(X1,X2)|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_53, plain, (k4_xboole_0(X1,k3_xboole_0(X1,X2))=k4_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_54, plain, (k2_xboole_0(X1,k4_xboole_0(X2,k2_xboole_0(X3,X1)))=k2_xboole_0(X1,k4_xboole_0(X2,X3))), inference(spm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_55, plain, (k4_xboole_0(X1,X1)=k1_xboole_0), inference(spm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_56, plain, (k3_xboole_0(X1,k4_xboole_0(X2,X1))=k1_xboole_0), inference(spm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_57, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (k3_xboole_0(X1,k3_xboole_0(X2,X1))=k3_xboole_0(X2,X1)), inference(spm,[status(thm)],[c_0_49, c_0_50])).
cnf(c_0_59, plain, (k3_xboole_0(X1,X2)=X1|k4_xboole_0(X1,X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_60, plain, (k4_xboole_0(X1,k3_xboole_0(X2,X1))=k4_xboole_0(X1,X2)), inference(spm,[status(thm)],[c_0_53, c_0_50])).
cnf(c_0_61, plain, (k2_xboole_0(X1,k4_xboole_0(k2_xboole_0(X2,X1),X2))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55]), c_0_28])).
cnf(c_0_62, plain, (k4_xboole_0(X1,k4_xboole_0(X2,X1))=X1), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_56]), c_0_33])).
cnf(c_0_63, plain, (k3_xboole_0(X1,k2_xboole_0(X2,X1))=X1), inference(spm,[status(thm)],[c_0_27, c_0_57])).
cnf(c_0_64, plain, (k3_xboole_0(X1,X2)=X2|k4_xboole_0(X2,X1)!=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59]), c_0_60])).
cnf(c_0_65, plain, (k2_xboole_0(X1,k4_xboole_0(k2_xboole_0(X1,X2),X2))=X1), inference(spm,[status(thm)],[c_0_61, c_0_57])).
cnf(c_0_66, plain, (k4_xboole_0(X1,k2_xboole_0(k4_xboole_0(X2,X1),X3))=k4_xboole_0(X1,X3)), inference(spm,[status(thm)],[c_0_43, c_0_62])).
fof(c_0_67, negated_conjecture, ~(![X1, X2]:k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(assume_negation,[status(cth)],[t48_xboole_1])).
cnf(c_0_68, plain, (k2_xboole_0(X1,X2)=X2|k4_xboole_0(k2_xboole_0(X1,X2),X2)!=k1_xboole_0), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_69, plain, (k2_xboole_0(X1,k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1)))=X1), inference(spm,[status(thm)],[c_0_65, c_0_42])).
cnf(c_0_70, plain, (k4_xboole_0(X1,k4_xboole_0(X2,k4_xboole_0(X3,X1)))=k4_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_42]), c_0_66])).
cnf(c_0_71, plain, (k4_xboole_0(X1,k2_xboole_0(X1,X2))=k1_xboole_0), inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_27]), c_0_55])).
fof(c_0_72, plain, ![X45, X46]:k2_xboole_0(X45,k3_xboole_0(X45,X46))=X45, inference(variable_rename,[status(thm)],[t22_xboole_1])).
fof(c_0_73, negated_conjecture, k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))!=k3_xboole_0(esk1_0,esk2_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_67])])])).
cnf(c_0_74, plain, (k4_xboole_0(k2_xboole_0(X1,X2),k4_xboole_0(X2,X1))=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_71])])).
cnf(c_0_75, plain, (k2_xboole_0(X1,k3_xboole_0(X1,X2))=X1), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_76, negated_conjecture, (k4_xboole_0(esk1_0,k4_xboole_0(esk1_0,esk2_0))!=k3_xboole_0(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_77, plain, (k4_xboole_0(X1,k4_xboole_0(X1,X2))=k3_xboole_0(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_53]), c_0_57]), c_0_75])).
cnf(c_0_78, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_77])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 79
# Proof object clause steps            : 44
# Proof object formula steps           : 35
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 17
# Proof object generating inferences   : 23
# Proof object simplifying inferences  : 13
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 18
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 9067
# ...of these trivial                  : 3020
# ...subsumed                          : 4449
# ...remaining for further processing  : 1598
# Other redundant clauses eliminated   : 3542
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 246
# Generated clauses                    : 527378
# ...of the previous two non-trivial   : 224826
# Contextual simplify-reflections      : 3
# Paramodulations                      : 523826
# Factorizations                       : 10
# Equation resolutions                 : 3542
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1337
#    Positive orientable unit clauses  : 1007
#    Positive unorientable unit clauses: 5
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 323
# Current number of unprocessed clauses: 215323
# ...number of literals in the above   : 392807
# Current number of archived formulas  : 0
# Current number of archived clauses   : 258
# Clause-clause subsumption calls (NU) : 26752
# Rec. Clause-clause subsumption calls : 22095
# Non-unit clause-clause subsumptions  : 2843
# Unit Clause-clause subsumption calls : 506
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 8505
# BW rewrite match successes           : 282
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4126947

# -------------------------------------------------
# User time                : 2.611 s
# System time              : 0.053 s
# Total time               : 2.664 s
# Maximum resident set size: 169796 pages
